Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvpSem
% (139091)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 153
% Passive clauses: 2349
% Generated clauses: 22211
% Final active clauses: 150
% Final passive clauses: 2079
% Input formulas: 133
% Initial clauses: 188
% 
% Duplicate literals: 605
% Trivial inequalities: 367
% Fw subsumption resolutions: 3040
% Fw demodulations: 3175
% Bw demodulations: 2
% 
% Simple tautologies: 360
% Equational tautologies: 51
% Forward subsumptions: 11710
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 227
% Factoring: 2
% Forward superposition: 7365
% Backward superposition: 6834
% Self superposition: 198
% Equality factoring: 105
% Equality resolution: 33
% 
% Split clauses: 250
% Split components: 529
% SAT solver clauses: 583
% SAT solver unit clauses: 12
% SAT solver binary clauses: 94
% 
% Sat splits: 250
% 
% TWLsolver clauses: 583
% TWLsolver calls for satisfiability: 109
% 
% Memory used [KB]: 14328
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvpSem
% (139105)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 279
% Passive clauses: 1004
% Generated clauses: 29140
% Final active clauses: 226
% Final passive clauses: 697
% Input formulas: 133
% Initial clauses: 188
% 
% Duplicate literals: 767
% Trivial inequalities: 1138
% Fw subsumption resolutions: 2840
% Fw demodulations: 4658
% Bw demodulations: 44
% Condensations: 1166
% Global subsumptions: 4
% 
% Simple tautologies: 483
% Equational tautologies: 50
% Forward subsumptions: 15193
% Fw demodulations to eq. taut.: 11
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 103
% Factoring: 2
% Forward superposition: 8152
% Backward superposition: 8402
% Self superposition: 143
% Equality factoring: 83
% Equality resolution: 62
% 
% Split clauses: 62
% Split components: 132
% SAT solver clauses: 1967
% SAT solver unit clauses: 294
% SAT solver binary clauses: 86
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 10
% 
% Sat splits: 62
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 1948
% TWLsolver calls for satisfiability: 1298
% 
% Memory used [KB]: 20340
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvpSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvpSem
% SZS output start Proof for thm.bvpSem
fof(f73408,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f423,f623,f621,f619,f617,f615,f613,f611,f609,f607,f605,f422,f622,f421,f620,f420,f618,f419,f616,f418,f614,f417,f612,f416,f610,f415,f608,f414,f606,f604,f641,f1694,f1695,f1686,f1699,f2004,f2007,f2005,f3070,f3076,f3078,f3082,f3086,f3088,f3092,f3094,f3096,f3098,f3103,f3105,f3201,f3202,f3203,f3204,f3205,f3206,f5184,f5244,f10365,f10463,f10507,f10509,f10510,f10511,f8797,f10508,f10512,f10514,f10515,f10516,f8317,f10513,f10517,f10519,f10520,f10521,f7937,f10518,f10522,f10524,f10525,f10526,f7569,f10523,f5370,f5330,f5329,f10543,f10542,f5369,f10539,f5299,f10547,f10548,f10549,f5361,f10551,f10553,f10541,f3561,f3578,f2344,f2355,f1697,f10726,f10728,f10732,f10733,f10734,f10735,f10736,f12939,f13485,f12985,f21586,f21569,f21564,f21563,f21559,f21550,f21671,f21549,f21713,f21548,f21714,f21547,f21715,f21462,f22005,f21461,f22006,f21460,f22007,f21453,f22010,f22220,f27783,f27950,f28480,f28583,f28469,f28611,f29232,f29360,f29220,f29389,f30062,f30193,f30049,f30223,f38276,f38279,f38265,f38277,f38278,f38353,f38356,f38358,f38367,f38370,f38390,f38393,f38406,f38427,f38433,f38435,f38437,f39044,f39238,f39240,f39242,f39244,f39246,f39248,f39250,f39252,f39253,f39254,f39255,f39256,f39257,f39258,f39259,f39260,f39261,f39263,f39266,f39269,f39271,f39273,f39275,f39282,f39285,f39287,f39289,f39290,f39291,f39292,f39293,f39294,f39295,f39296,f39297,f39300,f39303,f39306,f39309,f39312,f39314,f43026,f43028,f43027,f43035,f43037,f43036,f43040,f10540,f43042,f3077,f43041,f43283,f43284,f43287,f43288,f43290,f43291,f43293,f43294,f43297,f2006,f43298,f52758,f52769,f57859,f57871,f57896,f58377,f58392,f59971,f59988,f60025,f60566,f60583,f60618,f60619,f60996,f61016,f61057,f63363,f63380,f66816,f69076,f5243,f71352,f71370,f71404,f73348,f73369])).
fof(f73369,plain,(
  $false | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f73368,f611])).
fof(f73368,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK12) | $spl0),
  inference(forward_demodulation,[],[f73367,f559])).
fof(f559,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.option.option'('type.num.num'),X53) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_global')).
fof(f73367,plain,(
  s('type.option.option'('type.num.num'),sK12) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | $spl0),
  inference(forward_demodulation,[],[f73181,f559])).
fof(f73181,plain,(
  s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | $spl0),
  inference(unit_resulting_resolution,[],[f604,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f407,plain,(
  ! [X0,X1,X2] : ((s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))) & s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))) & s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2))))) & (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) != s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))) | s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) != s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))) | s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) != s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))) | s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X2))) | s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) != s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))) | s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2))) | s('type.bvpSem.state'(X0),X1) = s('type.bvpSem.state'(X0),X2)))),
  inference(flattening,[],[f406])).
fof(f406,plain,(
  ! [X0,X1,X2] : ((s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))) & s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))) & s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2))))) & ((s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) != s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))) | s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) != s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))) | s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) != s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))) | s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X2))) | s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) != s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))) | s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2)))) | s('type.bvpSem.state'(X0),X1) = s('type.bvpSem.state'(X0),X2)))),
  inference(nnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : (s('type.bvpSem.state'(X0),X1) = s('type.bvpSem.state'(X0),X2) <=> (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))) & s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))) & s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2)))))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0] : ! [X1,X2] : (s('type.bvpSem.state'(X0),X1) = s('type.bvpSem.state'(X0),X2) <=> (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))) & s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))) & s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))) & s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X50,X49] : (s('type.bvpSem.state'(X6),X49) = s('type.bvpSem.state'(X6),X50) <=> (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X49))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X50))) & s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X49))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X50))) & s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X49))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X50))) & s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X49))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X50))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X49))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X50))) & s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X49))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X50))) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X49))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X50))) & s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X49))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X50))) & s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X49))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X50)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_component_equality')).
fof(f73348,plain,(
  $false | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f73347,f611])).
fof(f73347,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK12) | $spl0),
  inference(forward_demodulation,[],[f73346,f559])).
fof(f73346,plain,(
  s('type.option.option'('type.num.num'),sK12) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | $spl0),
  inference(forward_demodulation,[],[f73190,f559])).
fof(f73190,plain,(
  s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | $spl0),
  inference(unit_resulting_resolution,[],[f604,f532])).
fof(f71404,plain,(
  $false | ($spl0 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f619,f71335])).
fof(f71335,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f71334,f560])).
fof(f560,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_code')).
fof(f71334,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f71194,f560])).
fof(f71194,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl6)),
  inference(unit_resulting_resolution,[],[f63411,f536])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f63411,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f604,f610])).
fof(f71370,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f71369,f607])).
fof(f71369,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),sK10) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f71368,f558])).
fof(f558,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.sptree.spt'('type.bvlSem.v'),X51) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_locals')).
fof(f71368,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK10) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f71179,f558])).
fof(f71179,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl6)),
  inference(unit_resulting_resolution,[],[f63411,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f71352,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f71351,f607])).
fof(f71351,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),sK10) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f71350,f558])).
fof(f71350,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK10) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f71188,f558])).
fof(f71188,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl6)),
  inference(unit_resulting_resolution,[],[f63411,f530])).
fof(f5243,plain,(
  ( ! [X50] : (s(bool,'T_0') = s(bool,X50)) ) | $spl34),
  inference(cnf_transformation,[],[f5243_D])).
fof(f5243_D,plain,(
  ( ! [X50] : (s(bool,'T_0') = s(bool,X50)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f69076,plain,(
  ( ! [X257,X259,X260,X256] : (p(s(bool,'|dollar|exists_1'(s(fun(X257,bool),'_dst_x0_1'(s(bool,X256)))))) | s(bool,X256) = s(bool,X259) | s(bool,X259) = s(bool,X260) | s(bool,'F_0') = s(bool,X260)) ) | $spl112),
  inference(cnf_transformation,[],[f69076_D])).
fof(f69076_D,plain,(
  ( ! [X257,X259,X260,X256] : (p(s(bool,'|dollar|exists_1'(s(fun(X257,bool),'_dst_x0_1'(s(bool,X256)))))) | s(bool,X256) = s(bool,X259) | s(bool,X259) = s(bool,X260) | s(bool,'F_0') = s(bool,X260)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f66816,plain,(
  ( ! [X257,X259,X260,X256,X258] : (p(s(bool,'|dollar|exists_1'(s(fun(X257,bool),'_dst_x0_1'(s(bool,X256)))))) | s(bool,'T_0') = s(bool,X258) | s(bool,X259) = s(bool,X260) | s(bool,'F_0') = s(bool,X260) | s(bool,X256) = s(bool,X259)) )),
  inference(duplicate_literal_removal,[],[f66118])).
fof(f66118,plain,(
  ( ! [X257,X259,X260,X256,X258] : (p(s(bool,'|dollar|exists_1'(s(fun(X257,bool),'_dst_x0_1'(s(bool,X256)))))) | s(bool,'T_0') = s(bool,X258) | s(bool,'T_0') = s(bool,X258) | s(bool,X259) = s(bool,X260) | s(bool,'F_0') = s(bool,X260) | s(bool,X256) = s(bool,X259)) )),
  inference(superposition,[],[f1823,f1477])).
fof(f1477,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1462,f424])).
fof(f424,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvpSem.state_|49||49|',aHL_TRUTH)).
fof(f1462,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1457,f659])).
fof(f659,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f430,f430])).
fof(f430,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvpSem.state_|49||49|',aHL_BOOL_CASES)).
fof(f1457,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f1456,f454])).
fof(f454,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X43] : ! [X45] : s(X43,X45) = s(X43,happ(s(fun(X43,X43),'_i_0_0'),s(X43,X45)))),
  file('thm.bvpSem.state_|49||49|','_i0_ax')).
fof(f1456,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f1430,f424])).
fof(f1430,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f575,f693])).
fof(f693,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f454,f430])).
fof(f575,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f384])).
fof(f384,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X35,X36] : ((p(s(bool,X35)) => p(s(bool,X36))) => ((p(s(bool,X36)) => p(s(bool,X35))) => s(bool,X35) = s(bool,X36)))),
  file('thm.bvpSem.state_|49||49|','thm.bool.IMP_ANTISYM_AX')).
fof(f1823,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1806,f424])).
fof(f1806,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1687,f659])).
fof(f1687,plain,(
  ( ! [X21,X22] : (~p(s(bool,X22)) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'_dst_x0_1'(s(bool,X22))))))) )),
  inference(superposition,[],[f529,f580])).
fof(f580,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X43,X41] : ! [X44,X46] : s(X41,X44) = s(X41,happ(s(fun(X43,X41),'_dst_x0_1'(s(X41,X44))),s(X43,X46)))),
  file('thm.bvpSem.state_|49||49|','_dst_x0_ax')).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK19(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f383])).
fof(f383,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X25] : ! [X47,X48] : (! [X27] : (s(X25,X27) = s(X25,X48) => p(s(bool,happ(s(fun(X25,bool),X47),s(X25,X27))))) => p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),X47)))))),
  file('thm.bvpSem.state_|49||49|','thm.bool.JRH_INDUCT_UTIL')).
fof(f63380,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl14)),
  inference(subsumption_resolution,[],[f63379,f607])).
fof(f63379,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),sK10) | ($spl0 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f63378,f558])).
fof(f63378,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK10) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f63193,f558])).
fof(f63193,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60026,f530])).
fof(f60026,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f58425,f610])).
fof(f58425,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f604,f618])).
fof(f63363,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl14)),
  inference(subsumption_resolution,[],[f63362,f607])).
fof(f63362,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),sK10) | ($spl0 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f63361,f558])).
fof(f63361,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK10) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f63202,f558])).
fof(f63202,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60026,f530])).
fof(f61057,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl6 | $spl14)),
  inference(subsumption_resolution,[],[f609,f61009])).
fof(f61009,plain,(
  s('type.list.list'('type.bvpSem.stack'),sK2) = s('type.list.list'('type.bvpSem.stack'),sK11) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f61008,f557])).
fof(f557,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.list.list'('type.bvpSem.stack'),X52) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_stack')).
fof(f61008,plain,(
  s('type.list.list'('type.bvpSem.stack'),sK11) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60853,f557])).
fof(f60853,plain,(
  s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60620,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f60620,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60026,f606])).
fof(f61016,plain,(
  $false | ($spl0 | $spl2 | $spl6 | $spl14 | $spl17)),
  inference(subsumption_resolution,[],[f61015,f621])).
fof(f61015,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),sK17) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f61014,f556])).
fof(f556,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.ffi.ffi_state'(X6),X56) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_ffi')).
fof(f61014,plain,(
  s('type.ffi.ffi_state'(sK0),sK17) = s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60850,f556])).
fof(f60850,plain,(
  s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60620,f537])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f60996,plain,(
  $false | ($spl0 | $spl2 | $spl6 | $spl14 | $spl17)),
  inference(subsumption_resolution,[],[f60995,f621])).
fof(f60995,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),sK17) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60994,f556])).
fof(f60994,plain,(
  s('type.ffi.ffi_state'(sK0),sK17) = s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60859,f556])).
fof(f60859,plain,(
  s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60620,f537])).
fof(f60619,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14 | $spl17)),
  inference(subsumption_resolution,[],[f621,f60561])).
fof(f60561,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),sK17) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60560,f556])).
fof(f60560,plain,(
  s('type.ffi.ffi_state'(sK0),sK17) = s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60425,f556])).
fof(f60425,plain,(
  s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.ffi.ffi_state'(sK0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60028,f537])).
fof(f60028,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60027,f606])).
fof(f60027,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60026,f608])).
fof(f60618,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl11 | $spl14)),
  inference(subsumption_resolution,[],[f615,f60568])).
fof(f60568,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60567,f561])).
fof(f561,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_refs')).
fof(f60567,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60422,f561])).
fof(f60422,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60028,f534])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f60583,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl13 | $spl14)),
  inference(subsumption_resolution,[],[f60582,f617])).
fof(f60582,plain,(
  s('type.num.num',sK6) = s('type.num.num',sK15) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60581,f554])).
fof(f554,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.num.num',X6) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X6) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X6) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.num.num',X54) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_clock')).
fof(f60581,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60414,f554])).
fof(f60414,plain,(
  s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60028,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f60566,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl13 | $spl14)),
  inference(subsumption_resolution,[],[f60565,f617])).
fof(f60565,plain,(
  s('type.num.num',sK6) = s('type.num.num',sK15) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60564,f554])).
fof(f60564,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(forward_demodulation,[],[f60423,f554])).
fof(f60423,plain,(
  s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl14)),
  inference(unit_resulting_resolution,[],[f60028,f535])).
fof(f60025,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl13 | $spl14)),
  inference(subsumption_resolution,[],[f617,f59977])).
fof(f59977,plain,(
  s('type.num.num',sK6) = s('type.num.num',sK15) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f59976,f554])).
fof(f59976,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f59835,f554])).
fof(f59835,plain,(
  s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(unit_resulting_resolution,[],[f58429,f535])).
fof(f58429,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f58428,f606])).
fof(f58428,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f58427,f608])).
fof(f58427,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f58426,f610])).
fof(f58426,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | ($spl0 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f58425,f614])).
fof(f59988,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl19)),
  inference(subsumption_resolution,[],[f59987,f623])).
fof(f59987,plain,(
  s('type.num.num',sK9) = s('type.num.num',sK18) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f59986,f553])).
fof(f553,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.num.num',X9) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X9) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X9) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.num.num',X57) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_space')).
fof(f59986,plain,(
  s('type.num.num',sK18) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f59829,f553])).
fof(f59829,plain,(
  s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(unit_resulting_resolution,[],[f58429,f538])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f59971,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl19)),
  inference(subsumption_resolution,[],[f59970,f623])).
fof(f59970,plain,(
  s('type.num.num',sK9) = s('type.num.num',sK18) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f59969,f553])).
fof(f59969,plain,(
  s('type.num.num',sK18) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(forward_demodulation,[],[f59838,f553])).
fof(f59838,plain,(
  s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14)),
  inference(unit_resulting_resolution,[],[f58429,f538])).
fof(f58392,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16 | $spl19)),
  inference(subsumption_resolution,[],[f58391,f623])).
fof(f58391,plain,(
  s('type.num.num',sK9) = s('type.num.num',sK18) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f58390,f553])).
fof(f58390,plain,(
  s('type.num.num',sK18) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f58237,f553])).
fof(f58237,plain,(
  s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(unit_resulting_resolution,[],[f57900,f538])).
fof(f57900,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57899,f606])).
fof(f57899,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57898,f608])).
fof(f57898,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57897,f610])).
fof(f57897,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl10 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f52792,f614])).
fof(f52792,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f52791,f618])).
fof(f52791,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl16)),
  inference(forward_demodulation,[],[f604,f620])).
fof(f58377,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16 | $spl19)),
  inference(subsumption_resolution,[],[f58376,f623])).
fof(f58376,plain,(
  s('type.num.num',sK9) = s('type.num.num',sK18) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f58375,f553])).
fof(f58375,plain,(
  s('type.num.num',sK18) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f58246,f553])).
fof(f58246,plain,(
  s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl14 | $spl16)),
  inference(unit_resulting_resolution,[],[f57900,f538])).
fof(f57896,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl19)),
  inference(subsumption_resolution,[],[f623,f57851])).
fof(f57851,plain,(
  s('type.num.num',sK9) = s('type.num.num',sK18) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57850,f553])).
fof(f57850,plain,(
  s('type.num.num',sK18) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57722,f553])).
fof(f57722,plain,(
  s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(unit_resulting_resolution,[],[f52797,f538])).
fof(f52797,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f52796,f606])).
fof(f52796,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f52795,f608])).
fof(f52795,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f52794,f610])).
fof(f52794,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f52793,f614])).
fof(f52793,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))) | ($spl0 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f52792,f616])).
fof(f57871,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f57870,f613])).
fof(f57870,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK13) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57869,f555])).
fof(f555,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.num.num',X4) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X4) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X4) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.num.num',X31) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_handler')).
fof(f57869,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57708,f555])).
fof(f57708,plain,(
  s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(unit_resulting_resolution,[],[f52797,f533])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) != s('type.bvpSem.state'(X0),X2) | s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f57859,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f57858,f613])).
fof(f57858,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK13) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57857,f555])).
fof(f57857,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f57717,f555])).
fof(f57717,plain,(
  s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK18))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16)),
  inference(unit_resulting_resolution,[],[f52797,f533])).
fof(f52769,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(subsumption_resolution,[],[f52768,f613])).
fof(f52768,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK13) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f52767,f555])).
fof(f52767,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f52610,f555])).
fof(f52610,plain,(
  s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(unit_resulting_resolution,[],[f649,f533])).
fof(f649,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f648,f606])).
fof(f648,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl0 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f647,f608])).
fof(f647,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl0 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f646,f610])).
fof(f646,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl0 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f645,f614])).
fof(f645,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl0 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f644,f616])).
fof(f644,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl0 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f643,f618])).
fof(f643,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl0 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f642,f620])).
fof(f642,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK9))) | ($spl0 | $spl18)),
  inference(forward_demodulation,[],[f604,f622])).
fof(f52758,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(subsumption_resolution,[],[f52757,f613])).
fof(f52757,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK13) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f52756,f555])).
fof(f52756,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f52619,f555])).
fof(f52619,plain,(
  s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(unit_resulting_resolution,[],[f649,f533])).
fof(f43298,plain,(
  ( ! [X177,X180,X178] : (s(bool,'F_0') = s(bool,sK19('const.pair.FST_1'(s('type.pair.prod'(bool,X177),'const.pair.,_2'(s(bool,'F_0'),s(X177,X178)))),X180,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X180)))) ) | $spl110),
  inference(cnf_transformation,[],[f43298_D])).
fof(f43298_D,plain,(
  ( ! [X177,X180,X178] : (s(bool,'F_0') = s(bool,sK19('const.pair.FST_1'(s('type.pair.prod'(bool,X177),'const.pair.,_2'(s(bool,'F_0'),s(X177,X178)))),X180,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X180)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2006,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | $spl28),
  inference(cnf_transformation,[],[f2006_D])).
fof(f2006_D,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f43297,plain,(
  ( ! [X177,X180,X178,X179] : (p(s(bool,X179)) | s(bool,'F_0') = s(bool,sK19('const.pair.FST_1'(s('type.pair.prod'(bool,X177),'const.pair.,_2'(s(bool,'F_0'),s(X177,X178)))),X180,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X180)))) )),
  inference(subsumption_resolution,[],[f43226,f431])).
fof(f431,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X48] : (s(bool,'T_0') = s(bool,X48) <=> p(s(bool,X48)))),
  file('thm.bvpSem.state_|49||49|','thm.bool.EQ_CLAUSES|split|0')).
fof(f43226,plain,(
  ( ! [X177,X180,X178,X179] : (p(s(bool,X179)) | s(bool,'F_0') = s(bool,sK19('const.pair.FST_1'(s('type.pair.prod'(bool,X177),'const.pair.,_2'(s(bool,'F_0'),s(X177,X178)))),X180,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X180))) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f1532,f12237])).
fof(f12237,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,'F_0'),s(X2,X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(trivial_inequality_removal,[],[f12230])).
fof(f12230,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,'F_0'),s(X2,X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f8626,f659])).
fof(f8626,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X1)))))) )),
  inference(unit_resulting_resolution,[],[f650,f7378])).
fof(f7378,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,X7),s(X6,X8))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f7262,f426])).
fof(f426,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvpSem.state_|49||49|',aHL_FALSITY)).
fof(f7262,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,X7),s(X6,X8))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f431,f1455])).
fof(f1455,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,X25),s(X24,X26))))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(forward_demodulation,[],[f1418,f454])).
fof(f1418,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,X25),s(X24,X26)))))) )),
  inference(superposition,[],[f693,f579])).
fof(f579,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f354])).
fof(f354,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X26,X25] : ! [X27,X28] : s(X25,X27) = s(X25,'const.pair.FST_1'(s('type.pair.prod'(X25,X26),'const.pair.,_2'(s(X25,X27),s(X26,X28)))))),
  file('thm.bvpSem.state_|49||49|','thm.pair.FST')).
fof(f650,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f426,f431])).
fof(f1532,plain,(
  ( ! [X45,X46] : (p(s(bool,X45)) | s(bool,'F_0') = s(bool,sK19(X45,X46,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X46)))) )),
  inference(superposition,[],[f436,f994])).
fof(f994,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK19(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f993,f426])).
fof(f993,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK19(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f528,f430])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK19(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f405])).
fof(f436,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X48] : (s(bool,'F_0') = s(bool,X48) <=> ~p(s(bool,X48)))),
  file('thm.bvpSem.state_|49||49|','thm.bool.EQ_CLAUSES|split|3')).
fof(f43294,plain,(
  ( ! [X165,X163,X162] : (s(bool,'F_0') = s(bool,sK19(happ(s(fun(X162,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X162,X163)),X165,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X165)))) ) | $spl108),
  inference(cnf_transformation,[],[f43294_D])).
fof(f43294_D,plain,(
  ( ! [X165,X163,X162] : (s(bool,'F_0') = s(bool,sK19(happ(s(fun(X162,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X162,X163)),X165,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X165)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f43293,plain,(
  ( ! [X165,X163,X164,X162] : (p(s(bool,X164)) | s(bool,'F_0') = s(bool,sK19(happ(s(fun(X162,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X162,X163)),X165,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X165)))) )),
  inference(subsumption_resolution,[],[f43222,f431])).
fof(f43222,plain,(
  ( ! [X165,X163,X164,X162] : (p(s(bool,X164)) | s(bool,'F_0') = s(bool,sK19(happ(s(fun(X162,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X162,X163)),X165,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X165))) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f1532,f12018])).
fof(f12018,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X2,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(trivial_inequality_removal,[],[f12011])).
fof(f12011,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X2,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f8148,f659])).
fof(f8148,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f650,f7114])).
fof(f7114,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,happ(s(fun(X9,bool),'_dst_x0_1'(s(bool,X10))),s(X9,X11))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f7002,f426])).
fof(f7002,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(X9,bool),'_dst_x0_1'(s(bool,X10))),s(X9,X11))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f431,f1454])).
fof(f1454,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,happ(s(fun(X21,bool),'_dst_x0_1'(s(bool,X22))),s(X21,X23))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(forward_demodulation,[],[f1417,f454])).
fof(f1417,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,happ(s(fun(X21,bool),'_dst_x0_1'(s(bool,X22))),s(X21,X23)))) )),
  inference(superposition,[],[f693,f580])).
fof(f43291,plain,(
  ( ! [X153] : (s(bool,'F_0') = s(bool,sK19(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')),X153,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X153)))) ) | $spl106),
  inference(cnf_transformation,[],[f43291_D])).
fof(f43291_D,plain,(
  ( ! [X153] : (s(bool,'F_0') = s(bool,sK19(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')),X153,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X153)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f43290,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | s(bool,'F_0') = s(bool,sK19(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')),X153,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X153)))) ) | $spl48),
  inference(subsumption_resolution,[],[f43219,f431])).
fof(f43219,plain,(
  ( ! [X152,X153] : (p(s(bool,X152)) | s(bool,'F_0') = s(bool,sK19(happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')),X153,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X153))) | s(bool,'T_0') = s(bool,X152)) ) | $spl48),
  inference(superposition,[],[f1532,f10539])).
fof(f43288,plain,(
  ( ! [X144,X142] : (s(bool,'F_0') = s(bool,sK19('const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.|48|_0')),X144,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X144)))) ) | $spl104),
  inference(cnf_transformation,[],[f43288_D])).
fof(f43288_D,plain,(
  ( ! [X144,X142] : (s(bool,'F_0') = s(bool,sK19('const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.|48|_0')),X144,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X144)))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f43287,plain,(
  ( ! [X144,X142,X143] : (p(s(bool,X143)) | s(bool,'F_0') = s(bool,sK19('const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.|48|_0')),X144,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X144)))) )),
  inference(subsumption_resolution,[],[f43215,f431])).
fof(f43215,plain,(
  ( ! [X144,X142,X143] : (p(s(bool,X143)) | s(bool,'F_0') = s(bool,sK19('const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.|48|_0')),X144,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X144))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f1532,f11779])).
fof(f11779,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(trivial_inequality_removal,[],[f11772])).
fof(f11772,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f7857,f659])).
fof(f7857,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f650,f6826])).
fof(f6826,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6742,f426])).
fof(f6742,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f431,f1453])).
fof(f1453,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X18),s(fun('type.num.num',bool),X19),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(forward_demodulation,[],[f1415,f454])).
fof(f1415,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X18),s(fun('type.num.num',bool),X19),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f693,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X25] : ! [X29,X30] : s(X25,X29) = s(X25,'const.ind_type.FCONS_3'(s(X25,X29),s(fun('type.num.num',X25),X30),s('type.num.num','const.num.|48|_0')))),
  file('thm.bvpSem.state_|49||49|','thm.ind_type.FCONS|split|0')).
fof(f43284,plain,(
  ( ! [X132,X130,X129] : (s(bool,'F_0') = s(bool,sK19('const.pair.SND_1'(s('type.pair.prod'(X129,bool),'const.pair.,_2'(s(X129,X130),s(bool,'F_0')))),X132,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X132)))) ) | $spl102),
  inference(cnf_transformation,[],[f43284_D])).
fof(f43284_D,plain,(
  ( ! [X132,X130,X129] : (s(bool,'F_0') = s(bool,sK19('const.pair.SND_1'(s('type.pair.prod'(X129,bool),'const.pair.,_2'(s(X129,X130),s(bool,'F_0')))),X132,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X132)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f43283,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,X131)) | s(bool,'F_0') = s(bool,sK19('const.pair.SND_1'(s('type.pair.prod'(X129,bool),'const.pair.,_2'(s(X129,X130),s(bool,'F_0')))),X132,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X132)))) )),
  inference(subsumption_resolution,[],[f43211,f431])).
fof(f43211,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,X131)) | s(bool,'F_0') = s(bool,sK19('const.pair.SND_1'(s('type.pair.prod'(X129,bool),'const.pair.,_2'(s(X129,X130),s(bool,'F_0')))),X132,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X132))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f1532,f11642])).
fof(f11642,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(trivial_inequality_removal,[],[f11635])).
fof(f11635,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f7421,f659])).
fof(f7421,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0')))))) )),
  inference(unit_resulting_resolution,[],[f650,f6664])).
fof(f6664,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,X8))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f6558,f426])).
fof(f6558,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,X8))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f431,f1452])).
fof(f1452,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X15,bool),'const.pair.,_2'(s(X15,X16),s(bool,X17))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(forward_demodulation,[],[f1414,f454])).
fof(f1414,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X15,bool),'const.pair.,_2'(s(X15,X16),s(bool,X17)))))) )),
  inference(superposition,[],[f693,f578])).
fof(f578,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f352])).
fof(f352,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X25,X26] : ! [X27,X28] : s(X26,X28) = s(X26,'const.pair.SND_1'(s('type.pair.prod'(X25,X26),'const.pair.,_2'(s(X25,X27),s(X26,X28)))))),
  file('thm.bvpSem.state_|49||49|','thm.pair.SND')).
fof(f43041,plain,(
  ( ! [X218] : (s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X218,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218)))) ) | $spl100),
  inference(cnf_transformation,[],[f43041_D])).
fof(f43041_D,plain,(
  ( ! [X218] : (s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X218,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f3077,plain,(
  ( ! [X29] : (~p(s(bool,X29))) ) | $spl32),
  inference(cnf_transformation,[],[f3077_D])).
fof(f3077_D,plain,(
  ( ! [X29] : (~p(s(bool,X29))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f43042,plain,(
  ( ! [X220,X221] : (~p(s(bool,X220)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X221,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X221)))) ) | $spl22),
  inference(subsumption_resolution,[],[f42967,f435])).
fof(f435,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f392])).
fof(f42967,plain,(
  ( ! [X220,X221] : (~p(s(bool,X220)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X221,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X221))) | s(bool,'F_0') = s(bool,X220)) ) | $spl22),
  inference(superposition,[],[f1530,f11463])).
fof(f11463,plain,(
  ( ! [X33] : (s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X33)) ) | $spl22),
  inference(subsumption_resolution,[],[f11444,f424])).
fof(f11444,plain,(
  ( ! [X33] : (~p(s(bool,'T_0')) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X33)) ) | $spl22),
  inference(superposition,[],[f1457,f10810])).
fof(f10810,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl22),
  inference(unit_resulting_resolution,[],[f650,f650,f1697,f1378])).
fof(f1378,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1152,f424])).
fof(f1152,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f575,f659])).
fof(f1530,plain,(
  ( ! [X41,X42] : (~p(s(bool,X41)) | s(bool,'T_0') = s(bool,sK19(X41,X42,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X42)))) )),
  inference(superposition,[],[f432,f994])).
fof(f432,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f390])).
fof(f10540,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,X11)) ) | $spl50),
  inference(cnf_transformation,[],[f10540_D])).
fof(f10540_D,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,X11)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f43040,plain,(
  ( ! [X218,X217] : (s(bool,'F_0') = s(bool,X217) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X218,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218)))) ) | ($spl22 | $spl24 | $spl36)),
  inference(subsumption_resolution,[],[f43039,f436])).
fof(f43039,plain,(
  ( ! [X218,X217] : (s(bool,'F_0') = s(bool,X217) | ~p(s(bool,X217)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X218,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218)))) ) | ($spl22 | $spl24 | $spl36)),
  inference(forward_demodulation,[],[f42966,f12999])).
fof(f12999,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK19(X0,'_i_0_0',bool))) ) | $spl24),
  inference(unit_resulting_resolution,[],[f426,f1699,f572])).
fof(f572,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f385])).
fof(f42966,plain,(
  ( ! [X218,X219,X217] : (~p(s(bool,X217)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X218,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218))) | s(bool,X217) = s(bool,sK19(X219,'_i_0_0',bool))) ) | ($spl22 | $spl24 | $spl36)),
  inference(superposition,[],[f1530,f21741])).
fof(f21741,plain,(
  ( ! [X140,X141] : (s(bool,X140) = s(bool,sK19(X141,'_i_0_0',bool)) | s(bool,X140) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl22 | $spl24 | $spl36)),
  inference(subsumption_resolution,[],[f20621,f650])).
fof(f20621,plain,(
  ( ! [X140,X141] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X140) = s(bool,sK19(X141,'_i_0_0',bool)) | s(bool,X140) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl22 | $spl24 | $spl36)),
  inference(superposition,[],[f20151,f10810])).
fof(f20151,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) = s(bool,sK19(X20,'_i_0_0',bool)) | s(bool,X19) = s(bool,X21) | s(bool,'F_0') = s(bool,X21)) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f19727,f5273])).
fof(f5273,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl36),
  inference(unit_resulting_resolution,[],[f650,f5244])).
fof(f19727,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,X19) = s(bool,sK19(X20,'_i_0_0',bool)) | s(bool,X19) = s(bool,X21) | s(bool,'F_0') = s(bool,X21)) ) | $spl24),
  inference(superposition,[],[f13336,f10557])).
fof(f10557,plain,(
  ( ! [X39,X41] : (s(bool,X39) = s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f10556,f424])).
fof(f10556,plain,(
  ( ! [X39,X41] : (~p(s(bool,'T_0')) | s(bool,X39) = s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(forward_demodulation,[],[f4364,f1846])).
fof(f1846,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f650,f1805,f1457])).
fof(f1805,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f424,f1687])).
fof(f4364,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'|dollar|exists_1'(s(fun(X40,bool),'_dst_x0_1'(s(bool,'T_0')))))) | s(bool,X39) = s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f1457,f1962])).
fof(f1962,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun(X25,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f1936,f424])).
fof(f1936,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun(X25,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f1457,f1846])).
fof(f13336,plain,(
  ( ! [X48,X49] : (s(bool,X48) = s(bool,sK19(X49,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X48)))) ) | $spl24),
  inference(superposition,[],[f13017,f454])).
fof(f13017,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,sK19(X1,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(subsumption_resolution,[],[f13003,f424])).
fof(f13003,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,sK19(X1,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(superposition,[],[f1699,f659])).
fof(f43036,plain,(
  ( ! [X210,X208] : (s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun('type.pair.prod'(bool,X208),bool),'const.pair.FST_0')),X210,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X210)))) ) | $spl98),
  inference(cnf_transformation,[],[f43036_D])).
fof(f43036_D,plain,(
  ( ! [X210,X208] : (s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun('type.pair.prod'(bool,X208),bool),'const.pair.FST_0')),X210,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X210)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f43037,plain,(
  ( ! [X212,X213,X214] : (~p(s(bool,X213)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun('type.pair.prod'(bool,X212),bool),'const.pair.FST_0')),X214,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X214)))) ) | $spl30),
  inference(subsumption_resolution,[],[f42964,f435])).
fof(f42964,plain,(
  ( ! [X212,X213,X214] : (~p(s(bool,X213)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun('type.pair.prod'(bool,X212),bool),'const.pair.FST_0')),X214,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X214))) | s(bool,'F_0') = s(bool,X213)) ) | $spl30),
  inference(superposition,[],[f1530,f10581])).
fof(f10581,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X26)) ) | $spl30),
  inference(subsumption_resolution,[],[f3299,f424])).
fof(f3299,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X26)) ) | $spl30),
  inference(superposition,[],[f1457,f3245])).
fof(f3245,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f650,f3076,f1457])).
fof(f43035,plain,(
  ( ! [X210,X208,X209] : (s(bool,'F_0') = s(bool,X209) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun('type.pair.prod'(bool,X208),bool),'const.pair.FST_0')),X210,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X210)))) ) | ($spl24 | $spl30 | $spl36)),
  inference(subsumption_resolution,[],[f43034,f436])).
fof(f43034,plain,(
  ( ! [X210,X208,X209] : (s(bool,'F_0') = s(bool,X209) | ~p(s(bool,X209)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun('type.pair.prod'(bool,X208),bool),'const.pair.FST_0')),X210,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X210)))) ) | ($spl24 | $spl30 | $spl36)),
  inference(forward_demodulation,[],[f42963,f12999])).
fof(f42963,plain,(
  ( ! [X210,X208,X211,X209] : (~p(s(bool,X209)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun('type.pair.prod'(bool,X208),bool),'const.pair.FST_0')),X210,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X210))) | s(bool,X209) = s(bool,sK19(X211,'_i_0_0',bool))) ) | ($spl24 | $spl30 | $spl36)),
  inference(superposition,[],[f1530,f21740])).
fof(f21740,plain,(
  ( ! [X134,X136,X135] : (s(bool,X135) = s(bool,sK19(X136,'_i_0_0',bool)) | s(bool,X135) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X134),bool),'const.pair.FST_0')))) ) | ($spl24 | $spl30 | $spl36)),
  inference(subsumption_resolution,[],[f20619,f650])).
fof(f20619,plain,(
  ( ! [X134,X136,X135] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X135) = s(bool,sK19(X136,'_i_0_0',bool)) | s(bool,X135) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X134),bool),'const.pair.FST_0')))) ) | ($spl24 | $spl30 | $spl36)),
  inference(superposition,[],[f20151,f3245])).
fof(f43027,plain,(
  ( ! [X187,X185] : (s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(X185,bool),'_dst_x0_1'(s(bool,'T_0')))),X187,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X187)))) ) | $spl96),
  inference(cnf_transformation,[],[f43027_D])).
fof(f43027_D,plain,(
  ( ! [X187,X185] : (s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(X185,bool),'_dst_x0_1'(s(bool,'T_0')))),X187,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X187)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f43028,plain,(
  ( ! [X191,X189,X190] : (~p(s(bool,X190)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(X189,bool),'_dst_x0_1'(s(bool,'T_0')))),X191,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X191)))) )),
  inference(subsumption_resolution,[],[f42957,f435])).
fof(f42957,plain,(
  ( ! [X191,X189,X190] : (~p(s(bool,X190)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(X189,bool),'_dst_x0_1'(s(bool,'T_0')))),X191,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X191))) | s(bool,'F_0') = s(bool,X190)) )),
  inference(superposition,[],[f1530,f1962])).
fof(f43026,plain,(
  ( ! [X187,X185,X186] : (s(bool,'F_0') = s(bool,X186) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(X185,bool),'_dst_x0_1'(s(bool,'T_0')))),X187,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X187)))) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f43025,f436])).
fof(f43025,plain,(
  ( ! [X187,X185,X186] : (s(bool,'F_0') = s(bool,X186) | ~p(s(bool,X186)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(X185,bool),'_dst_x0_1'(s(bool,'T_0')))),X187,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X187)))) ) | ($spl24 | $spl36)),
  inference(forward_demodulation,[],[f42956,f12999])).
fof(f42956,plain,(
  ( ! [X187,X185,X188,X186] : (~p(s(bool,X186)) | s(bool,'T_0') = s(bool,sK19('|dollar|exists_1'(s(fun(X185,bool),'_dst_x0_1'(s(bool,'T_0')))),X187,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X187))) | s(bool,X186) = s(bool,sK19(X188,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f1530,f21738])).
fof(f21738,plain,(
  ( ! [X125,X123,X124] : (s(bool,X124) = s(bool,sK19(X125,'_i_0_0',bool)) | s(bool,X124) = s(bool,'|dollar|exists_1'(s(fun(X123,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f20616,f650])).
fof(f20616,plain,(
  ( ! [X125,X123,X124] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X124) = s(bool,sK19(X125,'_i_0_0',bool)) | s(bool,X124) = s(bool,'|dollar|exists_1'(s(fun(X123,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f20151,f1846])).
fof(f39314,plain,(
  $false | ($spl24 | $spl60 | $spl70 | $spl92)),
  inference(global_subsumption,[],[f39313,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39272,plain,(
  ( ! [X82] : (s(bool,'T_0') = s(bool,X82)) ) | $spl92),
  inference(subsumption_resolution,[],[f38593,f38277])).
fof(f38593,plain,(
  ( ! [X83,X84,X82] : (s(bool,sK19(X83,'_i_0_0',bool)) != s(bool,'|dollar|exists_1'(s(fun(X84,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X82)) ) | $spl92),
  inference(superposition,[],[f1961,f38277])).
fof(f38532,plain,(
  ( ! [X33,X32] : (s(bool,X32) = s(bool,X33)) ) | $spl92),
  inference(superposition,[],[f38277,f38277])).
fof(f37789,plain,(
  ( ! [X99,X100] : (s(bool,X99) != s(bool,happ(s(fun('type.pair.prod'(bool,bool),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X99))))) | s(bool,X99) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f31374,f659])).
fof(f37737,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X13),s(X14,X15))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f31374,f430])).
fof(f37735,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(bool,X6),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,X5),s(X6,X7))))) | s(bool,X5) = s(bool,X8) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f31374,f10557])).
fof(f32044,plain,(
  ( ! [X233,X234,X232] : (s(bool,X232) != s(bool,sK19(X234,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK19(X232,X233,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X233)))) ) | $spl60),
  inference(superposition,[],[f13580,f994])).
fof(f31857,plain,(
  ( ! [X313,X315,X316,X314] : (p(s(bool,X313)) | s(bool,X315) = s(bool,sK19(X313,X314,bool)) | s(bool,X315) = s(bool,X316) | s(bool,'F_0') = s(bool,X316) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X314)))) )),
  inference(superposition,[],[f14495,f994])).
fof(f31902,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X204) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X201),'const.pair.,_2'(s(bool,X202),s(X201,X203))))) | s(bool,X204) = s(bool,X205) | s(bool,'F_0') = s(bool,X205) | s(bool,'T_0') = s(bool,X202)) )),
  inference(subsumption_resolution,[],[f31830,f426])).
fof(f31830,plain,(
  ( ! [X204,X202,X205,X203,X201] : (p(s(bool,'F_0')) | s(bool,X204) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X201),'const.pair.,_2'(s(bool,X202),s(X201,X203))))) | s(bool,X204) = s(bool,X205) | s(bool,'F_0') = s(bool,X205) | s(bool,'T_0') = s(bool,X202)) )),
  inference(superposition,[],[f14495,f1455])).
fof(f31901,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,X199) = s(bool,happ(s(fun('type.pair.prod'(bool,X197),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X197),'const.pair.,_2'(s(bool,'F_0'),s(X197,X198))))) | s(bool,X199) = s(bool,X200) | s(bool,'F_0') = s(bool,X200)) )),
  inference(subsumption_resolution,[],[f31829,f426])).
fof(f31829,plain,(
  ( ! [X198,X200,X199,X197] : (p(s(bool,'F_0')) | s(bool,X199) = s(bool,happ(s(fun('type.pair.prod'(bool,X197),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X197),'const.pair.,_2'(s(bool,'F_0'),s(X197,X198))))) | s(bool,X199) = s(bool,X200) | s(bool,'F_0') = s(bool,X200)) )),
  inference(superposition,[],[f14495,f12877])).
fof(f31899,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,X185) = s(bool,happ(s(fun(X182,bool),'_dst_x0_1'(s(bool,X183))),s(X182,X184))) | s(bool,X185) = s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,'T_0') = s(bool,X183)) )),
  inference(subsumption_resolution,[],[f31826,f426])).
fof(f31826,plain,(
  ( ! [X185,X182,X186,X184,X183] : (p(s(bool,'F_0')) | s(bool,X185) = s(bool,happ(s(fun(X182,bool),'_dst_x0_1'(s(bool,X183))),s(X182,X184))) | s(bool,X185) = s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,'T_0') = s(bool,X183)) )),
  inference(superposition,[],[f14495,f1454])).
fof(f31895,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,X163) = s(bool,'const.ind_type.FCONS_3'(s(bool,X161),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X162))))) | s(bool,X163) = s(bool,X164) | s(bool,'F_0') = s(bool,X164)) )),
  inference(subsumption_resolution,[],[f31820,f426])).
fof(f31820,plain,(
  ( ! [X163,X161,X164,X162] : (p(s(bool,'F_0')) | s(bool,X163) = s(bool,'const.ind_type.FCONS_3'(s(bool,X161),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X162))))) | s(bool,X163) = s(bool,X164) | s(bool,'F_0') = s(bool,X164)) )),
  inference(superposition,[],[f14495,f18667])).
fof(f31894,plain,(
  ( ! [X156,X154,X155,X153] : (s(bool,X155) = s(bool,'const.ind_type.FCONS_3'(s(bool,X153),s(fun('type.num.num',bool),X154),s('type.num.num','const.num.|48|_0'))) | s(bool,X155) = s(bool,X156) | s(bool,'F_0') = s(bool,X156) | s(bool,'T_0') = s(bool,X153)) )),
  inference(subsumption_resolution,[],[f31818,f426])).
fof(f31818,plain,(
  ( ! [X156,X154,X155,X153] : (p(s(bool,'F_0')) | s(bool,X155) = s(bool,'const.ind_type.FCONS_3'(s(bool,X153),s(fun('type.num.num',bool),X154),s('type.num.num','const.num.|48|_0'))) | s(bool,X155) = s(bool,X156) | s(bool,'F_0') = s(bool,X156) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f14495,f1453])).
fof(f31892,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,X141) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X138,bool),'const.pair.,_2'(s(X138,X139),s(bool,X140))))) | s(bool,X141) = s(bool,X142) | s(bool,'F_0') = s(bool,X142) | s(bool,'T_0') = s(bool,X140)) )),
  inference(subsumption_resolution,[],[f31815,f426])).
fof(f31815,plain,(
  ( ! [X142,X140,X138,X141,X139] : (p(s(bool,'F_0')) | s(bool,X141) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X138,bool),'const.pair.,_2'(s(X138,X139),s(bool,X140))))) | s(bool,X141) = s(bool,X142) | s(bool,'F_0') = s(bool,X142) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f14495,f1452])).
fof(f31594,plain,(
  ( ! [X313,X315,X316,X314] : (~p(s(bool,X313)) | s(bool,X315) != s(bool,sK19(X313,X314,bool)) | s(bool,X315) = s(bool,X316) | s(bool,'F_0') = s(bool,X316) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X314)))) )),
  inference(superposition,[],[f13960,f994])).
fof(f31736,plain,(
  ( ! [X235,X233,X234,X232] : (s(bool,X234) != s(bool,'|dollar|exists_1'(s(fun(X232,bool),'_dst_x0_1'(s(bool,X233))))) | s(bool,X234) = s(bool,X235) | s(bool,'F_0') = s(bool,X235) | s(bool,'F_0') = s(bool,X233)) )),
  inference(subsumption_resolution,[],[f31574,f424])).
fof(f31574,plain,(
  ( ! [X235,X233,X234,X232] : (~p(s(bool,'T_0')) | s(bool,X234) != s(bool,'|dollar|exists_1'(s(fun(X232,bool),'_dst_x0_1'(s(bool,X233))))) | s(bool,X234) = s(bool,X235) | s(bool,'F_0') = s(bool,X235) | s(bool,'F_0') = s(bool,X233)) )),
  inference(superposition,[],[f13960,f2008])).
fof(f31733,plain,(
  ( ! [X218,X216,X220,X219,X217] : (s(bool,X218) != s(bool,'|dollar|exists_1'(s(fun(X216,bool),X217))) | s(bool,X218) = s(bool,X219) | s(bool,'F_0') = s(bool,X219) | s(X216,X220) = s(X216,sK19(X220,X217,X216))) )),
  inference(subsumption_resolution,[],[f31570,f424])).
fof(f31570,plain,(
  ( ! [X218,X216,X220,X219,X217] : (~p(s(bool,'T_0')) | s(bool,X218) != s(bool,'|dollar|exists_1'(s(fun(X216,bool),X217))) | s(bool,X218) = s(bool,X219) | s(bool,'F_0') = s(bool,X219) | s(X216,X220) = s(X216,sK19(X220,X217,X216))) )),
  inference(superposition,[],[f13960,f994])).
fof(f31730,plain,(
  ( ! [X210,X208,X206,X209,X207] : (s(bool,X209) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X206),'const.pair.,_2'(s(bool,X207),s(X206,X208))))) | s(bool,X209) = s(bool,X210) | s(bool,'F_0') = s(bool,X210) | s(bool,'F_0') = s(bool,X207)) )),
  inference(subsumption_resolution,[],[f31568,f424])).
fof(f31568,plain,(
  ( ! [X210,X208,X206,X209,X207] : (~p(s(bool,'T_0')) | s(bool,X209) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X206),'const.pair.,_2'(s(bool,X207),s(X206,X208))))) | s(bool,X209) = s(bool,X210) | s(bool,'F_0') = s(bool,X210) | s(bool,'F_0') = s(bool,X207)) )),
  inference(superposition,[],[f13960,f868])).
fof(f31727,plain,(
  ( ! [X191,X189,X187,X190,X188] : (s(bool,X190) != s(bool,happ(s(fun(X187,bool),'_dst_x0_1'(s(bool,X188))),s(X187,X189))) | s(bool,X190) = s(bool,X191) | s(bool,'F_0') = s(bool,X191) | s(bool,'F_0') = s(bool,X188)) )),
  inference(subsumption_resolution,[],[f31564,f424])).
fof(f31564,plain,(
  ( ! [X191,X189,X187,X190,X188] : (~p(s(bool,'T_0')) | s(bool,X190) != s(bool,happ(s(fun(X187,bool),'_dst_x0_1'(s(bool,X188))),s(X187,X189))) | s(bool,X190) = s(bool,X191) | s(bool,'F_0') = s(bool,X191) | s(bool,'F_0') = s(bool,X188)) )),
  inference(superposition,[],[f13960,f932])).
fof(f31724,plain,(
  ( ! [X177,X178,X176] : (s(bool,X177) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X176))) | s(bool,X177) = s(bool,X178) | s(bool,'F_0') = s(bool,X178) | s(bool,'F_0') = s(bool,X176)) )),
  inference(subsumption_resolution,[],[f31561,f424])).
fof(f31561,plain,(
  ( ! [X177,X178,X176] : (~p(s(bool,'T_0')) | s(bool,X177) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X176))) | s(bool,X177) = s(bool,X178) | s(bool,'F_0') = s(bool,X178) | s(bool,'F_0') = s(bool,X176)) )),
  inference(superposition,[],[f13960,f693])).
fof(f31720,plain,(
  ( ! [X158,X159,X157,X160] : (s(bool,X159) != s(bool,'const.ind_type.FCONS_3'(s(bool,X157),s(fun('type.num.num',bool),X158),s('type.num.num','const.num.|48|_0'))) | s(bool,X159) = s(bool,X160) | s(bool,'F_0') = s(bool,X160) | s(bool,'F_0') = s(bool,X157)) )),
  inference(subsumption_resolution,[],[f31556,f424])).
fof(f31556,plain,(
  ( ! [X158,X159,X157,X160] : (~p(s(bool,'T_0')) | s(bool,X159) != s(bool,'const.ind_type.FCONS_3'(s(bool,X157),s(fun('type.num.num',bool),X158),s('type.num.num','const.num.|48|_0'))) | s(bool,X159) = s(bool,X160) | s(bool,'F_0') = s(bool,X160) | s(bool,'F_0') = s(bool,X157)) )),
  inference(superposition,[],[f13960,f752])).
fof(f31717,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,X146) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X143,bool),'const.pair.,_2'(s(X143,X144),s(bool,X145))))) | s(bool,X146) = s(bool,X147) | s(bool,'F_0') = s(bool,X147) | s(bool,'F_0') = s(bool,X145)) )),
  inference(subsumption_resolution,[],[f31553,f424])).
fof(f31553,plain,(
  ( ! [X146,X144,X147,X145,X143] : (~p(s(bool,'T_0')) | s(bool,X146) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X143,bool),'const.pair.,_2'(s(X143,X144),s(bool,X145))))) | s(bool,X146) = s(bool,X147) | s(bool,'F_0') = s(bool,X147) | s(bool,'F_0') = s(bool,X145)) )),
  inference(superposition,[],[f13960,f808])).
fof(f31426,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.bvpSem.state'(X31),X32) != s('type.bvpSem.state'(X31),X34) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X31),X34))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X31),sK19(X32,X33,'type.bvpSem.state'(X31))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X31),bool),X33)))) )),
  inference(superposition,[],[f536,f994])).
fof(f31374,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5)))))) )),
  inference(subsumption_resolution,[],[f31210,f426])).
fof(f31210,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5)))))) )),
  inference(superposition,[],[f431,f12877])).
fof(f31363,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f31185,f650])).
fof(f31185,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X6) = s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f12877,f659])).
fof(f31031,plain,(
  ( ! [X14,X15,X13] : (s(bool,X13) = s(bool,happ(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X13),s(X14,X15))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f12877,f430])).
fof(f31029,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(bool,X6),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,X5),s(X6,X7))))) | s(bool,X5) = s(bool,X8) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f12877,f10557])).
fof(f30839,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.bvpSem.state'(X31),X32) != s('type.bvpSem.state'(X31),X34) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X31),X34))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X31),sK19(X32,X33,'type.bvpSem.state'(X31))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X31),bool),X33)))) )),
  inference(superposition,[],[f534,f994])).
fof(f29580,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X7) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X6),s(X8,X9))))) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X10) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f12237,f10557])).
fof(f29578,plain,(
  ( ! [X308,X306,X309,X307] : (s(bool,X306) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X308),'const.pair.,_2'(s(bool,'F_0'),s(X308,X309))))) | s(bool,'T_0') = s(bool,sK19(X306,X307,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X307)))) )),
  inference(superposition,[],[f12237,f994])).
fof(f29577,plain,(
  ( ! [X302,X304,X301,X303,X305] : (s(bool,X302) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X304),'const.pair.,_2'(s(bool,'F_0'),s(X304,X305))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X301),'const.pair.,_2'(s(bool,X302),s(X301,X303)))))) )),
  inference(superposition,[],[f12237,f579])).
fof(f29575,plain,(
  ( ! [X292,X294,X296,X293,X295] : (s(bool,X293) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X295),'const.pair.,_2'(s(bool,'F_0'),s(X295,X296))))) | s(bool,'T_0') = s(bool,happ(s(fun(X292,bool),'_dst_x0_1'(s(bool,X293))),s(X292,X294)))) )),
  inference(superposition,[],[f12237,f580])).
fof(f29572,plain,(
  ( ! [X280,X282,X281,X283] : (s(bool,X280) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X282),'const.pair.,_2'(s(bool,'F_0'),s(X282,X283))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X280),s(fun('type.num.num',bool),X281),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f12237,f457])).
fof(f29571,plain,(
  ( ! [X277,X279,X275,X276,X278] : (s(bool,X277) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X278),'const.pair.,_2'(s(bool,'F_0'),s(X278,X279))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X275,bool),'const.pair.,_2'(s(X275,X276),s(bool,X277)))))) )),
  inference(superposition,[],[f12237,f578])).
fof(f28760,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X7) = s(bool,happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,X6))),s(X8,X9))) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X10) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f12018,f10557])).
fof(f28758,plain,(
  ( ! [X296,X298,X297,X299] : (s(bool,X296) = s(bool,happ(s(fun(X298,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X298,X299))) | s(bool,'T_0') = s(bool,sK19(X296,X297,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X297)))) )),
  inference(superposition,[],[f12018,f994])).
fof(f28757,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s(bool,X292) = s(bool,happ(s(fun(X294,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X294,X295))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X291),'const.pair.,_2'(s(bool,X292),s(X291,X293)))))) )),
  inference(superposition,[],[f12018,f579])).
fof(f28755,plain,(
  ( ! [X282,X285,X283,X284,X286] : (s(bool,X283) = s(bool,happ(s(fun(X285,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X285,X286))) | s(bool,'T_0') = s(bool,happ(s(fun(X282,bool),'_dst_x0_1'(s(bool,X283))),s(X282,X284)))) )),
  inference(superposition,[],[f12018,f580])).
fof(f28752,plain,(
  ( ! [X271,X273,X270,X272] : (s(bool,X270) = s(bool,happ(s(fun(X272,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X272,X273))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X270),s(fun('type.num.num',bool),X271),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f12018,f457])).
fof(f28751,plain,(
  ( ! [X269,X265,X267,X268,X266] : (s(bool,X267) = s(bool,happ(s(fun(X268,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X268,X269))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X265,bool),'const.pair.,_2'(s(X265,X266),s(bool,X267)))))) )),
  inference(superposition,[],[f12018,f578])).
fof(f28142,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X6) = s(bool,'const.ind_type.FCONS_3'(s(bool,X5),s(fun('type.num.num',bool),X7),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X6) | s(bool,X5) = s(bool,X8) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f11779,f10557])).
fof(f28140,plain,(
  ( ! [X218,X216,X217] : (s(bool,X216) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X218),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,sK19(X216,X217,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X217)))) )),
  inference(superposition,[],[f11779,f994])).
fof(f28139,plain,(
  ( ! [X212,X215,X213,X214] : (s(bool,X213) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X215),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X212),'const.pair.,_2'(s(bool,X213),s(X212,X214)))))) )),
  inference(superposition,[],[f11779,f579])).
fof(f28137,plain,(
  ( ! [X208,X206,X207,X205] : (s(bool,X206) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X208),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X205,bool),'_dst_x0_1'(s(bool,X206))),s(X205,X207)))) )),
  inference(superposition,[],[f11779,f580])).
fof(f28134,plain,(
  ( ! [X198,X196,X197] : (s(bool,X196) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X198),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X196),s(fun('type.num.num',bool),X197),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f11779,f457])).
fof(f28133,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,X194) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X195),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X192,bool),'const.pair.,_2'(s(X192,X193),s(bool,X194)))))) )),
  inference(superposition,[],[f11779,f578])).
fof(f27471,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X7) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,X9),s(bool,X6))))) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X10) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f11642,f10557])).
fof(f27374,plain,(
  ( ! [X123,X121,X122] : (s(bool,X122) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X121),s(bool,X121))))) | s(bool,'T_0') = s(bool,X122) | s(bool,X121) = s(bool,X123) | s(bool,'F_0') = s(bool,X123)) )),
  inference(superposition,[],[f11642,f10557])).
fof(f27338,plain,(
  ( ! [X280,X279,X281,X278] : (s(bool,X278) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X280,bool),'const.pair.,_2'(s(X280,X281),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK19(X278,X279,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X279)))) )),
  inference(superposition,[],[f11642,f994])).
fof(f27337,plain,(
  ( ! [X277,X273,X275,X276,X274] : (s(bool,X274) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X276,bool),'const.pair.,_2'(s(X276,X277),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X273),'const.pair.,_2'(s(bool,X274),s(X273,X275)))))) )),
  inference(superposition,[],[f11642,f579])).
fof(f27335,plain,(
  ( ! [X265,X267,X268,X264,X266] : (s(bool,X265) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X267,bool),'const.pair.,_2'(s(X267,X268),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(X264,bool),'_dst_x0_1'(s(bool,X265))),s(X264,X266)))) )),
  inference(superposition,[],[f11642,f580])).
fof(f27332,plain,(
  ( ! [X255,X253,X254,X252] : (s(bool,X252) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X254,bool),'const.pair.,_2'(s(X254,X255),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X252),s(fun('type.num.num',bool),X253),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f11642,f457])).
fof(f27331,plain,(
  ( ! [X251,X249,X250,X248,X247] : (s(bool,X249) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X250,bool),'const.pair.,_2'(s(X250,X251),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X247,bool),'const.pair.,_2'(s(X247,X248),s(bool,X249)))))) )),
  inference(superposition,[],[f11642,f578])).
fof(f1693,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f1674,f431])).
fof(f1674,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f529,f994])).
fof(f2298,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,happ(s(fun(X22,bool),X23),s(X22,X25)))) | p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),sK19(X23,X24,fun(X22,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X22,bool),bool),X24)))) )),
  inference(superposition,[],[f1693,f994])).
fof(f2359,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2339,f424])).
fof(f2339,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1693,f659])).
fof(f10555,plain,(
  ( ! [X101,X102,X100] : (p(s(bool,'|dollar|exists_1'(s(fun(X100,bool),X101)))) | s(bool,'F_0') = s(bool,happ(s(fun(X100,bool),X101),s(X100,X102)))) )),
  inference(subsumption_resolution,[],[f10554,f424])).
fof(f10554,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X100,bool),X101)))) | s(bool,'F_0') = s(bool,happ(s(fun(X100,bool),X101),s(X100,X102)))) )),
  inference(forward_demodulation,[],[f4381,f1846])).
fof(f4381,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,'|dollar|exists_1'(s(fun(X103,bool),'_dst_x0_1'(s(bool,'T_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(X100,bool),X101)))) | s(bool,'F_0') = s(bool,happ(s(fun(X100,bool),X101),s(X100,X102)))) )),
  inference(superposition,[],[f1693,f1962])).
fof(f16373,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27)))) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),sK19(X27,X28,fun(X26,bool))),s(X26,X29))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f10555,f994])).
fof(f18815,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f550,f430])).
fof(f18816,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',X9))) | s(bool,X10) = s(bool,'const.ind_type.FCONS_3'(s(bool,X7),s(fun('type.num.num',bool),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f550,f659])).
fof(f18817,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,happ(s(fun('type.num.num',bool),X12),s('type.num.num',X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X11),s(fun('type.num.num',bool),X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13)))))) )),
  inference(superposition,[],[f550,f659])).
fof(f18822,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X31),s('type.num.num',X32))) | s(bool,X33) = s(bool,'const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f550,f10557])).
fof(f18824,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X41) = s(bool,happ(s(fun('type.num.num',bool),X39),s('type.num.num',X40))) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X38),s(fun('type.num.num',bool),X39),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40)))))) )),
  inference(superposition,[],[f550,f10557])).
fof(f18863,plain,(
  ( ! [X156,X154,X155] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X155),s('type.num.num',X156))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X154),s(fun('type.num.num',bool),X155),s('type.num.num','const.num.SUC_1'(s('type.num.num',X156)))))) )),
  inference(superposition,[],[f693,f550])).
fof(f18865,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun('type.num.num',bool),X162),s('type.num.num',X163))),s(fun('type.num.num',bool),X164),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X161),s(fun('type.num.num',bool),X162),s('type.num.num','const.num.SUC_1'(s('type.num.num',X163)))))) )),
  inference(superposition,[],[f752,f550])).
fof(f18866,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X168,bool),'const.pair.,_2'(s(X168,X169),s(bool,happ(s(fun('type.num.num',bool),X166),s('type.num.num',X167))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X165),s(fun('type.num.num',bool),X166),s('type.num.num','const.num.SUC_1'(s('type.num.num',X167)))))) )),
  inference(superposition,[],[f808,f550])).
fof(f18867,plain,(
  ( ! [X173,X171,X174,X172,X170] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X173),'const.pair.,_2'(s(bool,happ(s(fun('type.num.num',bool),X171),s('type.num.num',X172))),s(X173,X174))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X170),s(fun('type.num.num',bool),X171),s('type.num.num','const.num.SUC_1'(s('type.num.num',X172)))))) )),
  inference(superposition,[],[f868,f550])).
fof(f18868,plain,(
  ( ! [X177,X175,X178,X176,X179] : (s(bool,'T_0') = s(bool,happ(s(fun(X178,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.num.num',bool),X176),s('type.num.num',X177))))),s(X178,X179))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X175),s(fun('type.num.num',bool),X176),s('type.num.num','const.num.SUC_1'(s('type.num.num',X177)))))) )),
  inference(superposition,[],[f932,f550])).
fof(f18870,plain,(
  ( ! [X189,X187,X185,X188,X186] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X188,bool),'const.pair.,_2'(s(X188,X189),s(bool,happ(s(fun('type.num.num',bool),X186),s('type.num.num',X187))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X185),s(fun('type.num.num',bool),X186),s('type.num.num','const.num.SUC_1'(s('type.num.num',X187)))))) )),
  inference(superposition,[],[f1452,f550])).
fof(f18871,plain,(
  ( ! [X191,X192,X190,X193] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun('type.num.num',bool),X191),s('type.num.num',X192))),s(fun('type.num.num',bool),X193),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X190),s(fun('type.num.num',bool),X191),s('type.num.num','const.num.SUC_1'(s('type.num.num',X192)))))) )),
  inference(superposition,[],[f1453,f550])).
fof(f18872,plain,(
  ( ! [X198,X196,X194,X197,X195] : (s(bool,'F_0') = s(bool,happ(s(fun(X197,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.num.num',bool),X195),s('type.num.num',X196))))),s(X197,X198))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X194),s(fun('type.num.num',bool),X195),s('type.num.num','const.num.SUC_1'(s('type.num.num',X196)))))) )),
  inference(superposition,[],[f1454,f550])).
fof(f18873,plain,(
  ( ! [X202,X200,X199,X203,X201] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X202),'const.pair.,_2'(s(bool,happ(s(fun('type.num.num',bool),X200),s('type.num.num',X201))),s(X202,X203))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X199),s(fun('type.num.num',bool),X200),s('type.num.num','const.num.SUC_1'(s('type.num.num',X201)))))) )),
  inference(superposition,[],[f1455,f550])).
fof(f18879,plain,(
  ( ! [X227,X225,X226,X224] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X227,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.num.num',bool),X225),s('type.num.num',X226))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X224),s(fun('type.num.num',bool),X225),s('type.num.num','const.num.SUC_1'(s('type.num.num',X226)))))) )),
  inference(superposition,[],[f2008,f550])).
fof(f18881,plain,(
  ( ! [X235,X233,X234,X232] : (s(bool,X235) != s(bool,happ(s(fun('type.num.num',bool),X233),s('type.num.num',X234))) | s(bool,'T_0') = s(bool,X235) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X232),s(fun('type.num.num',bool),X233),s('type.num.num','const.num.SUC_1'(s('type.num.num',X234)))))) )),
  inference(superposition,[],[f4407,f550])).
fof(f18882,plain,(
  ( ! [X239,X237,X238,X236] : (s(bool,X239) != s(bool,happ(s(fun('type.num.num',bool),X237),s('type.num.num',X238))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X236),s(fun('type.num.num',bool),X237),s('type.num.num','const.num.SUC_1'(s('type.num.num',X238))))) | s(bool,'F_0') = s(bool,X239)) )),
  inference(superposition,[],[f4407,f550])).
fof(f18884,plain,(
  ( ! [X243,X246,X244,X247,X245] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X246,bool),'const.pair.,_2'(s(X246,X247),s(bool,happ(s(fun('type.num.num',bool),X244),s('type.num.num',X245))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X243),s(fun('type.num.num',bool),X244),s('type.num.num','const.num.SUC_1'(s('type.num.num',X245)))))) )),
  inference(superposition,[],[f6664,f550])).
fof(f18885,plain,(
  ( ! [X251,X249,X250,X248] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun('type.num.num',bool),X249),s('type.num.num',X250))),s(fun('type.num.num',bool),X251),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X248),s(fun('type.num.num',bool),X249),s('type.num.num','const.num.SUC_1'(s('type.num.num',X250)))))) )),
  inference(superposition,[],[f6826,f550])).
fof(f18886,plain,(
  ( ! [X255,X253,X254,X256,X252] : (s(bool,'T_0') != s(bool,happ(s(fun(X255,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.num.num',bool),X253),s('type.num.num',X254))))),s(X255,X256))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X252),s(fun('type.num.num',bool),X253),s('type.num.num','const.num.SUC_1'(s('type.num.num',X254)))))) )),
  inference(superposition,[],[f7114,f550])).
fof(f18887,plain,(
  ( ! [X261,X257,X259,X260,X258] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X260),'const.pair.,_2'(s(bool,happ(s(fun('type.num.num',bool),X258),s('type.num.num',X259))),s(X260,X261))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X257),s(fun('type.num.num',bool),X258),s('type.num.num','const.num.SUC_1'(s('type.num.num',X259)))))) )),
  inference(superposition,[],[f7378,f550])).
fof(f18901,plain,(
  ( ! [X308,X310,X309,X311] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X309),s('type.num.num',X310))) | s(bool,X311) = s(bool,'const.ind_type.FCONS_3'(s(bool,X308),s(fun('type.num.num',bool),X309),s('type.num.num','const.num.SUC_1'(s('type.num.num',X310))))) | s(bool,'F_0') = s(bool,X311)) )),
  inference(superposition,[],[f13961,f550])).
fof(f19022,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X24),s('type.num.num',X25))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X23),s(fun('type.num.num',bool),X24),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))) )),
  inference(forward_demodulation,[],[f18820,f1846])).
fof(f18820,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'|dollar|exists_1'(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))))) = s(bool,happ(s(fun('type.num.num',bool),X24),s('type.num.num',X25))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X23),s(fun('type.num.num',bool),X24),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f550,f1962])).
fof(f19040,plain,(
  ( ! [X218,X216,X217] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X217),s('type.num.num',X218))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X216),s(fun('type.num.num',bool),X217),s('type.num.num','const.num.SUC_1'(s('type.num.num',X218)))))) )),
  inference(forward_demodulation,[],[f18877,f1846])).
fof(f18877,plain,(
  ( ! [X218,X216,X219,X217] : (s(bool,'|dollar|exists_1'(s(fun(X219,bool),'_dst_x0_1'(s(bool,'T_0'))))) != s(bool,happ(s(fun('type.num.num',bool),X217),s('type.num.num',X218))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X216),s(fun('type.num.num',bool),X217),s('type.num.num','const.num.SUC_1'(s('type.num.num',X218)))))) )),
  inference(superposition,[],[f1961,f550])).
fof(f2957,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4)))) )),
  inference(superposition,[],[f576,f430])).
fof(f2958,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X5),X6))) | s(bool,X7) = s(bool,happ(s(fun('type.pair.prod'(bool,X5),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X5),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f576,f659])).
fof(f2959,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X8),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X8),X9)))) )),
  inference(superposition,[],[f576,f659])).
fof(f2991,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),X80))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X79),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X79),X80)))) )),
  inference(superposition,[],[f693,f576])).
fof(f4105,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X77,bool),'_dst_x0_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),X76))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X75),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X75),X76)))) )),
  inference(superposition,[],[f2008,f576])).
fof(f4285,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X73),X74))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X73),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X73),X74)))) )),
  inference(forward_demodulation,[],[f4239,f1846])).
fof(f4239,plain,(
  ( ! [X74,X75,X73] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X73),X74))) != s(bool,'|dollar|exists_1'(s(fun(X75,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X73),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X73),X74)))) )),
  inference(superposition,[],[f1961,f576])).
fof(f4523,plain,(
  ( ! [X180,X179] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X179),X180))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X179),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X179),X180)))) )),
  inference(forward_demodulation,[],[f4404,f1846])).
fof(f4404,plain,(
  ( ! [X180,X181,X179] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X179),X180))) = s(bool,'|dollar|exists_1'(s(fun(X181,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X179),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X179),X180)))) )),
  inference(superposition,[],[f576,f1962])).
fof(f5476,plain,(
  ( ! [X92,X90,X91] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X90),X91))),s(fun('type.num.num',bool),X92),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X90),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X90),X91)))) )),
  inference(superposition,[],[f752,f576])).
fof(f5703,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X131,bool),'const.pair.,_2'(s(X131,X132),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)))) )),
  inference(superposition,[],[f808,f576])).
fof(f5858,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X136),'const.pair.,_2'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X134),X135))),s(X136,X137))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X134),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X134),X135)))) )),
  inference(superposition,[],[f868,f576])).
fof(f6110,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'T_0') = s(bool,happ(s(fun(X141,bool),'_dst_x0_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X139),X140))))),s(X141,X142))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X139),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X139),X140)))) )),
  inference(superposition,[],[f932,f576])).
fof(f6542,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X146,bool),'const.pair.,_2'(s(X146,X147),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X144),X145))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X144),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X144),X145)))) )),
  inference(superposition,[],[f1452,f576])).
fof(f6719,plain,(
  ( ! [X111,X109,X110] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X109),X110))),s(fun('type.num.num',bool),X111),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X109),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X109),X110)))) )),
  inference(superposition,[],[f1453,f576])).
fof(f6911,plain,(
  ( ! [X156,X154,X155,X153] : (s(bool,'F_0') = s(bool,happ(s(fun(X155,bool),'_dst_x0_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X153),X154))))),s(X155,X156))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X153),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X153),X154)))) )),
  inference(superposition,[],[f1454,f576])).
fof(f7171,plain,(
  ( ! [X158,X161,X159,X160] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X160),'const.pair.,_2'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X158),X159))),s(X160,X161))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X158),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X158),X159)))) )),
  inference(superposition,[],[f1455,f576])).
fof(f7541,plain,(
  ( ! [X158,X161,X159,X160] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X160,bool),'const.pair.,_2'(s(X160,X161),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X158),X159))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X158),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X158),X159)))) )),
  inference(superposition,[],[f6664,f576])).
fof(f7902,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X117),X118))),s(fun('type.num.num',bool),X119),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X117),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X117),X118)))) )),
  inference(superposition,[],[f6826,f576])).
fof(f8194,plain,(
  ( ! [X158,X161,X159,X160] : (s(bool,'T_0') != s(bool,happ(s(fun(X160,bool),'_dst_x0_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X158),X159))))),s(X160,X161))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X158),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X158),X159)))) )),
  inference(superposition,[],[f7114,f576])).
fof(f8673,plain,(
  ( ! [X158,X161,X159,X160] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X160),'const.pair.,_2'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X158),X159))),s(X160,X161))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X158),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X158),X159)))) )),
  inference(superposition,[],[f7378,f576])).
fof(f12911,plain,(
  ( ! [X57,X58,X56] : (s(bool,X58) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X56),X57))) | s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X56),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X56),X57)))) )),
  inference(superposition,[],[f4407,f576])).
fof(f12935,plain,(
  ( ! [X54,X56,X55] : (s(bool,X56) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X54),X55))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X54),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X54),X55))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f4407,f576])).
fof(f13679,plain,(
  ( ! [X74,X72,X73] : (s(bool,X74) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f10557,f576])).
fof(f13742,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X59),X60))) | s(bool,X61) = s(bool,happ(s(fun('type.pair.prod'(bool,X59),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X59),X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f10557,f576])).
fof(f14864,plain,(
  ( ! [X80,X81,X82] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X80),X81))) | s(bool,X82) = s(bool,happ(s(fun('type.pair.prod'(bool,X80),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X80),X81))) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f13961,f576])).
fof(f988,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X3)))) | s(X0,X4) = s(X0,sK19(X4,'const.pair.SND_1'(s('type.pair.prod'(X1,fun(X0,bool)),'const.pair.,_2'(s(X1,X2),s(fun(X0,bool),X3)))),X0))) )),
  inference(superposition,[],[f528,f578])).
fof(f989,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6)))) | s(X5,X8) = s(X5,sK19(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))) )),
  inference(superposition,[],[f528,f457])).
fof(f990,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X10)))) | s(X9,X11) = s(X9,sK19(X11,happ(s(fun(fun(X9,bool),fun(X9,bool)),'_i_0_0'),s(fun(X9,bool),X10)),X9))) )),
  inference(superposition,[],[f528,f454])).
fof(f991,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | s(X12,X16) = s(X12,sK19(X16,happ(s(fun(X13,fun(X12,bool)),'_dst_x0_1'(s(fun(X12,bool),X14))),s(X13,X15)),X12))) )),
  inference(superposition,[],[f528,f580])).
fof(f992,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X19)))) | s(X17,X21) = s(X17,sK19(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17))) )),
  inference(superposition,[],[f528,f579])).
fof(f1511,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X3))) | s(X0,X4) = s(X0,sK19(X4,'const.pair.SND_1'(s('type.pair.prod'(X1,fun(X0,bool)),'const.pair.,_2'(s(X1,X2),s(fun(X0,bool),X3)))),X0))) )),
  inference(superposition,[],[f994,f578])).
fof(f1512,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6))) | s(X5,X8) = s(X5,sK19(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))) )),
  inference(superposition,[],[f994,f457])).
fof(f1513,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X9,bool),X10))) | s(X9,X11) = s(X9,sK19(X11,happ(s(fun(fun(X9,bool),fun(X9,bool)),'_i_0_0'),s(fun(X9,bool),X10)),X9))) )),
  inference(superposition,[],[f994,f454])).
fof(f1514,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14))) | s(X12,X16) = s(X12,sK19(X16,happ(s(fun(X13,fun(X12,bool)),'_dst_x0_1'(s(fun(X12,bool),X14))),s(X13,X15)),X12))) )),
  inference(superposition,[],[f994,f580])).
fof(f1515,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X17,bool),X19))) | s(X17,X21) = s(X17,sK19(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17))) )),
  inference(superposition,[],[f994,f579])).
fof(f1516,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(X22,X25) = s(X22,sK19(X25,sK19(X23,X24,fun(X22,bool)),X22)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X22,bool),bool),X24)))) )),
  inference(superposition,[],[f994,f994])).
fof(f1547,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,'|dollar|exists_1'(s(fun(X91,bool),X92)))) | s(X91,X94) = s(X91,sK19(X94,sK19(X92,X93,fun(X91,bool)),X91)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X91,bool),bool),X93)))) )),
  inference(superposition,[],[f528,f994])).
fof(f1585,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X29)) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))) | s(X27,X30) = s(X27,sK19(X30,X28,X27))) )),
  inference(subsumption_resolution,[],[f1561,f424])).
fof(f1561,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))) | s(X27,X30) = s(X27,sK19(X30,X28,X27))) )),
  inference(superposition,[],[f575,f994])).
fof(f1826,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))))))) | s(X24,X27) = s(X24,sK19(X27,X25,X24))) )),
  inference(subsumption_resolution,[],[f1815,f424])).
fof(f1815,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))))))) | s(X24,X27) = s(X24,sK19(X27,X25,X24))) )),
  inference(superposition,[],[f1687,f994])).
fof(f16402,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) )),
  inference(subsumption_resolution,[],[f16382,f426])).
fof(f16382,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) )),
  inference(superposition,[],[f10555,f430])).
fof(f16406,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'F_0') = s(bool,happ(s(fun(X46,bool),X47),s(X46,X48))) | s(bool,X49) = s(bool,'|dollar|exists_1'(s(fun(X46,bool),X47))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f16389,f426])).
fof(f16389,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun(X46,bool),X47),s(X46,X48))) | s(bool,X49) = s(bool,'|dollar|exists_1'(s(fun(X46,bool),X47))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(superposition,[],[f10555,f10557])).
fof(f18945,plain,(
  ( ! [X524,X526,X522,X525,X523] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X522,bool),happ(s(fun('type.num.num',fun(X522,bool)),X524),s('type.num.num',X525))))) | s(X522,X526) = s(X522,sK19(X526,'const.ind_type.FCONS_3'(s(fun(X522,bool),X523),s(fun('type.num.num',fun(X522,bool)),X524),s('type.num.num','const.num.SUC_1'(s('type.num.num',X525)))),X522))) )),
  inference(superposition,[],[f994,f550])).
fof(f18944,plain,(
  ( ! [X520,X518,X521,X517,X519] : (p(s(bool,'|dollar|exists_1'(s(fun(X517,bool),happ(s(fun('type.num.num',fun(X517,bool)),X519),s('type.num.num',X520)))))) | s(X517,X521) = s(X517,sK19(X521,'const.ind_type.FCONS_3'(s(fun(X517,bool),X518),s(fun('type.num.num',fun(X517,bool)),X519),s('type.num.num','const.num.SUC_1'(s('type.num.num',X520)))),X517))) )),
  inference(superposition,[],[f528,f550])).
fof(f3012,plain,(
  ( ! [X152,X151,X149,X150] : (p(s(bool,'|dollar|exists_1'(s(fun(X149,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X149,bool),X150),X151)))))) | s(X149,X152) = s(X149,sK19(X152,happ(s(fun('type.pair.prod'(fun(X149,bool),X150),fun(X149,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X149,bool),X150),X151)),X149))) )),
  inference(superposition,[],[f528,f576])).
fof(f3011,plain,(
  ( ! [X146,X147,X145,X148] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X145,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X145,bool),X146),X147))))) | s(X145,X148) = s(X145,sK19(X148,happ(s(fun('type.pair.prod'(fun(X145,bool),X146),fun(X145,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X145,bool),X146),X147)),X145))) )),
  inference(superposition,[],[f994,f576])).
fof(f1825,plain,(
  ( ! [X23,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,X22)))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(forward_demodulation,[],[f1824,f454])).
fof(f1824,plain,(
  ( ! [X23,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1814,f424])).
fof(f1814,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1687,f693])).
fof(f1984,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun(X50,bool),'_dst_x0_1'(s(bool,X48)))))) | s(bool,'F_0') = s(bool,sK19(X48,X49,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1825,f994])).
fof(f4178,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X26)) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun(X24,bool),'_dst_x0_1'(s(bool,X25))))) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f4126,f424])).
fof(f4126,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun(X24,bool),'_dst_x0_1'(s(bool,X25))))) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f575,f2008])).
fof(f4180,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun(X50,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X48,bool),'_dst_x0_1'(s(bool,X49)))))))))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f4133,f424])).
fof(f4133,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X50,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X48,bool),'_dst_x0_1'(s(bool,X49)))))))))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(superposition,[],[f1687,f2008])).
fof(f1822,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(X48,bool),'_dst_x0_1'(s(bool,sK19(X46,X47,bool))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X47)))) )),
  inference(superposition,[],[f1687,f994])).
fof(f1529,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') != s(bool,sK19(X39,X40,bool)) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X40)))) )),
  inference(superposition,[],[f431,f994])).
fof(f1531,plain,(
  ( ! [X43,X44] : (~p(s(bool,sK19(X43,X44,bool))) | s(bool,'F_0') != s(bool,X43) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X44)))) )),
  inference(superposition,[],[f435,f994])).
fof(f1533,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | p(s(bool,X49)) | s(bool,X49) = s(bool,sK19(X47,X48,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X48)))) )),
  inference(superposition,[],[f572,f994])).
fof(f1535,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X55)) | s(bool,X55) = s(bool,sK19(X53,X54,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) )),
  inference(superposition,[],[f575,f994])).
fof(f1540,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,X68)) | s(bool,X70) != s(bool,sK19(X68,X69,bool)) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X69)))) )),
  inference(superposition,[],[f661,f994])).
fof(f1542,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,X73)) | s(bool,X75) = s(bool,sK19(X73,X74,bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X74)))) )),
  inference(superposition,[],[f715,f994])).
fof(f1543,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,X76)) | s(bool,X78) = s(bool,sK19(X76,X77,bool)) | s(bool,'F_0') = s(bool,X78) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) )),
  inference(superposition,[],[f1457,f994])).
fof(f6357,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,X150)) | s(bool,X152) = s(bool,sK19(X150,X151,bool)) | s(bool,X152) = s(bool,X153) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X151)))) )),
  inference(superposition,[],[f1378,f994])).
fof(f1787,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.bvpSem.state'(X22),X23) != s('type.bvpSem.state'(X22),X25) | s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X22),X25))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X22),sK19(X23,X24,'type.bvpSem.state'(X22))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X22),bool),X24)))) )),
  inference(superposition,[],[f533,f994])).
fof(f1893,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.bvpSem.state'(X22),X23) != s('type.bvpSem.state'(X22),X25) | s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X22),X25))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X22),sK19(X23,X24,'type.bvpSem.state'(X22))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X22),bool),X24)))) )),
  inference(superposition,[],[f535,f994])).
fof(f2024,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.bvpSem.state'(X22),X23) != s('type.bvpSem.state'(X22),X25) | s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X22),X25))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X22),sK19(X23,X24,'type.bvpSem.state'(X22))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X22),bool),X24)))) )),
  inference(superposition,[],[f538,f994])).
fof(f9963,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.bvpSem.state'(X26),X27) != s('type.bvpSem.state'(X26),X29) | s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X26),X29))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X26),sK19(X27,X28,'type.bvpSem.state'(X26))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X26),bool),X28)))) )),
  inference(superposition,[],[f530,f994])).
fof(f11209,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.bvpSem.state'(X26),X27) != s('type.bvpSem.state'(X26),X29) | s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X26),X29))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X26),sK19(X27,X28,'type.bvpSem.state'(X26))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X26),bool),X28)))) )),
  inference(superposition,[],[f531,f994])).
fof(f11488,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.bvpSem.state'(X26),X27) != s('type.bvpSem.state'(X26),X29) | s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X26),X29))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X26),sK19(X27,X28,'type.bvpSem.state'(X26))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X26),bool),X28)))) )),
  inference(superposition,[],[f532,f994])).
fof(f11850,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.bvpSem.state'(X26),X27) != s('type.bvpSem.state'(X26),X29) | s('type.ffi.ffi_state'(X26),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X26),X29))) = s('type.ffi.ffi_state'(X26),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X26),sK19(X27,X28,'type.bvpSem.state'(X26))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.bvpSem.state'(X26),bool),X28)))) )),
  inference(superposition,[],[f537,f994])).
fof(f19452,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X9),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f19281,f430])).
fof(f19450,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X1),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f19281,f10557])).
fof(f19419,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X51),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,X51))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f19281,f430])).
fof(f19417,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X45),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,X45))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))))) | s(bool,X45) = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f19281,f10557])).
fof(f19415,plain,(
  ( ! [X39,X41,X40] : (s(bool,X39) != s(bool,'const.ind_type.FCONS_3'(s(bool,X39),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))))) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f19281,f659])).
fof(f19281,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X2),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(subsumption_resolution,[],[f19174,f426])).
fof(f19174,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X2),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f431,f18667])).
fof(f19280,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f19162,f650])).
fof(f19162,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X6) = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f18667,f659])).
fof(f19121,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,'const.ind_type.FCONS_3'(s(bool,X9),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f18667,f430])).
fof(f19119,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.ind_type.FCONS_3'(s(bool,X1),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f18667,f10557])).
fof(f19088,plain,(
  ( ! [X52,X51] : (s(bool,X51) = s(bool,'const.ind_type.FCONS_3'(s(bool,X51),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,X51))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f18667,f430])).
fof(f19086,plain,(
  ( ! [X47,X45,X46] : (s(bool,X45) = s(bool,'const.ind_type.FCONS_3'(s(bool,X45),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,X45))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))))) | s(bool,X45) = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f18667,f10557])).
fof(f18667,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),'_dst_x0_1'(s(bool,'F_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f8148,f550,f4407])).
fof(f18413,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X76),s(bool,X75))))) | s(bool,'T_0') = s(bool,X76) | s(bool,X75) = s(bool,X77) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f1455,f10557])).
fof(f18094,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,happ(s(fun(bool,bool),'_dst_x0_1'(s(bool,X76))),s(bool,X75))) | s(bool,'T_0') = s(bool,X76) | s(bool,X75) = s(bool,X77) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f1454,f10557])).
fof(f17510,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X75),s(bool,X76))))) | s(bool,'T_0') = s(bool,X76) | s(bool,X75) = s(bool,X77) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f1452,f10557])).
fof(f14866,plain,(
  ( ! [X88,X87,X89] : (s(bool,'T_0') != s(bool,X87) | s(bool,X89) = s(bool,sK19(X87,X88,bool)) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f13961,f994])).
fof(f14865,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,'T_0') != s(bool,X84) | s(bool,X86) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X83),'const.pair.,_2'(s(bool,X84),s(X83,X85))))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f13961,f579])).
fof(f14863,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,X79) = s(bool,happ(s(fun(X76,bool),'_dst_x0_1'(s(bool,X77))),s(X76,X78))) | s(bool,'F_0') = s(bool,X79)) )),
  inference(superposition,[],[f13961,f580])).
fof(f14862,plain,(
  ( ! [X74,X75] : (s(bool,'T_0') != s(bool,X74) | s(bool,X75) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'F_0') = s(bool,X75)) )),
  inference(superposition,[],[f13961,f454])).
fof(f14861,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X71) | s(bool,X73) = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(superposition,[],[f13961,f457])).
fof(f14860,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'T_0') != s(bool,X69) | s(bool,X70) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X67,bool),'const.pair.,_2'(s(X67,X68),s(bool,X69))))) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f13961,f578])).
fof(f14827,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f13961,f659])).
fof(f14495,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f14009,f426])).
fof(f14009,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f572,f10557])).
fof(f14494,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) != s(bool,X81) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f13985,f650])).
fof(f13985,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f4407,f10557])).
fof(f13966,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,X23),s(X24,X25))))) | s(bool,X23) = s(bool,X26) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f8626,f10557])).
fof(f13965,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') != s(bool,happ(s(fun(X20,bool),'_dst_x0_1'(s(bool,X19))),s(X20,X21))) | s(bool,X19) = s(bool,X22) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f8148,f10557])).
fof(f13964,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.|48|_0'))) | s(bool,X16) = s(bool,X18) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f7857,f10557])).
fof(f13963,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X13,bool),'const.pair.,_2'(s(X13,X14),s(bool,X12))))) | s(bool,X12) = s(bool,X15) | s(bool,'F_0') = s(bool,X15)) )),
  inference(superposition,[],[f7421,f10557])).
fof(f13961,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f650,f10557])).
fof(f13960,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X5) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f435,f10557])).
fof(f13744,plain,(
  ( ! [X68,X66,X67] : (s(bool,'F_0') = s(bool,X66) | s(bool,X68) = s(bool,sK19(X66,X67,bool)) | s(bool,'F_0') = s(bool,X68) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X67)))) )),
  inference(superposition,[],[f10557,f994])).
fof(f13712,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f10557,f10557])).
fof(f13681,plain,(
  ( ! [X80,X81,X79] : (s(bool,X79) = s(bool,X81) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,sK19(X79,X80,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X80)))) )),
  inference(superposition,[],[f10557,f994])).
fof(f13680,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X76) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),'const.pair.,_2'(s(bool,X76),s(X75,X77)))))) )),
  inference(superposition,[],[f10557,f579])).
fof(f13678,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,happ(s(fun(X68,bool),'_dst_x0_1'(s(bool,X69))),s(X68,X70)))) )),
  inference(superposition,[],[f10557,f580])).
fof(f13677,plain,(
  ( ! [X66,X67] : (s(bool,X66) = s(bool,X67) | s(bool,'F_0') = s(bool,X67) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X66)))) )),
  inference(superposition,[],[f10557,f454])).
fof(f13676,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X63),s(fun('type.num.num',bool),X64),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f10557,f457])).
fof(f13675,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,'F_0') = s(bool,X62) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X59,bool),'const.pair.,_2'(s(X59,X60),s(bool,X61)))))) )),
  inference(superposition,[],[f10557,f578])).
fof(f13584,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) != s(bool,sK19(X63,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,sK19(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) ) | $spl60),
  inference(superposition,[],[f12985,f994])).
fof(f13583,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X58) != s(bool,sK19(X60,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,X58),s(X57,X59)))))) ) | $spl60),
  inference(superposition,[],[f12985,f579])).
fof(f13581,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X51) != s(bool,sK19(X53,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,happ(s(fun(X50,bool),'_dst_x0_1'(s(bool,X51))),s(X50,X52)))) ) | $spl60),
  inference(superposition,[],[f12985,f580])).
fof(f13580,plain,(
  ( ! [X48,X49] : (s(bool,X48) != s(bool,sK19(X49,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X48)))) ) | $spl60),
  inference(superposition,[],[f12985,f454])).
fof(f13579,plain,(
  ( ! [X47,X45,X46] : (s(bool,X45) != s(bool,sK19(X47,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X45),s(fun('type.num.num',bool),X46),s('type.num.num','const.num.|48|_0')))) ) | $spl60),
  inference(superposition,[],[f12985,f457])).
fof(f13578,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X43) != s(bool,sK19(X44,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43)))))) ) | $spl60),
  inference(superposition,[],[f12985,f578])).
fof(f12937,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) != s(bool,X63) | s(bool,'T_0') = s(bool,sK19(X61,X62,bool)) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f4407,f994])).
fof(f12936,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X58) != s(bool,X60) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,X58),s(X57,X59))))) | s(bool,'F_0') = s(bool,X60)) )),
  inference(superposition,[],[f4407,f579])).
fof(f12934,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X51) != s(bool,X53) | s(bool,'T_0') = s(bool,happ(s(fun(X50,bool),'_dst_x0_1'(s(bool,X51))),s(X50,X52))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f4407,f580])).
fof(f12933,plain,(
  ( ! [X48,X49] : (s(bool,X48) != s(bool,X49) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X48))) | s(bool,'F_0') = s(bool,X49)) )),
  inference(superposition,[],[f4407,f454])).
fof(f12932,plain,(
  ( ! [X47,X45,X46] : (s(bool,X45) != s(bool,X47) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X45),s(fun('type.num.num',bool),X46),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f4407,f457])).
fof(f12931,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X43) != s(bool,X44) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43))))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(superposition,[],[f4407,f578])).
fof(f12913,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) != s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,sK19(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f4407,f994])).
fof(f12912,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,X60) != s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X59),'const.pair.,_2'(s(bool,X60),s(X59,X61)))))) )),
  inference(superposition,[],[f4407,f579])).
fof(f12910,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X53) != s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X53))),s(X52,X54)))) )),
  inference(superposition,[],[f4407,f580])).
fof(f12909,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X50)))) )),
  inference(superposition,[],[f4407,f454])).
fof(f12908,plain,(
  ( ! [X47,X48,X49] : (s(bool,X47) != s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X47),s(fun('type.num.num',bool),X48),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f4407,f457])).
fof(f12907,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,X45)))))) )),
  inference(superposition,[],[f4407,f578])).
fof(f12877,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X1)))))) )),
  inference(unit_resulting_resolution,[],[f576,f8626,f4407])).
fof(f12182,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X34))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f8626,f659])).
fof(f11963,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,happ(s(fun(bool,bool),'_dst_x0_1'(s(bool,'F_0'))),s(bool,X34))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f8148,f659])).
fof(f11581,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X34),s(bool,'F_0'))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f7421,f659])).
fof(f752,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f457,f430])).
fof(f808,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6))))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f578,f430])).
fof(f868,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6))))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f579,f430])).
fof(f932,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f580,f430])).
fof(f1458,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f1433,f424])).
fof(f1433,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | s(bool,X18) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(superposition,[],[f661,f693])).
fof(f1482,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1470,f424])).
fof(f1470,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1457,f693])).
fof(f1510,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16))) | s(bool,'T_0') = s(bool,sK19(X15,X16,bool))) )),
  inference(superposition,[],[f994,f430])).
fof(f1541,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X71))) | s(bool,'F_0') = s(bool,sK19(X71,X72,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X72)))) )),
  inference(superposition,[],[f693,f994])).
fof(f1587,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X56) = s(bool,'|dollar|exists_1'(s(fun(X54,bool),X55))) | s(bool,'F_0') = s(bool,X56) | s(X54,X57) = s(X54,sK19(X57,X55,X54))) )),
  inference(subsumption_resolution,[],[f1567,f424])).
fof(f1567,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'T_0')) | s(bool,X56) = s(bool,'|dollar|exists_1'(s(fun(X54,bool),X55))) | s(bool,'F_0') = s(bool,X56) | s(X54,X57) = s(X54,sK19(X57,X55,X54))) )),
  inference(superposition,[],[f1457,f994])).
fof(f2008,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),'_dst_x0_1'(s(bool,X17))))) | s(bool,'F_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f1991,f426])).
fof(f1991,plain,(
  ( ! [X17,X16] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),'_dst_x0_1'(s(bool,X17)))))) )),
  inference(superposition,[],[f1825,f430])).
fof(f4101,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X65,bool),'_dst_x0_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X62,bool),'const.pair.,_2'(s(X62,X63),s(bool,X64)))))) )),
  inference(superposition,[],[f2008,f578])).
fof(f4102,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X68,bool),'_dst_x0_1'(s(bool,X66))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X66),s(fun('type.num.num',bool),X67),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2008,f457])).
fof(f4104,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X74,bool),'_dst_x0_1'(s(bool,X72))))) | s(bool,'F_0') = s(bool,happ(s(fun(X71,bool),'_dst_x0_1'(s(bool,X72))),s(X71,X73)))) )),
  inference(superposition,[],[f2008,f580])).
fof(f4106,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X81,bool),'_dst_x0_1'(s(bool,X79))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X79),s(X78,X80)))))) )),
  inference(superposition,[],[f2008,f579])).
fof(f4107,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X84,bool),'_dst_x0_1'(s(bool,X82))))) | s(bool,'F_0') = s(bool,sK19(X82,X83,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X83)))) )),
  inference(superposition,[],[f2008,f994])).
fof(f4179,plain,(
  ( ! [X39,X37,X38] : (s(bool,X39) != s(bool,'|dollar|exists_1'(s(fun(X37,bool),'_dst_x0_1'(s(bool,X38))))) | s(bool,'T_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4129,f424])).
fof(f4129,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | s(bool,X39) != s(bool,'|dollar|exists_1'(s(fun(X37,bool),'_dst_x0_1'(s(bool,X38))))) | s(bool,'T_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f661,f2008])).
fof(f4318,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X70) = s(bool,'|dollar|exists_1'(s(fun(X71,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X70)))))) )),
  inference(superposition,[],[f1962,f578])).
fof(f4319,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) = s(bool,'|dollar|exists_1'(s(fun(X74,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1962,f457])).
fof(f4320,plain,(
  ( ! [X76,X75] : (s(bool,X75) = s(bool,'|dollar|exists_1'(s(fun(X76,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X75)))) )),
  inference(superposition,[],[f1962,f454])).
fof(f4321,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,X78) = s(bool,'|dollar|exists_1'(s(fun(X80,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,happ(s(fun(X77,bool),'_dst_x0_1'(s(bool,X78))),s(X77,X79)))) )),
  inference(superposition,[],[f1962,f580])).
fof(f4323,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X85) = s(bool,'|dollar|exists_1'(s(fun(X87,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X84),'const.pair.,_2'(s(bool,X85),s(X84,X86)))))) )),
  inference(superposition,[],[f1962,f579])).
fof(f4324,plain,(
  ( ! [X90,X88,X89] : (s(bool,X88) = s(bool,'|dollar|exists_1'(s(fun(X90,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,sK19(X88,X89,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X89)))) )),
  inference(superposition,[],[f1962,f994])).
fof(f4325,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1962,f659])).
fof(f5472,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X79),s(fun('type.num.num',bool),X80),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X77,bool),'const.pair.,_2'(s(X77,X78),s(bool,X79)))))) )),
  inference(superposition,[],[f752,f578])).
fof(f5473,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X81),s(fun('type.num.num',bool),X83),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X81),s(fun('type.num.num',bool),X82),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f752,f457])).
fof(f5475,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X87),s(fun('type.num.num',bool),X89),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X86,bool),'_dst_x0_1'(s(bool,X87))),s(X86,X88)))) )),
  inference(superposition,[],[f752,f580])).
fof(f5477,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X94),s(fun('type.num.num',bool),X96),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,X94),s(X93,X95)))))) )),
  inference(superposition,[],[f752,f579])).
fof(f5478,plain,(
  ( ! [X99,X97,X98] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X97),s(fun('type.num.num',bool),X99),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK19(X97,X98,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X98)))) )),
  inference(superposition,[],[f752,f994])).
fof(f5570,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) != s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f5508,f424])).
fof(f5508,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X31) != s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X29)) )),
  inference(superposition,[],[f661,f752])).
fof(f5699,plain,(
  ( ! [X116,X114,X112,X115,X113] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X115,bool),'const.pair.,_2'(s(X115,X116),s(bool,X114))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X112,bool),'const.pair.,_2'(s(X112,X113),s(bool,X114)))))) )),
  inference(superposition,[],[f808,f578])).
fof(f5700,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X119,bool),'const.pair.,_2'(s(X119,X120),s(bool,X117))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X117),s(fun('type.num.num',bool),X118),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f808,f457])).
fof(f5702,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,X128),s(bool,X125))))) | s(bool,'F_0') = s(bool,happ(s(fun(X124,bool),'_dst_x0_1'(s(bool,X125))),s(X124,X126)))) )),
  inference(superposition,[],[f808,f580])).
fof(f5704,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X136,bool),'const.pair.,_2'(s(X136,X137),s(bool,X134))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X133),'const.pair.,_2'(s(bool,X134),s(X133,X135)))))) )),
  inference(superposition,[],[f808,f579])).
fof(f5705,plain,(
  ( ! [X140,X138,X141,X139] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X140,bool),'const.pair.,_2'(s(X140,X141),s(bool,X138))))) | s(bool,'F_0') = s(bool,sK19(X138,X139,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X139)))) )),
  inference(superposition,[],[f808,f994])).
fof(f5813,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X44) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f5728,f424])).
fof(f5728,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | s(bool,X44) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f661,f808])).
fof(f5854,plain,(
  ( ! [X121,X118,X120,X119,X117] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X120),'const.pair.,_2'(s(bool,X119),s(X120,X121))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X117,bool),'const.pair.,_2'(s(X117,X118),s(bool,X119)))))) )),
  inference(superposition,[],[f868,f578])).
fof(f5855,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,X122),s(X124,X125))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X122),s(fun('type.num.num',bool),X123),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f868,f457])).
fof(f5857,plain,(
  ( ! [X132,X130,X133,X131,X129] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X130),s(X132,X133))))) | s(bool,'F_0') = s(bool,happ(s(fun(X129,bool),'_dst_x0_1'(s(bool,X130))),s(X129,X131)))) )),
  inference(superposition,[],[f868,f580])).
fof(f5859,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X141),'const.pair.,_2'(s(bool,X139),s(X141,X142))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X138),'const.pair.,_2'(s(bool,X139),s(X138,X140)))))) )),
  inference(superposition,[],[f868,f579])).
fof(f5860,plain,(
  ( ! [X146,X144,X145,X143] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X145),'const.pair.,_2'(s(bool,X143),s(X145,X146))))) | s(bool,'F_0') = s(bool,sK19(X143,X144,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X144)))) )),
  inference(superposition,[],[f868,f994])).
fof(f6039,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X44) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,X42),s(X41,X43))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f5950,f424])).
fof(f5950,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | s(bool,X44) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,X42),s(X41,X43))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X42)) )),
  inference(superposition,[],[f661,f868])).
fof(f6106,plain,(
  ( ! [X125,X123,X126,X124,X122] : (s(bool,'T_0') = s(bool,happ(s(fun(X125,bool),'_dst_x0_1'(s(bool,X124))),s(X125,X126))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X122,bool),'const.pair.,_2'(s(X122,X123),s(bool,X124)))))) )),
  inference(superposition,[],[f932,f578])).
fof(f6107,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'T_0') = s(bool,happ(s(fun(X129,bool),'_dst_x0_1'(s(bool,X127))),s(X129,X130))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X127),s(fun('type.num.num',bool),X128),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f932,f457])).
fof(f6109,plain,(
  ( ! [X134,X138,X136,X135,X137] : (s(bool,'T_0') = s(bool,happ(s(fun(X137,bool),'_dst_x0_1'(s(bool,X135))),s(X137,X138))) | s(bool,'F_0') = s(bool,happ(s(fun(X134,bool),'_dst_x0_1'(s(bool,X135))),s(X134,X136)))) )),
  inference(superposition,[],[f932,f580])).
fof(f6111,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,'T_0') = s(bool,happ(s(fun(X146,bool),'_dst_x0_1'(s(bool,X144))),s(X146,X147))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X143),'const.pair.,_2'(s(bool,X144),s(X143,X145)))))) )),
  inference(superposition,[],[f932,f579])).
fof(f6112,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') = s(bool,happ(s(fun(X150,bool),'_dst_x0_1'(s(bool,X148))),s(X150,X151))) | s(bool,'F_0') = s(bool,sK19(X148,X149,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X149)))) )),
  inference(superposition,[],[f932,f994])).
fof(f6298,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X47) != s(bool,happ(s(fun(X44,bool),'_dst_x0_1'(s(bool,X45))),s(X44,X46))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f6205,f424])).
fof(f6205,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'T_0')) | s(bool,X47) != s(bool,happ(s(fun(X44,bool),'_dst_x0_1'(s(bool,X45))),s(X44,X46))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f661,f932])).
fof(f6383,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,X71) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X70))))) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f6336,f424])).
fof(f6336,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'T_0')) | s(bool,X71) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X70))))) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f1378,f808])).
fof(f6384,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X75) = s(bool,'const.ind_type.FCONS_3'(s(bool,X73),s(fun('type.num.num',bool),X74),s('type.num.num','const.num.|48|_0'))) | s(bool,X75) = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f6337,f424])).
fof(f6337,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | s(bool,X75) = s(bool,'const.ind_type.FCONS_3'(s(bool,X73),s(fun('type.num.num',bool),X74),s('type.num.num','const.num.|48|_0'))) | s(bool,X75) = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X73)) )),
  inference(superposition,[],[f1378,f752])).
fof(f6385,plain,(
  ( ! [X80,X81,X82] : (s(bool,X81) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X80))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f6339,f424])).
fof(f6339,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,'T_0')) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X80))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X80)) )),
  inference(superposition,[],[f1378,f693])).
fof(f6386,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X89) = s(bool,happ(s(fun(X86,bool),'_dst_x0_1'(s(bool,X87))),s(X86,X88))) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f6341,f424])).
fof(f6341,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,'T_0')) | s(bool,X89) = s(bool,happ(s(fun(X86,bool),'_dst_x0_1'(s(bool,X87))),s(X86,X88))) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X87)) )),
  inference(superposition,[],[f1378,f932])).
fof(f6387,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X94) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X91),'const.pair.,_2'(s(bool,X92),s(X91,X93))))) | s(bool,X94) = s(bool,X95) | s(bool,'T_0') = s(bool,X95) | s(bool,'F_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f6342,f424])).
fof(f6342,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,'T_0')) | s(bool,X94) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X91),'const.pair.,_2'(s(bool,X92),s(X91,X93))))) | s(bool,X94) = s(bool,X95) | s(bool,'T_0') = s(bool,X95) | s(bool,'F_0') = s(bool,X92)) )),
  inference(superposition,[],[f1378,f868])).
fof(f6392,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,X110) = s(bool,'|dollar|exists_1'(s(fun(X108,bool),'_dst_x0_1'(s(bool,X109))))) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111) | s(bool,'F_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f6346,f424])).
fof(f6346,plain,(
  ( ! [X111,X109,X110,X108] : (~p(s(bool,'T_0')) | s(bool,X110) = s(bool,'|dollar|exists_1'(s(fun(X108,bool),'_dst_x0_1'(s(bool,X109))))) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111) | s(bool,'F_0') = s(bool,X109)) )),
  inference(superposition,[],[f1378,f2008])).
fof(f6544,plain,(
  ( ! [X156,X154,X155,X153] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),'const.pair.,_2'(s(X155,X156),s(bool,X153))))) | s(bool,'T_0') = s(bool,sK19(X153,X154,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X154)))) )),
  inference(superposition,[],[f1452,f994])).
fof(f6721,plain,(
  ( ! [X118,X116,X117] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X116),s(fun('type.num.num',bool),X118),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,sK19(X116,X117,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X117)))) )),
  inference(superposition,[],[f1453,f994])).
fof(f6913,plain,(
  ( ! [X165,X163,X164,X162] : (s(bool,'F_0') = s(bool,happ(s(fun(X164,bool),'_dst_x0_1'(s(bool,X162))),s(X164,X165))) | s(bool,'T_0') = s(bool,sK19(X162,X163,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X163)))) )),
  inference(superposition,[],[f1454,f994])).
fof(f7173,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X169),'const.pair.,_2'(s(bool,X167),s(X169,X170))))) | s(bool,'T_0') = s(bool,sK19(X167,X168,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X168)))) )),
  inference(superposition,[],[f1455,f994])).
fof(f10529,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X72) = s(bool,happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,X70))),s(X69,X71))) | s(bool,'F_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f6211,f424])).
fof(f6211,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | s(bool,X72) = s(bool,happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,X70))),s(X69,X71))) | s(bool,'F_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f1457,f932])).
fof(f10532,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X61),'const.pair.,_2'(s(bool,X62),s(X61,X63))))) | s(bool,'F_0') = s(bool,X64) | s(bool,'F_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f5955,f424])).
fof(f5955,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X61),'const.pair.,_2'(s(bool,X62),s(X61,X63))))) | s(bool,'F_0') = s(bool,X64) | s(bool,'F_0') = s(bool,X62)) )),
  inference(superposition,[],[f1457,f868])).
fof(f10535,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X59) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X56,bool),'const.pair.,_2'(s(X56,X57),s(bool,X58))))) | s(bool,'F_0') = s(bool,X59) | s(bool,'F_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f5732,f424])).
fof(f5732,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | s(bool,X59) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X56,bool),'const.pair.,_2'(s(X56,X57),s(bool,X58))))) | s(bool,'F_0') = s(bool,X59) | s(bool,'F_0') = s(bool,X58)) )),
  inference(superposition,[],[f1457,f808])).
fof(f10538,plain,(
  ( ! [X39,X37,X38] : (s(bool,X39) = s(bool,'const.ind_type.FCONS_3'(s(bool,X37),s(fun('type.num.num',bool),X38),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f5511,f424])).
fof(f5511,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | s(bool,X39) = s(bool,'const.ind_type.FCONS_3'(s(bool,X37),s(fun('type.num.num',bool),X38),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X37)) )),
  inference(superposition,[],[f1457,f752])).
fof(f10572,plain,(
  ( ! [X47,X45,X46] : (s(bool,X47) = s(bool,'|dollar|exists_1'(s(fun(X45,bool),'_dst_x0_1'(s(bool,X46))))) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4132,f424])).
fof(f4132,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | s(bool,X47) = s(bool,'|dollar|exists_1'(s(fun(X45,bool),'_dst_x0_1'(s(bool,X46))))) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f1457,f2008])).
fof(f4407,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f1961,f1962])).
fof(f8696,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X67),s(bool,X66))))) | s(bool,'T_0') = s(bool,X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f7378,f659])).
fof(f8675,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X169),'const.pair.,_2'(s(bool,X167),s(X169,X170))))) | s(bool,'T_0') = s(bool,sK19(X167,X168,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X168)))) )),
  inference(superposition,[],[f7378,f994])).
fof(f8674,plain,(
  ( ! [X165,X163,X166,X164,X162] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X165),'const.pair.,_2'(s(bool,X163),s(X165,X166))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X162),'const.pair.,_2'(s(bool,X163),s(X162,X164)))))) )),
  inference(superposition,[],[f7378,f579])).
fof(f8672,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X156),'const.pair.,_2'(s(bool,X154),s(X156,X157))))) | s(bool,'T_0') = s(bool,happ(s(fun(X153,bool),'_dst_x0_1'(s(bool,X154))),s(X153,X155)))) )),
  inference(superposition,[],[f7378,f580])).
fof(f8670,plain,(
  ( ! [X146,X149,X147,X148] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X148),'const.pair.,_2'(s(bool,X146),s(X148,X149))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X146),s(fun('type.num.num',bool),X147),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f7378,f457])).
fof(f8669,plain,(
  ( ! [X144,X142,X145,X143,X141] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X144),'const.pair.,_2'(s(bool,X143),s(X144,X145))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))))) )),
  inference(superposition,[],[f7378,f578])).
fof(f8217,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_dst_x0_1'(s(bool,X67))),s(bool,X66))) | s(bool,'T_0') = s(bool,X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f7114,f659])).
fof(f8196,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,'T_0') != s(bool,happ(s(fun(X169,bool),'_dst_x0_1'(s(bool,X167))),s(X169,X170))) | s(bool,'T_0') = s(bool,sK19(X167,X168,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X168)))) )),
  inference(superposition,[],[f7114,f994])).
fof(f8195,plain,(
  ( ! [X165,X163,X166,X164,X162] : (s(bool,'T_0') != s(bool,happ(s(fun(X165,bool),'_dst_x0_1'(s(bool,X163))),s(X165,X166))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X162),'const.pair.,_2'(s(bool,X163),s(X162,X164)))))) )),
  inference(superposition,[],[f7114,f579])).
fof(f8193,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,'T_0') != s(bool,happ(s(fun(X156,bool),'_dst_x0_1'(s(bool,X154))),s(X156,X157))) | s(bool,'T_0') = s(bool,happ(s(fun(X153,bool),'_dst_x0_1'(s(bool,X154))),s(X153,X155)))) )),
  inference(superposition,[],[f7114,f580])).
fof(f8191,plain,(
  ( ! [X146,X149,X147,X148] : (s(bool,'T_0') != s(bool,happ(s(fun(X148,bool),'_dst_x0_1'(s(bool,X146))),s(X148,X149))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X146),s(fun('type.num.num',bool),X147),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f7114,f457])).
fof(f8190,plain,(
  ( ! [X144,X142,X145,X143,X141] : (s(bool,'T_0') != s(bool,happ(s(fun(X144,bool),'_dst_x0_1'(s(bool,X143))),s(X144,X145))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))))) )),
  inference(superposition,[],[f7114,f578])).
fof(f7904,plain,(
  ( ! [X125,X126,X124] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X124),s(fun('type.num.num',bool),X126),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,sK19(X124,X125,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X125)))) )),
  inference(superposition,[],[f6826,f994])).
fof(f7903,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X121),s(fun('type.num.num',bool),X123),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X120),'const.pair.,_2'(s(bool,X121),s(X120,X122)))))) )),
  inference(superposition,[],[f6826,f579])).
fof(f7901,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X114),s(fun('type.num.num',bool),X116),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),'_dst_x0_1'(s(bool,X114))),s(X113,X115)))) )),
  inference(superposition,[],[f6826,f580])).
fof(f7899,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X110),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X109),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6826,f457])).
fof(f7898,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X106),s(fun('type.num.num',bool),X107),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X104,bool),'const.pair.,_2'(s(X104,X105),s(bool,X106)))))) )),
  inference(superposition,[],[f6826,f578])).
fof(f7543,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X169,bool),'const.pair.,_2'(s(X169,X170),s(bool,X167))))) | s(bool,'T_0') = s(bool,sK19(X167,X168,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X168)))) )),
  inference(superposition,[],[f6664,f994])).
fof(f7542,plain,(
  ( ! [X165,X163,X166,X164,X162] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X165,bool),'const.pair.,_2'(s(X165,X166),s(bool,X163))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X162),'const.pair.,_2'(s(bool,X163),s(X162,X164)))))) )),
  inference(superposition,[],[f6664,f579])).
fof(f7540,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X156,bool),'const.pair.,_2'(s(X156,X157),s(bool,X154))))) | s(bool,'T_0') = s(bool,happ(s(fun(X153,bool),'_dst_x0_1'(s(bool,X154))),s(X153,X155)))) )),
  inference(superposition,[],[f6664,f580])).
fof(f7538,plain,(
  ( ! [X146,X149,X147,X148] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X148,bool),'const.pair.,_2'(s(X148,X149),s(bool,X146))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X146),s(fun('type.num.num',bool),X147),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6664,f457])).
fof(f7537,plain,(
  ( ! [X144,X142,X145,X143,X141] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X144,bool),'const.pair.,_2'(s(X144,X145),s(bool,X143))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))))) )),
  inference(superposition,[],[f6664,f578])).
fof(f7444,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X66),s(bool,X67))))) | s(bool,'T_0') = s(bool,X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f6664,f659])).
fof(f6391,plain,(
  ( ! [X107,X105,X106] : (s(bool,X106) = s(bool,'|dollar|exists_1'(s(fun(X105,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X106) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f6345,f424])).
fof(f6345,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,'T_0')) | s(bool,X106) = s(bool,'|dollar|exists_1'(s(fun(X105,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X106) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f1378,f1846])).
fof(f6388,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X98) = s(bool,'|dollar|exists_1'(s(fun(X96,bool),X97))) | s(bool,X98) = s(bool,X99) | s(bool,'T_0') = s(bool,X99) | s(X96,X100) = s(X96,sK19(X100,X97,X96))) )),
  inference(subsumption_resolution,[],[f6343,f424])).
fof(f6343,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,'T_0')) | s(bool,X98) = s(bool,'|dollar|exists_1'(s(fun(X96,bool),X97))) | s(bool,X98) = s(bool,X99) | s(bool,'T_0') = s(bool,X99) | s(X96,X100) = s(X96,sK19(X100,X97,X96))) )),
  inference(superposition,[],[f1378,f994])).
fof(f6133,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) = s(bool,happ(s(fun(bool,bool),'_dst_x0_1'(s(bool,X67))),s(bool,X66))) | s(bool,'F_0') = s(bool,X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f932,f659])).
fof(f5881,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X67),s(bool,X66))))) | s(bool,'F_0') = s(bool,X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f868,f659])).
fof(f5622,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X66),s(bool,X67))))) | s(bool,'F_0') = s(bool,X67) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f808,f659])).
fof(f4242,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1961,f659])).
fof(f4241,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) != s(bool,'|dollar|exists_1'(s(fun(X82,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,sK19(X80,X81,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X81)))) )),
  inference(superposition,[],[f1961,f994])).
fof(f4240,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X77) != s(bool,'|dollar|exists_1'(s(fun(X79,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X76),'const.pair.,_2'(s(bool,X77),s(X76,X78)))))) )),
  inference(superposition,[],[f1961,f579])).
fof(f4238,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X70) != s(bool,'|dollar|exists_1'(s(fun(X72,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,X70))),s(X69,X71)))) )),
  inference(superposition,[],[f1961,f580])).
fof(f4237,plain,(
  ( ! [X68,X67] : (s(bool,X67) != s(bool,'|dollar|exists_1'(s(fun(X68,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X67)))) )),
  inference(superposition,[],[f1961,f454])).
fof(f4236,plain,(
  ( ! [X66,X64,X65] : (s(bool,X64) != s(bool,'|dollar|exists_1'(s(fun(X66,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X64),s(fun('type.num.num',bool),X65),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1961,f457])).
fof(f4235,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X62) != s(bool,'|dollar|exists_1'(s(fun(X63,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X60,bool),'const.pair.,_2'(s(X60,X61),s(bool,X62)))))) )),
  inference(superposition,[],[f1961,f578])).
fof(f2977,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),X44)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f1693,f576])).
fof(f3207,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0'))))) )),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1541,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1693,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f1962,f2004,f2007,f2008,f2298,f2359,f2957,f2958,f2959,f2977,f2991,f3070,f1787,f1893,f2024,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f1984,f1825,f1823,f1687,f3011,f3012,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f2976,f3071])).
fof(f3071,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3061,f424])).
fof(f3061,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2977,f659])).
fof(f2976,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),sK19(X42,'const.pair.FST_0','type.pair.prod'(bool,X41)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X41),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f529,f576])).
fof(f1961,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,'|dollar|exists_1'(s(fun(X20,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1933,f424])).
fof(f1933,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,'|dollar|exists_1'(s(fun(X20,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f661,f1846])).
fof(f1913,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1846,f659])).
fof(f1847,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f1805,f1687])).
fof(f1698,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK19(X3,X2,X1)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1681,f424])).
fof(f1681,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK19(X3,X2,X1)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f529,f659])).
fof(f1673,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK19(X25,sK19(X23,X24,fun(X22,bool)),X22))))) | p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),sK19(X23,X24,fun(X22,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X22,bool),bool),X24)))) )),
  inference(superposition,[],[f529,f994])).
fof(f1692,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X19)))) | ~p(s(bool,happ(s(fun(X17,bool),X19),s(X17,sK19(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17)))))) )),
  inference(forward_demodulation,[],[f1672,f579])).
fof(f1672,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,happ(s(fun(X17,bool),X19),s(X17,sK19(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17))))) | p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20))))))))) )),
  inference(superposition,[],[f529,f579])).
fof(f1691,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK19(X16,happ(s(fun(X13,fun(X12,bool)),'_dst_x0_1'(s(fun(X12,bool),X14))),s(X13,X15)),X12)))))) )),
  inference(forward_demodulation,[],[f1671,f580])).
fof(f1671,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK19(X16,happ(s(fun(X13,fun(X12,bool)),'_dst_x0_1'(s(fun(X12,bool),X14))),s(X13,X15)),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),happ(s(fun(X13,fun(X12,bool)),'_dst_x0_1'(s(fun(X12,bool),X14))),s(X13,X15))))))) )),
  inference(superposition,[],[f529,f580])).
fof(f1690,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK19(X11,happ(s(fun(fun(X9,bool),fun(X9,bool)),'_i_0_0'),s(fun(X9,bool),X10)),X9)))))) )),
  inference(forward_demodulation,[],[f1670,f454])).
fof(f1670,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK19(X11,happ(s(fun(fun(X9,bool),fun(X9,bool)),'_i_0_0'),s(fun(X9,bool),X10)),X9))))) | p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),happ(s(fun(fun(X9,bool),fun(X9,bool)),'_i_0_0'),s(fun(X9,bool),X10))))))) )),
  inference(superposition,[],[f529,f454])).
fof(f1689,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6)))) | ~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK19(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5)))))) )),
  inference(forward_demodulation,[],[f1669,f457])).
fof(f1669,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK19(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))))) | p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f529,f457])).
fof(f1688,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK19(X4,'const.pair.SND_1'(s('type.pair.prod'(X1,fun(X0,bool)),'const.pair.,_2'(s(X1,X2),s(fun(X0,bool),X3)))),X0)))))) )),
  inference(forward_demodulation,[],[f1668,f578])).
fof(f1668,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK19(X4,'const.pair.SND_1'(s('type.pair.prod'(X1,fun(X0,bool)),'const.pair.,_2'(s(X1,X2),s(fun(X0,bool),X3)))),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'const.pair.SND_1'(s('type.pair.prod'(X1,fun(X0,bool)),'const.pair.,_2'(s(X1,X2),s(fun(X0,bool),X3))))))))) )),
  inference(superposition,[],[f529,f578])).
fof(f1586,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))) | s(bool,'T_0') = s(bool,X45) | s(X43,X46) = s(X43,sK19(X46,X44,X43))) )),
  inference(subsumption_resolution,[],[f1564,f424])).
fof(f1564,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | s(bool,X45) != s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))) | s(bool,'T_0') = s(bool,X45) | s(X43,X46) = s(X43,sK19(X46,X44,X43))) )),
  inference(superposition,[],[f661,f994])).
fof(f1508,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,X11) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X10))) | s(bool,'T_0') = s(bool,sK19(X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f994,f659])).
fof(f1506,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,X5) = s(bool,sK19(X3,X4,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f994,f659])).
fof(f1379,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1153,f424])).
fof(f1153,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f661,f659])).
fof(f1054,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f659,f579])).
fof(f1053,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,happ(s(fun(X23,bool),'_dst_x0_1'(s(bool,X24))),s(X23,X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f659,f580])).
fof(f1052,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f659,f454])).
fof(f1051,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X20) = s(bool,'const.ind_type.FCONS_3'(s(bool,X18),s(fun('type.num.num',bool),X19),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f659,f457])).
fof(f1050,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f659,f578])).
fof(f1040,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f659,f659])).
fof(f1026,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X33),'const.pair.,_2'(s(bool,X34),s(X33,X35)))))) )),
  inference(superposition,[],[f659,f579])).
fof(f1025,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),'_dst_x0_1'(s(bool,X30))),s(X29,X31)))) )),
  inference(superposition,[],[f659,f580])).
fof(f1024,plain,(
  ( ! [X28,X27] : (s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27)))) )),
  inference(superposition,[],[f659,f454])).
fof(f1023,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f659,f457])).
fof(f1022,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X20,bool),'const.pair.,_2'(s(X20,X21),s(bool,X22)))))) )),
  inference(superposition,[],[f659,f578])).
fof(f715,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f710,f426])).
fof(f710,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f572,f430])).
fof(f661,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f435,f430])).
fof(f601,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f380])).
fof(f380,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X41,X42,X43] : ! [X44,X45,X46] : s(X41,happ(s(fun(X42,X41),X44),s(X42,happ(s(fun(X43,X42),X45),s(X43,X46))))) = s(X41,happ(s(fun(X43,X41),'_dst_x0x1_2'(s(fun(X42,X41),X44),s(fun(X43,X42),X45))),s(X43,X46)))),
  file('thm.bvpSem.state_|49||49|','_dst_x0x1_ax')).
fof(f591,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.sptree.spt'('type.bvlSem.v'),X2) = s('type.sptree.spt'('type.bvlSem.v'),X11) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : ((s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) | (s('type.sptree.spt'('type.bvlSem.v'),X2) = s('type.sptree.spt'('type.bvlSem.v'),X11) & s('type.list.list'('type.bvpSem.stack'),X3) = s('type.list.list'('type.bvpSem.stack'),X12) & s('type.option.option'('type.num.num'),X4) = s('type.option.option'('type.num.num'),X13) & s('type.num.num',X5) = s('type.num.num',X14) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15) & s('type.num.num',X7) = s('type.num.num',X16) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17) & s('type.ffi.ffi_state'(X1),X9) = s('type.ffi.ffi_state'(X1),X18) & s('type.num.num',X10) = s('type.num.num',X19))) & (s('type.sptree.spt'('type.bvlSem.v'),X2) != s('type.sptree.spt'('type.bvlSem.v'),X11) | s('type.list.list'('type.bvpSem.stack'),X3) != s('type.list.list'('type.bvpSem.stack'),X12) | s('type.option.option'('type.num.num'),X4) != s('type.option.option'('type.num.num'),X13) | s('type.num.num',X5) != s('type.num.num',X14) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15) | s('type.num.num',X7) != s('type.num.num',X16) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17) | s('type.ffi.ffi_state'(X1),X9) != s('type.ffi.ffi_state'(X1),X18) | s('type.num.num',X10) != s('type.num.num',X19) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) = s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))))),
  inference(flattening,[],[f412])).
fof(f412,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : ((s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) | (s('type.sptree.spt'('type.bvlSem.v'),X2) = s('type.sptree.spt'('type.bvlSem.v'),X11) & s('type.list.list'('type.bvpSem.stack'),X3) = s('type.list.list'('type.bvpSem.stack'),X12) & s('type.option.option'('type.num.num'),X4) = s('type.option.option'('type.num.num'),X13) & s('type.num.num',X5) = s('type.num.num',X14) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15) & s('type.num.num',X7) = s('type.num.num',X16) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17) & s('type.ffi.ffi_state'(X1),X9) = s('type.ffi.ffi_state'(X1),X18) & s('type.num.num',X10) = s('type.num.num',X19))) & ((s('type.sptree.spt'('type.bvlSem.v'),X2) != s('type.sptree.spt'('type.bvlSem.v'),X11) | s('type.list.list'('type.bvpSem.stack'),X3) != s('type.list.list'('type.bvpSem.stack'),X12) | s('type.option.option'('type.num.num'),X4) != s('type.option.option'('type.num.num'),X13) | s('type.num.num',X5) != s('type.num.num',X14) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15) | s('type.num.num',X7) != s('type.num.num',X16) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17) | s('type.ffi.ffi_state'(X1),X9) != s('type.ffi.ffi_state'(X1),X18) | s('type.num.num',X10) != s('type.num.num',X19)) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) = s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))))),
  inference(nnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : (s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) = s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) <=> (s('type.sptree.spt'('type.bvlSem.v'),X2) = s('type.sptree.spt'('type.bvlSem.v'),X11) & s('type.list.list'('type.bvpSem.stack'),X3) = s('type.list.list'('type.bvpSem.stack'),X12) & s('type.option.option'('type.num.num'),X4) = s('type.option.option'('type.num.num'),X13) & s('type.num.num',X5) = s('type.num.num',X14) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15) & s('type.num.num',X7) = s('type.num.num',X16) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17) & s('type.ffi.ffi_state'(X1),X9) = s('type.ffi.ffi_state'(X1),X18) & s('type.num.num',X10) = s('type.num.num',X19)))),
  inference(flattening,[],[f378])).
fof(f378,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : (s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) = s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) <=> (s('type.sptree.spt'('type.bvlSem.v'),X2) = s('type.sptree.spt'('type.bvlSem.v'),X11) & s('type.list.list'('type.bvpSem.stack'),X3) = s('type.list.list'('type.bvpSem.stack'),X12) & s('type.option.option'('type.num.num'),X4) = s('type.option.option'('type.num.num'),X13) & s('type.num.num',X5) = s('type.num.num',X14) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15) & s('type.num.num',X7) = s('type.num.num',X16) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17) & s('type.ffi.ffi_state'(X1),X9) = s('type.ffi.ffi_state'(X1),X18) & s('type.num.num',X10) = s('type.num.num',X19)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6,X59] : ! [X60,X61,X62,X63,X64,X65,X66,X58,X57,X67,X68,X69,X70,X71,X72,X73,X74,X75] : (s('type.bvpSem.state'(X59),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X60))),s('type.bvpSem.state'(X59),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X61))),s('type.bvpSem.state'(X59),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X62))),s('type.bvpSem.state'(X59),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X63))),s('type.bvpSem.state'(X59),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X64))),s('type.bvpSem.state'(X59),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X65))),s('type.bvpSem.state'(X59),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X66))),s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),'const.combin.K_1'(s('type.ffi.ffi_state'(X59),X58))),s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X57))),s('type.bvpSem.state'(X6),'const.bool.ARB_0'))))))))))))))))))) = s('type.bvpSem.state'(X59),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X67))),s('type.bvpSem.state'(X59),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X68))),s('type.bvpSem.state'(X59),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X69))),s('type.bvpSem.state'(X59),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X70))),s('type.bvpSem.state'(X59),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X71))),s('type.bvpSem.state'(X59),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X72))),s('type.bvpSem.state'(X59),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X73))),s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),'const.combin.K_1'(s('type.ffi.ffi_state'(X59),X74))),s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X75))),s('type.bvpSem.state'(X6),'const.bool.ARB_0'))))))))))))))))))) <=> (s('type.sptree.spt'('type.bvlSem.v'),X60) = s('type.sptree.spt'('type.bvlSem.v'),X67) & s('type.list.list'('type.bvpSem.stack'),X61) = s('type.list.list'('type.bvpSem.stack'),X68) & s('type.option.option'('type.num.num'),X62) = s('type.option.option'('type.num.num'),X69) & s('type.num.num',X63) = s('type.num.num',X70) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X64) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X71) & s('type.num.num',X65) = s('type.num.num',X72) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X66) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X73) & s('type.ffi.ffi_state'(X59),X58) = s('type.ffi.ffi_state'(X59),X74) & s('type.num.num',X57) = s('type.num.num',X75)))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_literal_|49||49|')).
fof(f592,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.list.list'('type.bvpSem.stack'),X3) = s('type.list.list'('type.bvpSem.stack'),X12) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f593,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.option.option'('type.num.num'),X4) = s('type.option.option'('type.num.num'),X13) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f594,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.num.num',X5) = s('type.num.num',X14) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f595,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f596,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.num.num',X7) = s('type.num.num',X16) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f597,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f598,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.ffi.ffi_state'(X1),X9) = s('type.ffi.ffi_state'(X1),X18) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f599,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.num.num',X10) = s('type.num.num',X19) | s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) != s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0')))))))))))))))))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f600,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X19,X17,X7,X5,X3,X1,X15,X13,X11,X9,X18,X16] : (s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X3))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X4))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X5))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X7))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X9))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) = s('type.bvpSem.state'(X1),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.K_1'(s('type.sptree.spt'('type.bvlSem.v'),X11))),s('type.bvpSem.state'(X1),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.K_1'(s('type.list.list'('type.bvpSem.stack'),X12))),s('type.bvpSem.state'(X1),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.K_1'(s('type.option.option'('type.num.num'),X13))),s('type.bvpSem.state'(X1),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.bvpSem.state'(X1),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15))),s('type.bvpSem.state'(X1),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X16))),s('type.bvpSem.state'(X1),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.K_1'(s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17))),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.K_1'(s('type.ffi.ffi_state'(X1),X18))),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X19))),s('type.bvpSem.state'(X0),'const.bool.ARB_0'))))))))))))))))))) | s('type.num.num',X10) != s('type.num.num',X19) | s('type.ffi.ffi_state'(X1),X9) != s('type.ffi.ffi_state'(X1),X18) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X17) | s('type.num.num',X7) != s('type.num.num',X16) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X15) | s('type.num.num',X5) != s('type.num.num',X14) | s('type.option.option'('type.num.num'),X4) != s('type.option.option'('type.num.num'),X13) | s('type.list.list'('type.bvpSem.stack'),X3) != s('type.list.list'('type.bvpSem.stack'),X12) | s('type.sptree.spt'('type.bvlSem.v'),X2) != s('type.sptree.spt'('type.bvlSem.v'),X11)) )),
  inference(cnf_transformation,[],[f413])).
fof(f590,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.ffi.ffi_state'(X1),X10))),s('type.num.num',X11))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X1),X10),s('type.num.num',X11)))))) )),
  inference(cnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.ffi.ffi_state'(X1),X10))),s('type.num.num',X11))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X1),X10),s('type.num.num',X11)))))),
  inference(flattening,[],[f376])).
fof(f376,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.ffi.ffi_state'(X1),X10))),s('type.num.num',X11))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X1),X10),s('type.num.num',X11)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X59,X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X59),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X59),happ(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X58),s('type.ffi.ffi_state'(X6),X56))),s('type.num.num',X57))) = s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_ffi_fupd')).
fof(f589,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK21(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK21(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK21(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK21(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21])],[f386])).
fof(f386,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f374])).
fof(f374,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.bvpSem.state_|49||49|',aHL_EXT)).
fof(f588,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X1),X2))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0,X1,X2,X3] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X1),X2))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f372])).
fof(f372,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X1),X2))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X59,X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_95_prove_ltb')).
fof(f587,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X1),X2))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0,X1,X2,X3] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X1),X2))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f370])).
fof(f370,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X1),X2))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X59,X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_79_prove_ltb')).
fof(f586,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X1),X2))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0,X1,X2,X3] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X1),X2))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f368])).
fof(f368,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X1),X2))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X59,X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_63_prove_ltb')).
fof(f585,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X2))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X2))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f366])).
fof(f366,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X2))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X59,X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_55_prove_ltb')).
fof(f584,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X1),X2))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0,X1,X2,X3] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X1),X2))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f364])).
fof(f364,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X1),X2))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X59,X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_47_prove_ltb')).
fof(f583,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f362])).
fof(f362,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X59,X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_7|49|_prove_ltb')).
fof(f582,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f360])).
fof(f360,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X59,X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_87_prove_ltb')).
fof(f581,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X1),X2))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X59,X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X59),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X6),'type.ffi.ffi_state'(X59)),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49|2_prove_ltb')).
fof(f577,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK20(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK20(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK20(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK20(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f351])).
fof(f351,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f350])).
fof(f350,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X26,X25] : ! [X37] : ? [X30] : ! [X38,X39,X40] : s(X26,happ(s(fun('type.ind_type.recspace'(X25),X26),X30),s('type.ind_type.recspace'(X25),'const.ind_type.CONSTR_3'(s('type.num.num',X38),s(X25,X39),s(fun('type.num.num','type.ind_type.recspace'(X25)),X40))))) = s(X26,happ(s(fun(fun('type.num.num',X26),X26),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X25)),fun(fun('type.num.num',X26),X26)),happ(s(fun(X25,fun(fun('type.num.num','type.ind_type.recspace'(X25)),fun(fun('type.num.num',X26),X26))),happ(s(fun('type.num.num',fun(X25,fun(fun('type.num.num','type.ind_type.recspace'(X25)),fun(fun('type.num.num',X26),X26)))),X37),s('type.num.num',X38))),s(X25,X39))),s(fun('type.num.num','type.ind_type.recspace'(X25)),X40))),s(fun('type.num.num',X26),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X25),X26),X30),s(fun('type.num.num','type.ind_type.recspace'(X25)),X40)))))),
  file('thm.bvpSem.state_|49||49|','thm.ind_type.CONSTR_REC')).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f348])).
fof(f348,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X25,X26] : ! [X32] : s(X25,'const.pair.FST_1'(s('type.pair.prod'(X25,X26),X32))) = s(X25,happ(s(fun('type.pair.prod'(X25,X26),X25),'const.pair.FST_0'),s('type.pair.prod'(X25,X26),X32)))),
  file('thm.bvpSem.state_|49||49|','ARITY_0_1_const.pair.FST')).
fof(f631,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f630,f571])).
fof(f630,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f629,f571])).
fof(f629,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f628,f571])).
fof(f628,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f627,f571])).
fof(f627,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f626,f571])).
fof(f626,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f625,f571])).
fof(f625,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f624,f571])).
fof(f624,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(backward_demodulation,[],[f571,f570])).
fof(f571,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X32,X33] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X32))),s('type.num.num',X33)))),
  file('thm.bvpSem.state_|49||49|','ARITY_1_2_const.arithmetic.+')).
fof(f570,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))))),s('type.num.num',X10)))))))))))))))))))),
  inference(flattening,[],[f344])).
fof(f344,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))))),s('type.num.num',X10)))))))))))))))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X30,X7,X8,X9,X10,X11,X12,X13,X14,X15] : s('type.num.num','const.bvpSem.state_size_2'(s(fun(X6,'type.num.num'),X30),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X7),s('type.list.list'('type.bvpSem.stack'),X8),s('type.option.option'('type.num.num'),X9),s('type.num.num',X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X11),s('type.num.num',X12),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X13),s('type.ffi.ffi_state'(X6),X14),s('type.num.num',X15))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X7))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X8))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X9))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X10))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X11))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X12))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X13))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X6,'type.num.num'),X30),s('type.ffi.ffi_state'(X6),X14))))),s('type.num.num',X15)))))))))))))))))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_size_def')).
fof(f569,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))) )),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57))))) = s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X58),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55))),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_code_fupd')).
fof(f568,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))) )),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))),
  inference(flattening,[],[f340])).
fof(f340,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57))))) = s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X58),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30))),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_refs_fupd')).
fof(f567,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f338])).
fof(f338,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X58),s('type.list.list'('type.bvpSem.stack'),X52))),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57))) = s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_stack_fupd')).
fof(f566,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f336])).
fof(f336,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X58),s('type.sptree.spt'('type.bvlSem.v'),X51))),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57))) = s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_locals_fupd')).
fof(f565,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f334])).
fof(f334,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X58),s('type.option.option'('type.num.num'),X53))),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57))) = s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_global_fupd')).
fof(f564,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f332])).
fof(f332,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X58),s('type.num.num',X57))))) = s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_space_fupd')).
fof(f563,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X58),s('type.num.num',X54))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57))) = s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_clock_fupd')).
fof(f562,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X58,X51,X52,X53,X31,X30,X54,X55,X56,X57] : s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X58),s('type.num.num',X31))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57))) = s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X58),s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X51),s('type.list.list'('type.bvpSem.stack'),X52),s('type.option.option'('type.num.num'),X53),s('type.num.num',X31),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X30),s('type.num.num',X54),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X55),s('type.ffi.ffi_state'(X6),X56),s('type.num.num',X57)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_handler_fupd')).
fof(f552,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X25] : ! [X32,X33,X34] : s(X25,'const.ind_type.FCONS_3'(s(X25,X32),s(fun('type.num.num',X25),X33),s('type.num.num',X34))) = s(X25,happ(s(fun('type.num.num',X25),happ(s(fun(fun('type.num.num',X25),fun('type.num.num',X25)),happ(s(fun(X25,fun(fun('type.num.num',X25),fun('type.num.num',X25))),'const.ind_type.FCONS_0'),s(X25,X32))),s(fun('type.num.num',X25),X33))),s('type.num.num',X34)))),
  file('thm.bvpSem.state_|49||49|','ARITY_0_3_const.ind_type.FCONS')).
fof(f551,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f306])).
fof(f306,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X25] : ! [X32,X33,X34] : s('type.ind_type.recspace'(X25),'const.ind_type.CONSTR_3'(s('type.num.num',X32),s(X25,X33),s(fun('type.num.num','type.ind_type.recspace'(X25)),X34))) = s('type.ind_type.recspace'(X25),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X25)),'type.ind_type.recspace'(X25)),happ(s(fun(X25,fun(fun('type.num.num','type.ind_type.recspace'(X25)),'type.ind_type.recspace'(X25))),'const.ind_type.CONSTR_1'(s('type.num.num',X32))),s(X25,X33))),s(fun('type.num.num','type.ind_type.recspace'(X25)),X34)))),
  file('thm.bvpSem.state_|49||49|','ARITY_1_3_const.ind_type.CONSTR')).
fof(f550,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X25] : ! [X29,X30,X31] : s(X25,happ(s(fun('type.num.num',X25),X30),s('type.num.num',X31))) = s(X25,'const.ind_type.FCONS_3'(s(X25,X29),s(fun('type.num.num',X25),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))),
  file('thm.bvpSem.state_|49||49|','thm.ind_type.FCONS|split|1')).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),X1) = s('type.bvpSem.state'(X0),X2) | s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) != s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X2))) | s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X2))) | s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) != s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))) | s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) != s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))) | s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) != s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))) | s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) != s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51))))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49|9_prove_ltb')).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51))))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49|7_prove_ltb')).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49|4_prove_ltb')).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49|5_prove_ltb')).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49|3_prove_ltb')).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_78_prove_ltb')).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_93_prove_ltb')).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_9|49|_prove_ltb')).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_89_prove_ltb')).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_9|48|_prove_ltb')).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_75_prove_ltb')).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_74_prove_ltb')).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_73_prove_ltb')).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X30),s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49|8_prove_ltb')).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_62_prove_ltb')).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_6|48|_prove_ltb')).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_54_prove_ltb')).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X6] : ! [X51,X30] : s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X51))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48|3_prove_ltb')).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X30),s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X51))))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49|6_prove_ltb')).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_52_prove_ltb')).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_46_prove_ltb')).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X6] : ! [X51,X30] : s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X51))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48||49|_prove_ltb')).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_44_prove_ltb')).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_85_prove_ltb')).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_8|48|_prove_ltb')).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_86_prove_ltb')).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_94_prove_ltb')).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_92_prove_ltb')).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49||49|_prove_ltb')).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_96_prove_ltb')).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48|9_prove_ltb')).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_76_prove_ltb')).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X6] : ! [X51,X30] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),X51))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_77_prove_ltb')).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_68_prove_ltb')).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_7|48|_prove_ltb')).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X6] : ! [X51,X30] : s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X51))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_97_prove_ltb')).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X6] : ! [X51,X30] : s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X51))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_98_prove_ltb')).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X6] : ! [X51,X30] : s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X51))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_99_prove_ltb')).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f114])).
fof(f114,axiom,(
  ! [X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_58_prove_ltb')).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_57_prove_ltb')).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f127])).
fof(f127,axiom,(
  ! [X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_49_prove_ltb')).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_5|48|_prove_ltb')).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_42_prove_ltb')).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_4|49|_prove_ltb')).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48|6_prove_ltb')).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f206])).
fof(f206,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_66_prove_ltb')).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_43_prove_ltb')).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_45_prove_ltb')).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X6] : ! [X51,X30] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),X51))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_48_prove_ltb')).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_5|49|_prove_ltb')).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_53_prove_ltb')).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_65_prove_ltb')).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X6] : ! [X51,X30] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),X51))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_56_prove_ltb')).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f112])).
fof(f112,axiom,(
  ! [X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_64_prove_ltb')).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_59_prove_ltb')).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X6] : ! [X51,X30] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),X51))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_6|49|_prove_ltb')).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_8|49|_prove_ltb')).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_82_prove_ltb')).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_83_prove_ltb')).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48|5_prove_ltb')).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48|7_prove_ltb')).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X6] : ! [X51,X30] : s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X51))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48||48|_prove_ltb')).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X6] : ! [X51,X30] : s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X51))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48|4_prove_ltb')).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : ! [X51,X30] : s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),X51))) = s('type.ffi.ffi_state'(X6),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48|2_prove_ltb')).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||48|8_prove_ltb')).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_|49||49||48|_prove_ltb')).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_88_prove_ltb')).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_84_prove_ltb')).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_72_prove_ltb')).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X6] : ! [X51,X30] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),X51))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X6),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X30),s('type.bvpSem.state'(X6),X51)))))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.bvpSem_69_prove_ltb')).
fof(f429,plain,(
  ( ! [X0] : (s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X0))) = s('type.num.num',happ(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X0)))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X0))) = s('type.num.num',happ(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X0)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X32] : s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X32))) = s('type.num.num',happ(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X32)))),
  file('thm.bvpSem.state_|49||49|','ARITY_0_1_const.bvpSem.stack_size')).
fof(f428,plain,(
  ( ! [X0] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X0))) = s('type.num.num',happ(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X0)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X0))) = s('type.num.num',happ(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X0)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X32] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X32))) = s('type.num.num',happ(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X32)))),
  file('thm.bvpSem.state_|49||49|','ARITY_0_1_const.bvlSem.v_size')).
fof(f427,plain,(
  ( ! [X0] : (s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X0))) = s('type.num.num',happ(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X0)))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X0))) = s('type.num.num',happ(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X32] : s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X32))) = s('type.num.num',happ(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X32)))),
  file('thm.bvpSem.state_|49||49|','ARITY_0_1_const.bvp.prog_size')).
fof(f39313,plain,(
  ( ! [X358,X359] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X359,bool),'_dst_x0_1'(s(bool,'F_0'))))) | s(bool,'T_0') != s(bool,X358)) ) | ($spl24 | $spl70 | $spl92)),
  inference(forward_demodulation,[],[f38883,f12999])).
fof(f38883,plain,(
  ( ! [X358,X357,X359] : (s(bool,'T_0') != s(bool,X358) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X359,bool),'_dst_x0_1'(s(bool,sK19(X357,'_i_0_0',bool))))))) ) | ($spl24 | $spl70 | $spl92)),
  inference(superposition,[],[f34723,f38277])).
fof(f34723,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X3)))))) ) | ($spl24 | $spl70)),
  inference(subsumption_resolution,[],[f34722,f426])).
fof(f34722,plain,(
  ( ! [X5,X3] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X3)))))) ) | ($spl24 | $spl70)),
  inference(forward_demodulation,[],[f34247,f12999])).
fof(f34247,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,sK19(X4,'_i_0_0',bool))) | s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl70),
  inference(superposition,[],[f431,f22005])).
fof(f39312,plain,(
  $false | ($spl22 | $spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39311,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39311,plain,(
  ( ! [X212] : (s(bool,'T_0') != s(bool,X212)) ) | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39310,f8626])).
fof(f39310,plain,(
  ( ! [X212,X213,X214] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X213),'const.pair.,_2'(s(bool,'F_0'),s(X213,X214))))) | s(bool,'T_0') != s(bool,X212)) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38841,f12999])).
fof(f38841,plain,(
  ( ! [X212,X213,X211,X214] : (s(bool,'T_0') != s(bool,X212) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X213),'const.pair.,_2'(s(bool,sK19(X211,'_i_0_0',bool)),s(X213,X214)))))) ) | ($spl22 | $spl92)),
  inference(superposition,[],[f12462,f38277])).
fof(f12462,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X34) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X33),'const.pair.,_2'(s(bool,X34),s(X33,X35)))))) ) | $spl22),
  inference(forward_demodulation,[],[f12418,f10810])).
fof(f12418,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X33),'const.pair.,_2'(s(bool,X34),s(X33,X35)))))) ) | $spl22),
  inference(superposition,[],[f11033,f579])).
fof(f11033,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(subsumption_resolution,[],[f11004,f424])).
fof(f11004,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(superposition,[],[f661,f10810])).
fof(f39309,plain,(
  $false | ($spl22 | $spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39308,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39308,plain,(
  ( ! [X208] : (s(bool,'T_0') != s(bool,X208)) ) | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39307,f8148])).
fof(f39307,plain,(
  ( ! [X210,X208,X209] : (s(bool,'T_0') = s(bool,happ(s(fun(X209,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X209,X210))) | s(bool,'T_0') != s(bool,X208)) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38840,f12999])).
fof(f38840,plain,(
  ( ! [X210,X208,X209,X207] : (s(bool,'T_0') != s(bool,X208) | s(bool,'T_0') = s(bool,happ(s(fun(X209,bool),'_dst_x0_1'(s(bool,sK19(X207,'_i_0_0',bool)))),s(X209,X210)))) ) | ($spl22 | $spl92)),
  inference(superposition,[],[f12460,f38277])).
fof(f12460,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),'_dst_x0_1'(s(bool,X29))),s(X28,X30)))) ) | $spl22),
  inference(forward_demodulation,[],[f12416,f10810])).
fof(f12416,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),'_dst_x0_1'(s(bool,X29))),s(X28,X30)))) ) | $spl22),
  inference(superposition,[],[f11033,f580])).
fof(f39306,plain,(
  $false | ($spl22 | $spl24 | $spl36 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39305,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39305,plain,(
  ( ! [X206] : (s(bool,'T_0') != s(bool,X206)) ) | ($spl22 | $spl24 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f39304,f5273])).
fof(f39304,plain,(
  ( ! [X206] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') != s(bool,X206)) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38839,f12999])).
fof(f38839,plain,(
  ( ! [X206,X205] : (s(bool,'T_0') != s(bool,X206) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK19(X205,'_i_0_0',bool))))) ) | ($spl22 | $spl92)),
  inference(superposition,[],[f12459,f38277])).
fof(f12459,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27)))) ) | $spl22),
  inference(forward_demodulation,[],[f12415,f10810])).
fof(f12415,plain,(
  ( ! [X27] : (s(bool,X27) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27)))) ) | $spl22),
  inference(superposition,[],[f11033,f454])).
fof(f39303,plain,(
  $false | ($spl22 | $spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39302,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39302,plain,(
  ( ! [X203] : (s(bool,'T_0') != s(bool,X203)) ) | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39301,f7857])).
fof(f39301,plain,(
  ( ! [X204,X203] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X204),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,X203)) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38838,f12999])).
fof(f38838,plain,(
  ( ! [X204,X202,X203] : (s(bool,'T_0') != s(bool,X203) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,sK19(X202,'_i_0_0',bool)),s(fun('type.num.num',bool),X204),s('type.num.num','const.num.|48|_0')))) ) | ($spl22 | $spl92)),
  inference(superposition,[],[f12458,f38277])).
fof(f12458,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X25),s(fun('type.num.num',bool),X26),s('type.num.num','const.num.|48|_0')))) ) | $spl22),
  inference(forward_demodulation,[],[f12414,f10810])).
fof(f12414,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X25),s(fun('type.num.num',bool),X26),s('type.num.num','const.num.|48|_0')))) ) | $spl22),
  inference(superposition,[],[f11033,f457])).
fof(f39300,plain,(
  $false | ($spl22 | $spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39299,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39299,plain,(
  ( ! [X199] : (s(bool,'T_0') != s(bool,X199)) ) | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39298,f7421])).
fof(f39298,plain,(
  ( ! [X200,X199,X201] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X200,bool),'const.pair.,_2'(s(X200,X201),s(bool,'F_0'))))) | s(bool,'T_0') != s(bool,X199)) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38837,f12999])).
fof(f38837,plain,(
  ( ! [X198,X200,X199,X201] : (s(bool,'T_0') != s(bool,X199) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X200,bool),'const.pair.,_2'(s(X200,X201),s(bool,sK19(X198,'_i_0_0',bool))))))) ) | ($spl22 | $spl92)),
  inference(superposition,[],[f12457,f38277])).
fof(f12457,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X24) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,X23),s(bool,X24)))))) ) | $spl22),
  inference(forward_demodulation,[],[f12413,f10810])).
fof(f12413,plain,(
  ( ! [X24,X23,X22] : (s(bool,X24) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,X23),s(bool,X24)))))) ) | $spl22),
  inference(superposition,[],[f11033,f578])).
fof(f39297,plain,(
  $false | ($spl24 | $spl36 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f38945,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f38945,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(condensation,[],[f38785])).
fof(f38785,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,'|dollar|exists_1'(s(fun(X42,bool),'_dst_x0_1'(s(bool,X41)))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(superposition,[],[f22013,f38277])).
fof(f22013,plain,(
  ( ! [X233,X234,X232] : (s(bool,X232) = s(bool,'|dollar|exists_1'(s(fun(X233,bool),'_dst_x0_1'(s(bool,X232))))) | s(bool,X232) = s(bool,sK19(X234,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f21061,f650])).
fof(f21061,plain,(
  ( ! [X233,X234,X232] : (s(bool,X232) = s(bool,'|dollar|exists_1'(s(fun(X233,bool),'_dst_x0_1'(s(bool,X232))))) | s(bool,X232) = s(bool,sK19(X234,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f1846,f20151])).
fof(f39296,plain,(
  $false | ($spl60 | $spl70 | $spl92)),
  inference(global_subsumption,[],[f38783,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f38783,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X34,bool),'_dst_x0_1'(s(bool,X33)))))) ) | ($spl70 | $spl92)),
  inference(superposition,[],[f22005,f38277])).
fof(f39295,plain,(
  $false | ($spl22 | $spl24 | $spl36 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f38946,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f38946,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl22 | $spl24 | $spl36 | $spl92)),
  inference(condensation,[],[f38782])).
fof(f38782,plain,(
  ( ! [X30,X29] : (s(bool,X29) = s(bool,X30) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl22 | $spl24 | $spl36 | $spl92)),
  inference(superposition,[],[f21741,f38277])).
fof(f39294,plain,(
  $false | ($spl24 | $spl30 | $spl36 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f38947,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f38947,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X1),bool),'const.pair.FST_0')))) ) | ($spl24 | $spl30 | $spl36 | $spl92)),
  inference(condensation,[],[f38781])).
fof(f38781,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0')))) ) | ($spl24 | $spl30 | $spl36 | $spl92)),
  inference(superposition,[],[f21740,f38277])).
fof(f39293,plain,(
  $false | ($spl24 | $spl36 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f38948,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f38948,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(condensation,[],[f38780])).
fof(f38780,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) = s(bool,X22) | s(bool,X22) = s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(superposition,[],[f21738,f38277])).
fof(f39292,plain,(
  $false | ($spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f38777,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f38777,plain,(
  ( ! [X12,X11] : (s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12)))) ) | ($spl24 | $spl92)),
  inference(superposition,[],[f13336,f38277])).
fof(f39291,plain,(
  $false | ($spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f38951,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f38951,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl92)),
  inference(condensation,[],[f38776])).
fof(f38776,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl24 | $spl92)),
  inference(superposition,[],[f13017,f38277])).
fof(f39290,plain,(
  $false | ($spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f38775,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f38775,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,X6)) ) | ($spl24 | $spl92)),
  inference(superposition,[],[f13001,f38277])).
fof(f13001,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK19(X0,'_i_0_0',bool))) ) | $spl24),
  inference(unit_resulting_resolution,[],[f1699,f431])).
fof(f39289,plain,(
  $false | ($spl22 | $spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39288,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39288,plain,(
  ( ! [X563] : (s(bool,'T_0') = s(bool,X563)) ) | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f38721,f13017])).
fof(f38721,plain,(
  ( ! [X563,X562] : (s(bool,X563) != s(bool,sK19(X562,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X563)) ) | ($spl22 | $spl92)),
  inference(superposition,[],[f11033,f38277])).
fof(f39287,plain,(
  $false | ($spl24 | $spl30 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39286,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39286,plain,(
  ( ! [X554] : (s(bool,'T_0') = s(bool,X554)) ) | ($spl24 | $spl30 | $spl92)),
  inference(subsumption_resolution,[],[f38717,f13017])).
fof(f38717,plain,(
  ( ! [X553,X554] : (s(bool,X554) != s(bool,sK19(X553,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X554)) ) | ($spl30 | $spl92)),
  inference(superposition,[],[f3325,f38277])).
fof(f3325,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X20),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X21)) ) | $spl30),
  inference(subsumption_resolution,[],[f3296,f424])).
fof(f3296,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X20),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X21)) ) | $spl30),
  inference(superposition,[],[f661,f3245])).
fof(f39285,plain,(
  $false | ($spl24 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39284,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39284,plain,(
  ( ! [X521] : (s(bool,'T_0') = s(bool,X521)) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f38707,f13017])).
fof(f38707,plain,(
  ( ! [X520,X521] : (s(bool,X521) != s(bool,sK19(X520,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X521)) ) | $spl92),
  inference(superposition,[],[f1961,f38277])).
fof(f39282,plain,(
  $false | ($spl22 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39281,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39281,plain,(
  ( ! [X134] : (s(bool,'T_0') = s(bool,X134)) ) | ($spl22 | $spl92)),
  inference(subsumption_resolution,[],[f38610,f38277])).
fof(f38610,plain,(
  ( ! [X134,X135] : (s(bool,sK19(X135,'_i_0_0',bool)) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X134)) ) | ($spl22 | $spl92)),
  inference(superposition,[],[f11033,f38277])).
fof(f39275,plain,(
  $false | ($spl30 | $spl60 | $spl92)),
  inference(global_subsumption,[],[f39274,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39274,plain,(
  ( ! [X91] : (s(bool,'T_0') = s(bool,X91)) ) | ($spl30 | $spl92)),
  inference(subsumption_resolution,[],[f38596,f38277])).
fof(f38596,plain,(
  ( ! [X92,X93,X91] : (s(bool,sK19(X92,'_i_0_0',bool)) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X91)) ) | ($spl30 | $spl92)),
  inference(superposition,[],[f3325,f38277])).
fof(f39273,plain,(
  $false | ($spl60 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f12985,f13578,f13579,f13580,f13581,f13583,f13584,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f32044,f37735,f37737,f37789,f38277,f38532,f39272])).
fof(f39271,plain,(
  $false | ($spl24 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525,f38532,f39262,f39265,f39268,f39270])).
fof(f39270,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f38980,f38426])).
fof(f38980,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,sK19(X1,'_i_0_0',bool)))) ) | $spl92),
  inference(condensation,[],[f38580])).
fof(f38580,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,sK19(X37,'_i_0_0',bool))) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl92),
  inference(superposition,[],[f715,f38277])).
fof(f39268,plain,(
  ( ! [X15,X13] : (s(bool,X13) = s(bool,X15)) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39267,f38426])).
fof(f39267,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,sK19(X14,'_i_0_0',bool))) | s(bool,X13) = s(bool,X15)) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f38572,f38426])).
fof(f38572,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,sK19(X14,'_i_0_0',bool))) | p(s(bool,X15)) | s(bool,X13) = s(bool,X15)) ) | $spl92),
  inference(superposition,[],[f572,f38277])).
fof(f39265,plain,(
  ( ! [X12,X10] : (s(bool,X10) = s(bool,X12)) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39264,f38426])).
fof(f39264,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,sK19(X11,'_i_0_0',bool))) | s(bool,X10) = s(bool,X12)) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f38571,f38426])).
fof(f38571,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,sK19(X11,'_i_0_0',bool))) | p(s(bool,X12)) | s(bool,X10) = s(bool,X12)) ) | $spl92),
  inference(superposition,[],[f572,f38277])).
fof(f39262,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f38567,f38426])).
fof(f38567,plain,(
  ( ! [X2,X3] : (p(s(bool,sK19(X3,'_i_0_0',bool))) | s(bool,'T_0') != s(bool,X2)) ) | $spl92),
  inference(superposition,[],[f431,f38277])).
fof(f38525,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X6)))) ) | ($spl24 | $spl92)),
  inference(superposition,[],[f38277,f13336])).
fof(f39016,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl92)),
  inference(condensation,[],[f38524])).
fof(f38524,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl24 | $spl92)),
  inference(superposition,[],[f38277,f13017])).
fof(f39237,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39225,f38426])).
fof(f39225,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | ($spl24 | $spl92)),
  inference(condensation,[],[f39143])).
fof(f39143,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl24 | $spl92)),
  inference(backward_demodulation,[],[f38392,f1823])).
fof(f39236,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK19(X2,X1,X0))) ) | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39108,f38426])).
fof(f39108,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK19(X2,X1,X0))) ) | ($spl24 | $spl92)),
  inference(backward_demodulation,[],[f38392,f528])).
fof(f38392,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f38277,f4407])).
fof(f38426,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f38277,f661])).
fof(f31723,plain,(
  ( ! [X171,X169,X172,X170] : (s(bool,X170) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X169))) | s(bool,X170) = s(bool,X171) | s(bool,'F_0') = s(bool,X171) | s(bool,X169) = s(bool,sK19(X172,'_i_0_0',bool))) ) | $spl24),
  inference(subsumption_resolution,[],[f31559,f424])).
fof(f31559,plain,(
  ( ! [X171,X169,X172,X170] : (~p(s(bool,'T_0')) | s(bool,X170) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X169))) | s(bool,X170) = s(bool,X171) | s(bool,'F_0') = s(bool,X171) | s(bool,X169) = s(bool,sK19(X172,'_i_0_0',bool))) ) | $spl24),
  inference(superposition,[],[f13960,f13336])).
fof(f30151,plain,(
  ( ! [X265,X267,X266] : (s(bool,X265) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X266),'const.pair.,_2'(s(bool,'F_0'),s(X266,X267))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X265)))) ) | $spl24),
  inference(subsumption_resolution,[],[f29568,f13001])).
fof(f29568,plain,(
  ( ! [X265,X267,X264,X266] : (s(bool,X265) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X266),'const.pair.,_2'(s(bool,'F_0'),s(X266,X267))))) | s(bool,'T_0') = s(bool,sK19(X264,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X265)))) ) | $spl24),
  inference(superposition,[],[f12237,f13336])).
fof(f29318,plain,(
  ( ! [X255,X257,X256] : (s(bool,X255) = s(bool,happ(s(fun(X256,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X256,X257))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X255)))) ) | $spl24),
  inference(subsumption_resolution,[],[f28748,f13001])).
fof(f28748,plain,(
  ( ! [X255,X257,X254,X256] : (s(bool,X255) = s(bool,happ(s(fun(X256,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X256,X257))) | s(bool,'T_0') = s(bool,sK19(X254,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X255)))) ) | $spl24),
  inference(superposition,[],[f12018,f13336])).
fof(f28541,plain,(
  ( ! [X185,X186] : (s(bool,X185) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X186),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X185)))) ) | $spl24),
  inference(subsumption_resolution,[],[f28130,f13001])).
fof(f28130,plain,(
  ( ! [X185,X186,X184] : (s(bool,X185) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X186),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,sK19(X184,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X185)))) ) | $spl24),
  inference(superposition,[],[f11779,f13336])).
fof(f27873,plain,(
  ( ! [X239,X237,X238] : (s(bool,X237) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X238,bool),'const.pair.,_2'(s(X238,X239),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X237)))) ) | $spl24),
  inference(subsumption_resolution,[],[f27328,f13001])).
fof(f27328,plain,(
  ( ! [X239,X237,X238,X236] : (s(bool,X237) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X238,bool),'const.pair.,_2'(s(X238,X239),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK19(X236,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X237)))) ) | $spl24),
  inference(superposition,[],[f11642,f13336])).
fof(f20273,plain,(
  ( ! [X273,X275,X272] : (s(bool,'F_0') = s(bool,happ(s(fun(X272,bool),X273),s(X272,X275))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X272,bool),X273)))))) ) | $spl24),
  inference(subsumption_resolution,[],[f20272,f426])).
fof(f20272,plain,(
  ( ! [X273,X275,X272] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun(X272,bool),X273),s(X272,X275))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X272,bool),X273)))))) ) | $spl24),
  inference(forward_demodulation,[],[f19867,f12999])).
fof(f19867,plain,(
  ( ! [X273,X275,X272,X274] : (p(s(bool,sK19(X274,'_i_0_0',bool))) | s(bool,'F_0') = s(bool,happ(s(fun(X272,bool),X273),s(X272,X275))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X272,bool),X273)))))) ) | $spl24),
  inference(superposition,[],[f10555,f13336])).
fof(f19700,plain,(
  ( ! [X146,X147,X148] : (s(bool,X146) = s(bool,sK19(X148,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK19(X146,X147,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) ) | $spl24),
  inference(superposition,[],[f13336,f994])).
fof(f20359,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,X62) = s(bool,sK19(X65,'_i_0_0',bool))) ) | $spl24),
  inference(subsumption_resolution,[],[f20019,f424])).
fof(f20019,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'T_0')) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,X62) = s(bool,sK19(X65,'_i_0_0',bool))) ) | $spl24),
  inference(superposition,[],[f1378,f13336])).
fof(f20358,plain,(
  ( ! [X39,X41,X40] : (s(bool,X40) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X39))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,sK19(X41,'_i_0_0',bool))) ) | $spl24),
  inference(subsumption_resolution,[],[f20012,f424])).
fof(f20012,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X40) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X39))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,sK19(X41,'_i_0_0',bool))) ) | $spl24),
  inference(superposition,[],[f661,f13336])).
fof(f20191,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) = s(bool,sK19(X22,'_i_0_0',bool)) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'F_0') = s(bool,X23)) ) | $spl24),
  inference(subsumption_resolution,[],[f19779,f650])).
fof(f19779,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X21) = s(bool,sK19(X22,'_i_0_0',bool)) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'F_0') = s(bool,X23)) ) | $spl24),
  inference(superposition,[],[f13336,f10557])).
fof(f19701,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X2))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X1)))) ) | $spl24),
  inference(superposition,[],[f13336,f13336])).
fof(f13340,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) = s(bool,sK19(X63,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,sK19(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) ) | $spl24),
  inference(superposition,[],[f13017,f994])).
fof(f13339,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X58) = s(bool,sK19(X60,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,X58),s(X57,X59)))))) ) | $spl24),
  inference(superposition,[],[f13017,f579])).
fof(f13337,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X51) = s(bool,sK19(X53,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X50,bool),'_dst_x0_1'(s(bool,X51))),s(X50,X52)))) ) | $spl24),
  inference(superposition,[],[f13017,f580])).
fof(f13335,plain,(
  ( ! [X47,X45,X46] : (s(bool,X45) = s(bool,sK19(X47,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X45),s(fun('type.num.num',bool),X46),s('type.num.num','const.num.|48|_0')))) ) | $spl24),
  inference(superposition,[],[f13017,f457])).
fof(f13334,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X43) = s(bool,sK19(X44,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43)))))) ) | $spl24),
  inference(superposition,[],[f13017,f578])).
fof(f13000,plain,(
  ( ! [X0,X1] : (s(bool,sK19(X0,'_i_0_0',bool)) = s(bool,sK19(X1,'_i_0_0',bool))) ) | $spl24),
  inference(unit_resulting_resolution,[],[f1699,f1699,f572])).
fof(f39269,plain,(
  $false | ($spl24 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525,f38532,f39262,f39265,f39268])).
fof(f39266,plain,(
  $false | ($spl24 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525,f38532,f39262,f39265])).
fof(f39263,plain,(
  $false | ($spl24 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525,f38532,f39262])).
fof(f39261,plain,(
  $false | ($spl24 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525,f38532])).
fof(f39260,plain,(
  $false | ($spl24 | $spl36 | $spl92)),
  inference(global_subsumption,[],[f39010,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525])).
fof(f39010,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(condensation,[],[f38531])).
fof(f38531,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,X29) | s(bool,X28) = s(bool,'|dollar|exists_1'(s(fun(X30,bool),'_dst_x0_1'(s(bool,X28)))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(superposition,[],[f38277,f22013])).
fof(f39259,plain,(
  $false | ($spl24 | $spl70 | $spl92)),
  inference(global_subsumption,[],[f38530,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525])).
fof(f38530,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),'_dst_x0_1'(s(bool,X24)))))) ) | ($spl70 | $spl92)),
  inference(superposition,[],[f38277,f22005])).
fof(f39258,plain,(
  $false | ($spl22 | $spl24 | $spl36 | $spl92)),
  inference(global_subsumption,[],[f39011,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525])).
fof(f39011,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl22 | $spl24 | $spl36 | $spl92)),
  inference(condensation,[],[f38529])).
fof(f38529,plain,(
  ( ! [X21,X22] : (s(bool,X21) = s(bool,X22) | s(bool,X21) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl22 | $spl24 | $spl36 | $spl92)),
  inference(superposition,[],[f38277,f21741])).
fof(f39257,plain,(
  $false | ($spl24 | $spl30 | $spl36 | $spl92)),
  inference(global_subsumption,[],[f39012,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525])).
fof(f39012,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X1),bool),'const.pair.FST_0')))) ) | ($spl24 | $spl30 | $spl36 | $spl92)),
  inference(condensation,[],[f38528])).
fof(f38528,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,X17) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X19),bool),'const.pair.FST_0')))) ) | ($spl24 | $spl30 | $spl36 | $spl92)),
  inference(superposition,[],[f38277,f21740])).
fof(f39256,plain,(
  $false | ($spl24 | $spl36 | $spl92)),
  inference(global_subsumption,[],[f39013,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525])).
fof(f39013,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(condensation,[],[f38527])).
fof(f38527,plain,(
  ( ! [X14,X15,X13] : (s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,'|dollar|exists_1'(s(fun(X15,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(superposition,[],[f38277,f21738])).
fof(f39255,plain,(
  $false | ($spl24 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016,f38525])).
fof(f39254,plain,(
  $false | ($spl24 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237,f39016])).
fof(f39253,plain,(
  $false | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f38359,f38426])).
fof(f38359,plain,(
  ( ! [X0] : (p(s(bool,sK19(X0,'_i_0_0',bool)))) ) | $spl92),
  inference(unit_resulting_resolution,[],[f38277,f431])).
fof(f39252,plain,(
  $false | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39251,f7421])).
fof(f39251,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0')))))) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38360,f12999])).
fof(f38360,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,sK19(X2,'_i_0_0',bool))))))) ) | ($spl22 | $spl92)),
  inference(unit_resulting_resolution,[],[f38277,f12457])).
fof(f39250,plain,(
  $false | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39249,f7857])).
fof(f39249,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38361,f12999])).
fof(f38361,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,sK19(X0,'_i_0_0',bool)),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) ) | ($spl22 | $spl92)),
  inference(unit_resulting_resolution,[],[f38277,f12458])).
fof(f39248,plain,(
  $false | ($spl22 | $spl24 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f39247,f5273])).
fof(f39247,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38362,f12999])).
fof(f38362,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK19(X0,'_i_0_0',bool))))) ) | ($spl22 | $spl92)),
  inference(unit_resulting_resolution,[],[f38277,f12459])).
fof(f39246,plain,(
  $false | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39245,f8148])).
fof(f39245,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X2)))) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38363,f12999])).
fof(f38363,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,sK19(X1,'_i_0_0',bool)))),s(X0,X2)))) ) | ($spl22 | $spl92)),
  inference(unit_resulting_resolution,[],[f38277,f12460])).
fof(f39244,plain,(
  $false | ($spl22 | $spl24 | $spl92)),
  inference(subsumption_resolution,[],[f39243,f8626])).
fof(f39243,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X2)))))) ) | ($spl22 | $spl24 | $spl92)),
  inference(forward_demodulation,[],[f38364,f12999])).
fof(f38364,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,sK19(X1,'_i_0_0',bool)),s(X0,X2)))))) ) | ($spl22 | $spl92)),
  inference(unit_resulting_resolution,[],[f38277,f12462])).
fof(f39242,plain,(
  $false | ($spl24 | $spl70 | $spl92)),
  inference(global_subsumption,[],[f39241,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237])).
fof(f39241,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0')))))) ) | ($spl24 | $spl70 | $spl92)),
  inference(forward_demodulation,[],[f38374,f12999])).
fof(f38374,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,sK19(X1,'_i_0_0',bool))))))) ) | ($spl24 | $spl70 | $spl92)),
  inference(unit_resulting_resolution,[],[f38277,f34723])).
fof(f39240,plain,(
  $false | ($spl24 | $spl36 | $spl92)),
  inference(global_subsumption,[],[f39239,f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237])).
fof(f39239,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f39218,f38426])).
fof(f39218,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))))))) ) | ($spl24 | $spl36 | $spl92)),
  inference(condensation,[],[f39161])).
fof(f39161,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,'F_0')))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | ($spl24 | $spl36 | $spl92)),
  inference(backward_demodulation,[],[f38392,f27184])).
fof(f27184,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X3)))))))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f27116,f424])).
fof(f27116,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X3)))))))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f24563,f659])).
fof(f24563,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,X92)) | p(s(bool,'|dollar|exists_1'(s(fun(X93,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X91,bool),'_dst_x0_1'(s(bool,X92))))))))))) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f24455,f435])).
fof(f24455,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,X92)) | p(s(bool,'|dollar|exists_1'(s(fun(X93,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X91,bool),'_dst_x0_1'(s(bool,X92)))))))))) | s(bool,'F_0') = s(bool,X92)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f1687,f21891])).
fof(f21891,plain,(
  ( ! [X233,X232] : (s(bool,X232) = s(bool,'|dollar|exists_1'(s(fun(X233,bool),'_dst_x0_1'(s(bool,X232))))) | s(bool,'F_0') = s(bool,X232)) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f20907,f13001])).
fof(f20907,plain,(
  ( ! [X233,X234,X232] : (s(bool,X232) = s(bool,'|dollar|exists_1'(s(fun(X233,bool),'_dst_x0_1'(s(bool,X232))))) | s(bool,'T_0') = s(bool,sK19(X234,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X232)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f1846,f20151])).
fof(f39238,plain,(
  $false | ($spl24 | $spl92)),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f10572,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f6913,f6721,f6544,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6109,f6107,f6106,f6039,f5860,f5859,f5857,f5855,f5854,f5813,f5705,f5704,f5702,f5700,f5699,f5570,f5478,f5477,f5475,f5473,f5472,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4179,f4107,f4106,f4104,f4102,f4101,f2008,f1962,f1587,f1541,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11581,f11642,f11779,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12912,f12913,f12931,f12932,f12933,f12934,f12936,f12937,f1699,f13001,f13000,f12999,f13017,f13334,f13335,f13336,f13337,f13339,f13340,f13485,f13675,f13676,f13677,f13678,f13680,f13681,f13712,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14865,f14866,f17510,f18094,f18413,f18667,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f19701,f20191,f20358,f20359,f11850,f11488,f11209,f9963,f2024,f1893,f1787,f19700,f6357,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f4180,f4178,f1984,f1825,f1823,f1687,f3011,f3012,f18944,f18945,f20273,f16406,f16402,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f14864,f13742,f13679,f12935,f12911,f8673,f8194,f7902,f7541,f7171,f6911,f6719,f6542,f6110,f5858,f5703,f5476,f4523,f4285,f4105,f2991,f2959,f2958,f2957,f19040,f19022,f18901,f18887,f18886,f18885,f18884,f18882,f18881,f18879,f18873,f18872,f18871,f18870,f18868,f18867,f18866,f18865,f18863,f18824,f18822,f18817,f18816,f18815,f16373,f10555,f2359,f2298,f1693,f27873,f27331,f27332,f27335,f27337,f27338,f27374,f27471,f28541,f28133,f28134,f28137,f28139,f28140,f28142,f28480,f29318,f28751,f28752,f28755,f28757,f28758,f28760,f29232,f30151,f29571,f29572,f29575,f29577,f29578,f29580,f30062,f30839,f31029,f31031,f31363,f31374,f31426,f31717,f31720,f31723,f31724,f31727,f31730,f31733,f31736,f31594,f31892,f31894,f31895,f31899,f31901,f31902,f31857,f37735,f37737,f37789,f38277,f38426,f38392,f39236,f39237])).
fof(f39044,plain,(
  $false | ($spl24 | $spl92)),
  inference(subsumption_resolution,[],[f38404,f650])).
fof(f38404,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f424,f13001,f38277,f13960])).
fof(f38437,plain,(
  $false | ($spl22 | $spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f38277,f11033])).
fof(f38435,plain,(
  $false | ($spl24 | $spl30 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f38277,f3325])).
fof(f38433,plain,(
  $false | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f38277,f1961])).
fof(f38427,plain,(
  $false | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f424,f13001,f38277,f661])).
fof(f38406,plain,(
  $false | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f424,f650,f13001,f38277,f13960])).
fof(f38393,plain,(
  $false | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f650,f13001,f38277,f4407])).
fof(f38390,plain,(
  $false | ($spl24 | $spl72 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f13001,f38277,f22006])).
fof(f38370,plain,(
  $false | ($spl24 | $spl36 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f13001,f38277,f21159])).
fof(f21159,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,X11) | s(bool,X12) = s(bool,sK19(X13,'_i_0_0',bool)) | s(bool,X11) = s(bool,X12)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f650,f20151])).
fof(f38367,plain,(
  $false | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f650,f13001,f38277,f13961])).
fof(f38358,plain,(
  $false | ($spl60 | $spl92)),
  inference(unit_resulting_resolution,[],[f650,f38277,f12985])).
fof(f38356,plain,(
  $false | ($spl24 | $spl76 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f38277,f22010])).
fof(f38353,plain,(
  $false | ($spl24 | $spl92)),
  inference(unit_resulting_resolution,[],[f13001,f38277])).
fof(f38278,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X36,bool),'_dst_x0_1'(s(bool,X35))))) | s(bool,'F_0') = s(bool,X35)) ) | $spl94),
  inference(cnf_transformation,[],[f38278_D])).
fof(f38278_D,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X36,bool),'_dst_x0_1'(s(bool,X35))))) | s(bool,'F_0') = s(bool,X35)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f38277,plain,(
  ( ! [X37,X34] : (s(bool,X34) = s(bool,sK19(X37,'_i_0_0',bool))) ) | $spl92),
  inference(cnf_transformation,[],[f38277_D])).
fof(f38277_D,plain,(
  ( ! [X37,X34] : (s(bool,X34) = s(bool,sK19(X37,'_i_0_0',bool))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f38265,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,sK19(X1,'_i_0_0',bool)) | s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,X3))))) | s(bool,'F_0') = s(bool,X3)) ) | ($spl24 | $spl36)),
  inference(condensation,[],[f38151])).
fof(f38151,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X83,bool),'_dst_x0_1'(s(bool,X82))))) | s(bool,X81) = s(bool,sK19(X84,'_i_0_0',bool)) | s(bool,X81) = s(bool,sK19(X85,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X82)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f36152,f20151])).
fof(f36152,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,sK19(X1,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(forward_demodulation,[],[f35794,f12999])).
fof(f35794,plain,(
  ( ! [X2,X0,X1] : (s(bool,sK19(X1,'_i_0_0',bool)) != s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,sK19(X1,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(equality_factoring,[],[f22013])).
fof(f38279,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X40,bool),'_dst_x0_1'(s(bool,X39))))) | s(bool,X38) = s(bool,sK19(X41,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X39)) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f38141,f21402])).
fof(f21402,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) = s(bool,sK19(X0,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f21378])).
fof(f21378,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) = s(bool,sK19(X0,'_i_0_0',bool)) | s(bool,X1) = s(bool,sK19(X0,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(equality_factoring,[],[f20151])).
fof(f38141,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X40,bool),'_dst_x0_1'(s(bool,X39))))) | s(bool,X38) = s(bool,sK19(X41,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X39)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f36152,f10557])).
fof(f38276,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X36,bool),'_dst_x0_1'(s(bool,X35))))) | s(bool,X34) = s(bool,sK19(X37,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X35)) ) | ($spl24 | $spl36)),
  inference(subsumption_resolution,[],[f38140,f21402])).
fof(f38140,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X36,bool),'_dst_x0_1'(s(bool,X35))))) | s(bool,X34) = s(bool,sK19(X37,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f36152,f10557])).
fof(f30223,plain,(
  ( ! [X292,X294,X293,X295] : (s(bool,X294) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X292),'const.pair.,_2'(s(bool,'F_0'),s(X292,X293))))) | s(bool,X294) = s(bool,sK19(X295,'_i_0_0',bool))) ) | $spl90),
  inference(cnf_transformation,[],[f30223_D])).
fof(f30223_D,plain,(
  ( ! [X292,X294,X293,X295] : (s(bool,X294) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X292),'const.pair.,_2'(s(bool,'F_0'),s(X292,X293))))) | s(bool,X294) = s(bool,sK19(X295,'_i_0_0',bool))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f30049,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s(bool,X294) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X292),'const.pair.,_2'(s(bool,'F_0'),s(X292,X293))))) | s(bool,X294) = s(bool,sK19(X295,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X291)) ) | $spl72),
  inference(duplicate_literal_removal,[],[f29836])).
fof(f29836,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s(bool,X294) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X292),'const.pair.,_2'(s(bool,'F_0'),s(X292,X293))))) | s(bool,X294) = s(bool,sK19(X295,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X291) | s(bool,'T_0') = s(bool,X291)) ) | $spl72),
  inference(superposition,[],[f22006,f12237])).
fof(f30193,plain,(
  ( ! [X132,X130,X131] : (s(bool,X132) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X130),'const.pair.,_2'(s(bool,'F_0'),s(X130,X131))))) | s(bool,'F_0') = s(bool,X132)) ) | $spl88),
  inference(cnf_transformation,[],[f30193_D])).
fof(f30193_D,plain,(
  ( ! [X132,X130,X131] : (s(bool,X132) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X130),'const.pair.,_2'(s(bool,'F_0'),s(X130,X131))))) | s(bool,'F_0') = s(bool,X132)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f30062,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X132) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X130),'const.pair.,_2'(s(bool,'F_0'),s(X130,X131))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'F_0') = s(bool,X132)) )),
  inference(duplicate_literal_removal,[],[f29797])).
fof(f29797,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X132) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X130),'const.pair.,_2'(s(bool,'F_0'),s(X130,X131))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'F_0') = s(bool,X132) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f4407,f12237])).
fof(f29389,plain,(
  ( ! [X288,X290,X287,X289] : (s(bool,X289) != s(bool,happ(s(fun(X287,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X287,X288))) | s(bool,X289) = s(bool,sK19(X290,'_i_0_0',bool))) ) | $spl86),
  inference(cnf_transformation,[],[f29389_D])).
fof(f29389_D,plain,(
  ( ! [X288,X290,X287,X289] : (s(bool,X289) != s(bool,happ(s(fun(X287,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X287,X288))) | s(bool,X289) = s(bool,sK19(X290,'_i_0_0',bool))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f29220,plain,(
  ( ! [X288,X290,X287,X289,X286] : (s(bool,X289) != s(bool,happ(s(fun(X287,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X287,X288))) | s(bool,X289) = s(bool,sK19(X290,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X286)) ) | $spl72),
  inference(duplicate_literal_removal,[],[f29011])).
fof(f29011,plain,(
  ( ! [X288,X290,X287,X289,X286] : (s(bool,X289) != s(bool,happ(s(fun(X287,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X287,X288))) | s(bool,X289) = s(bool,sK19(X290,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X286) | s(bool,'T_0') = s(bool,X286)) ) | $spl72),
  inference(superposition,[],[f22006,f12018])).
fof(f29360,plain,(
  ( ! [X132,X130,X131] : (s(bool,X132) != s(bool,happ(s(fun(X130,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X130,X131))) | s(bool,'F_0') = s(bool,X132)) ) | $spl84),
  inference(cnf_transformation,[],[f29360_D])).
fof(f29360_D,plain,(
  ( ! [X132,X130,X131] : (s(bool,X132) != s(bool,happ(s(fun(X130,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X130,X131))) | s(bool,'F_0') = s(bool,X132)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f29232,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X132) != s(bool,happ(s(fun(X130,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X130,X131))) | s(bool,'T_0') = s(bool,X129) | s(bool,'F_0') = s(bool,X132)) )),
  inference(duplicate_literal_removal,[],[f28973])).
fof(f28973,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X132) != s(bool,happ(s(fun(X130,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X130,X131))) | s(bool,'T_0') = s(bool,X129) | s(bool,'F_0') = s(bool,X132) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f4407,f12018])).
fof(f28611,plain,(
  ( ! [X216,X215,X214] : (s(bool,X215) != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X214),s('type.num.num','const.num.|48|_0'))) | s(bool,X215) = s(bool,sK19(X216,'_i_0_0',bool))) ) | $spl82),
  inference(cnf_transformation,[],[f28611_D])).
fof(f28611_D,plain,(
  ( ! [X216,X215,X214] : (s(bool,X215) != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X214),s('type.num.num','const.num.|48|_0'))) | s(bool,X215) = s(bool,sK19(X216,'_i_0_0',bool))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f28469,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,X215) != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X214),s('type.num.num','const.num.|48|_0'))) | s(bool,X215) = s(bool,sK19(X216,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X213)) ) | $spl72),
  inference(duplicate_literal_removal,[],[f28263])).
fof(f28263,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,X215) != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X214),s('type.num.num','const.num.|48|_0'))) | s(bool,X215) = s(bool,sK19(X216,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | $spl72),
  inference(superposition,[],[f22006,f11779])).
fof(f28583,plain,(
  ( ! [X99,X98] : (s(bool,X99) != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X99)) ) | $spl80),
  inference(cnf_transformation,[],[f28583_D])).
fof(f28583_D,plain,(
  ( ! [X99,X98] : (s(bool,X99) != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X99)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f28480,plain,(
  ( ! [X99,X97,X98] : (s(bool,X99) != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X97) | s(bool,'F_0') = s(bool,X99)) )),
  inference(duplicate_literal_removal,[],[f28226])).
fof(f28226,plain,(
  ( ! [X99,X97,X98] : (s(bool,X99) != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X97) | s(bool,'F_0') = s(bool,X99) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f4407,f11779])).
fof(f27950,plain,(
  ( ! [X280,X279,X281,X278] : (s(bool,X280) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X278,bool),'const.pair.,_2'(s(X278,X279),s(bool,'F_0'))))) | s(bool,X280) = s(bool,sK19(X281,'_i_0_0',bool))) ) | $spl78),
  inference(cnf_transformation,[],[f27950_D])).
fof(f27950_D,plain,(
  ( ! [X280,X279,X281,X278] : (s(bool,X280) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X278,bool),'const.pair.,_2'(s(X278,X279),s(bool,'F_0'))))) | s(bool,X280) = s(bool,sK19(X281,'_i_0_0',bool))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f27783,plain,(
  ( ! [X280,X277,X279,X281,X278] : (s(bool,X280) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X278,bool),'const.pair.,_2'(s(X278,X279),s(bool,'F_0'))))) | s(bool,X280) = s(bool,sK19(X281,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X277)) ) | $spl72),
  inference(duplicate_literal_removal,[],[f27582])).
fof(f27582,plain,(
  ( ! [X280,X277,X279,X281,X278] : (s(bool,X280) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X278,bool),'const.pair.,_2'(s(X278,X279),s(bool,'F_0'))))) | s(bool,X280) = s(bool,sK19(X281,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X277) | s(bool,'T_0') = s(bool,X277)) ) | $spl72),
  inference(superposition,[],[f22006,f11642])).
fof(f22220,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f427,f428,f429,f430,f432,f431,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f715,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1477,f1506,f1508,f1586,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f2004,f2007,f3070,f2976,f3071,f3207,f2977,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f5622,f5881,f6133,f6388,f6391,f6664,f6826,f7114,f7378,f7444,f7537,f7538,f7540,f7542,f7543,f7898,f7899,f7901,f7903,f7904,f8190,f8191,f8193,f8195,f8196,f8217,f8669,f8670,f8672,f8674,f8675,f8696,f10507,f10509,f10510,f10511,f8797,f10512,f10514,f10515,f10516,f8317,f10517,f10519,f10520,f10521,f7937,f10522,f10524,f10525,f10526,f7569,f4407,f10557,f3561,f3578,f2344,f2355,f1787,f1893,f2024,f9963,f6357,f1542,f1540,f1535,f1533,f1530,f1529,f1822,f1823,f1687,f3011,f3012,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f8673,f8194,f7902,f7541,f4285,f2959,f2958,f1693,f2298,f2359,f10572,f10555,f10538,f10535,f10532,f10529,f8626,f8148,f7857,f7421,f7173,f7171,f6913,f6911,f6721,f6719,f6544,f6542,f6392,f6387,f6386,f6385,f6384,f6383,f6298,f6112,f6111,f6110,f6109,f6107,f6106,f6039,f5860,f5859,f5858,f5857,f5855,f5854,f5813,f5705,f5704,f5703,f5702,f5700,f5699,f5570,f5478,f5477,f5476,f5475,f5473,f5472,f4523,f4325,f4324,f4323,f4321,f4320,f4319,f4318,f4180,f4179,f4178,f4107,f4106,f4105,f4104,f4102,f4101,f2991,f2957,f2008,f1984,f1962,f1825,f1587,f1543,f1541,f1532,f1531,f1510,f1482,f1458,f1457,f1455,f1454,f1453,f1452,f932,f868,f808,f752,f693,f436,f435,f426,f11209,f11488,f11581,f11642,f11779,f11850,f11963,f12018,f12182,f12237,f12877,f12907,f12908,f12909,f12910,f12911,f12912,f12913,f12931,f12932,f12933,f12934,f12935,f12936,f12937,f13675,f13676,f13677,f13678,f13679,f13680,f13681,f13712,f13742,f13744,f13960,f13961,f13963,f13964,f13965,f13966,f14494,f14495,f14827,f14860,f14861,f14862,f14863,f14864,f14865,f14866,f16373,f16402,f16406,f17510,f18094,f18413,f18667,f18815,f18816,f18817,f19022,f18822,f18824,f18863,f18865,f18866,f18867,f18868,f18870,f18871,f18872,f18873,f19040,f18879,f18881,f18882,f18884,f18885,f18886,f18887,f18901,f18944,f18945,f19086,f19088,f19119,f19121,f19280,f19281,f19415,f19417,f19419,f19450,f19452,f10540,f22217,f22218,f22219])).
fof(f22219,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl50),
  inference(subsumption_resolution,[],[f22206,f22218])).
fof(f22206,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl50),
  inference(condensation,[],[f22179])).
fof(f22179,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl50),
  inference(backward_demodulation,[],[f10540,f1823])).
fof(f22218,plain,(
  ( ! [X22] : (~p(s(bool,X22))) ) | $spl50),
  inference(subsumption_resolution,[],[f22178,f426])).
fof(f22178,plain,(
  ( ! [X22] : (p(s(bool,'F_0')) | ~p(s(bool,X22))) ) | $spl50),
  inference(backward_demodulation,[],[f10540,f1687])).
fof(f22217,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK19(X2,X1,X0))) ) | $spl50),
  inference(subsumption_resolution,[],[f22154,f426])).
fof(f22154,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK19(X2,X1,X0))) ) | $spl50),
  inference(backward_demodulation,[],[f10540,f528])).
fof(f22010,plain,(
  ( ! [X196,X194,X195] : (s(bool,X194) != s(bool,sK19(X195,'_i_0_0',bool)) | s(bool,X194) = s(bool,sK19(X196,'_i_0_0',bool))) ) | $spl76),
  inference(cnf_transformation,[],[f22010_D])).
fof(f22010_D,plain,(
  ( ! [X196,X194,X195] : (s(bool,X194) != s(bool,sK19(X195,'_i_0_0',bool)) | s(bool,X194) = s(bool,sK19(X196,'_i_0_0',bool))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f21453,plain,(
  ( ! [X196,X194,X195,X193] : (s(bool,X194) != s(bool,sK19(X195,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X193) | s(bool,X194) = s(bool,sK19(X196,'_i_0_0',bool))) ) | ($spl24 | $spl36 | $spl60)),
  inference(duplicate_literal_removal,[],[f21051])).
fof(f21051,plain,(
  ( ! [X196,X194,X195,X193] : (s(bool,X194) != s(bool,sK19(X195,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X193) | s(bool,X194) = s(bool,sK19(X196,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X193)) ) | ($spl24 | $spl36 | $spl60)),
  inference(superposition,[],[f12985,f20151])).
fof(f22007,plain,(
  ( ! [X134,X135] : (s(bool,X134) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,X134) = s(bool,sK19(X135,'_i_0_0',bool))) ) | $spl74),
  inference(cnf_transformation,[],[f22007_D])).
fof(f22007_D,plain,(
  ( ! [X134,X135] : (s(bool,X134) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,X134) = s(bool,sK19(X135,'_i_0_0',bool))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f21460,plain,(
  ( ! [X134,X135,X133] : (s(bool,X134) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X133) | s(bool,X134) = s(bool,sK19(X135,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f21035])).
fof(f21035,plain,(
  ( ! [X134,X135,X133] : (s(bool,X134) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X133) | s(bool,X134) = s(bool,sK19(X135,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X133)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f5244,f20151])).
fof(f22006,plain,(
  ( ! [X127,X128,X126] : (s(bool,X126) != s(bool,X127) | s(bool,X126) = s(bool,sK19(X128,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X127)) ) | $spl72),
  inference(cnf_transformation,[],[f22006_D])).
fof(f22006_D,plain,(
  ( ! [X127,X128,X126] : (s(bool,X126) != s(bool,X127) | s(bool,X126) = s(bool,sK19(X128,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X127)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f21461,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,X126) != s(bool,X127) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,X125) | s(bool,X126) = s(bool,sK19(X128,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f21033])).
fof(f21033,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,X126) != s(bool,X127) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,X125) | s(bool,X126) = s(bool,sK19(X128,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X125)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f4407,f20151])).
fof(f22005,plain,(
  ( ! [X118,X120,X119] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X119,bool),'_dst_x0_1'(s(bool,X118))))) | s(bool,X118) = s(bool,sK19(X120,'_i_0_0',bool))) ) | $spl70),
  inference(cnf_transformation,[],[f22005_D])).
fof(f22005_D,plain,(
  ( ! [X118,X120,X119] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X119,bool),'_dst_x0_1'(s(bool,X118))))) | s(bool,X118) = s(bool,sK19(X120,'_i_0_0',bool))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f21462,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X119,bool),'_dst_x0_1'(s(bool,X118))))) | s(bool,'F_0') = s(bool,X117) | s(bool,X118) = s(bool,sK19(X120,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f21031])).
fof(f21031,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X119,bool),'_dst_x0_1'(s(bool,X118))))) | s(bool,'F_0') = s(bool,X117) | s(bool,X118) = s(bool,sK19(X120,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X117)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f2008,f20151])).
fof(f21715,plain,(
  ( ! [X47,X48] : (s(bool,X47) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,X47) = s(bool,sK19(X48,'_i_0_0',bool))) ) | $spl68),
  inference(cnf_transformation,[],[f21715_D])).
fof(f21715_D,plain,(
  ( ! [X47,X48] : (s(bool,X47) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,X47) = s(bool,sK19(X48,'_i_0_0',bool))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f21547,plain,(
  ( ! [X138,X139,X137] : (s(bool,'F_0') = s(bool,X137) | s(bool,X138) = s(bool,sK19(X139,'_i_0_0',bool)) | s(bool,X138) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) ) | ($spl22 | $spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f20620])).
fof(f20620,plain,(
  ( ! [X138,X139,X137] : (s(bool,'F_0') = s(bool,X137) | s(bool,X138) = s(bool,sK19(X139,'_i_0_0',bool)) | s(bool,X138) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X137)) ) | ($spl22 | $spl24 | $spl36)),
  inference(superposition,[],[f20151,f11463])).
fof(f21714,plain,(
  ( ! [X45,X43,X44] : (s(bool,X44) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))) | s(bool,X44) = s(bool,sK19(X45,'_i_0_0',bool))) ) | $spl66),
  inference(cnf_transformation,[],[f21714_D])).
fof(f21714_D,plain,(
  ( ! [X45,X43,X44] : (s(bool,X44) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))) | s(bool,X44) = s(bool,sK19(X45,'_i_0_0',bool))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f21548,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,'F_0') = s(bool,X131) | s(bool,X132) = s(bool,sK19(X133,'_i_0_0',bool)) | s(bool,X132) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X130),bool),'const.pair.FST_0')))) ) | ($spl24 | $spl30 | $spl36)),
  inference(duplicate_literal_removal,[],[f20618])).
fof(f20618,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,'F_0') = s(bool,X131) | s(bool,X132) = s(bool,sK19(X133,'_i_0_0',bool)) | s(bool,X132) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X130),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X131)) ) | ($spl24 | $spl30 | $spl36)),
  inference(superposition,[],[f20151,f10581])).
fof(f21713,plain,(
  ( ! [X21,X22,X20] : (s(bool,X21) = s(bool,'|dollar|exists_1'(s(fun(X20,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X21) = s(bool,sK19(X22,'_i_0_0',bool))) ) | $spl64),
  inference(cnf_transformation,[],[f21713_D])).
fof(f21713_D,plain,(
  ( ! [X21,X22,X20] : (s(bool,X21) = s(bool,'|dollar|exists_1'(s(fun(X20,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X21) = s(bool,sK19(X22,'_i_0_0',bool))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f21549,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,X120) | s(bool,X121) = s(bool,sK19(X122,'_i_0_0',bool)) | s(bool,X121) = s(bool,'|dollar|exists_1'(s(fun(X119,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f20615])).
fof(f20615,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,X120) | s(bool,X121) = s(bool,sK19(X122,'_i_0_0',bool)) | s(bool,X121) = s(bool,'|dollar|exists_1'(s(fun(X119,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X120)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f20151,f1962])).
fof(f21671,plain,(
  ( ! [X76,X74,X72,X75] : (s(bool,X72) = s(bool,sK19(X74,'_i_0_0',bool)) | s(bool,X72) = s(bool,X75) | s(bool,X75) = s(bool,sK19(X76,'_i_0_0',bool))) ) | $spl62),
  inference(cnf_transformation,[],[f21671_D])).
fof(f21671_D,plain,(
  ( ! [X76,X74,X72,X75] : (s(bool,X72) = s(bool,sK19(X74,'_i_0_0',bool)) | s(bool,X72) = s(bool,X75) | s(bool,X75) = s(bool,sK19(X76,'_i_0_0',bool))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f21550,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,X56) = s(bool,sK19(X57,'_i_0_0',bool)) | s(bool,X54) = s(bool,X56) | s(bool,X54) = s(bool,sK19(X58,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f20605])).
fof(f20605,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,X56) = s(bool,sK19(X57,'_i_0_0',bool)) | s(bool,X54) = s(bool,X56) | s(bool,X54) = s(bool,sK19(X58,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X55)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f20151,f20151])).
fof(f21559,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,X68) = s(bool,X69) | s(bool,X69) = s(bool,sK19(X70,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X67) | s(bool,X68) = s(bool,sK19(X71,'_i_0_0',bool))) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f20535])).
fof(f20535,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,X68) = s(bool,X69) | s(bool,X69) = s(bool,sK19(X70,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X67) | s(bool,X68) = s(bool,sK19(X71,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X67)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f20151,f20151])).
fof(f21563,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,X47) = s(bool,sK19(X48,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X46)) ) | ($spl22 | $spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f20530])).
fof(f20530,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,X47) = s(bool,sK19(X48,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X46)) ) | ($spl22 | $spl24 | $spl36)),
  inference(superposition,[],[f20151,f11463])).
fof(f21564,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X44) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))) | s(bool,X44) = s(bool,sK19(X45,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X42)) ) | ($spl24 | $spl30 | $spl36)),
  inference(duplicate_literal_removal,[],[f20529])).
fof(f20529,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X44) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X43),bool),'const.pair.FST_0'))) | s(bool,X44) = s(bool,sK19(X45,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | ($spl24 | $spl30 | $spl36)),
  inference(superposition,[],[f20151,f10581])).
fof(f21569,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,'|dollar|exists_1'(s(fun(X20,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X21) = s(bool,sK19(X22,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X19)) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f20523])).
fof(f20523,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,'|dollar|exists_1'(s(fun(X20,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X21) = s(bool,sK19(X22,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X19)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f20151,f1962])).
fof(f21586,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,sK19(X74,'_i_0_0',bool)) | s(bool,X75) = s(bool,sK19(X76,'_i_0_0',bool)) | s(bool,X72) = s(bool,X75)) ) | ($spl24 | $spl36)),
  inference(duplicate_literal_removal,[],[f20478])).
fof(f20478,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'F_0') = s(bool,X73) | s(bool,X72) = s(bool,sK19(X74,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X73) | s(bool,X75) = s(bool,sK19(X76,'_i_0_0',bool)) | s(bool,X72) = s(bool,X75)) ) | ($spl24 | $spl36)),
  inference(superposition,[],[f20151,f20151])).
fof(f12985,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,sK19(X15,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X16)) ) | $spl60),
  inference(cnf_transformation,[],[f12985_D])).
fof(f12985_D,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,sK19(X15,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f13485,plain,(
  ( ! [X61,X59,X60] : (s(bool,X61) != s(bool,sK19(X60,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,X61)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f13372])).
fof(f13372,plain,(
  ( ! [X61,X59,X60] : (s(bool,X61) != s(bool,sK19(X60,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X59)) ) | $spl24),
  inference(superposition,[],[f4407,f13017])).
fof(f12939,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) != s(bool,sK19(X15,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X16)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f12919])).
fof(f12919,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) != s(bool,sK19(X15,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X14)) ) | $spl20),
  inference(superposition,[],[f4407,f10870])).
fof(f10870,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK19(X2,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl20),
  inference(subsumption_resolution,[],[f10831,f650])).
fof(f10831,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK19(X2,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl20),
  inference(superposition,[],[f1695,f659])).
fof(f10736,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f10686,f5243])).
fof(f10686,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1)))))) ) | $spl34),
  inference(backward_demodulation,[],[f5243,f8626])).
fof(f10735,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f10685,f5243])).
fof(f10685,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)))) ) | $spl34),
  inference(backward_demodulation,[],[f5243,f8148])).
fof(f10734,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f10684,f5243])).
fof(f10684,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'T_0'),s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) ) | $spl34),
  inference(backward_demodulation,[],[f5243,f7857])).
fof(f10733,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f10683,f5243])).
fof(f10683,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'T_0')))))) ) | $spl34),
  inference(backward_demodulation,[],[f5243,f7421])).
fof(f10732,plain,(
  $false | ($spl34 | $spl36)),
  inference(subsumption_resolution,[],[f10629,f5243])).
fof(f10629,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | ($spl34 | $spl36)),
  inference(backward_demodulation,[],[f5243,f5273])).
fof(f10728,plain,(
  $false | $spl34),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1541,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f1962,f2004,f2007,f2008,f3070,f1787,f1893,f2024,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f1984,f1825,f1823,f1687,f3011,f3012,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f2976,f3071,f3207,f2991,f2977,f2959,f2958,f2957,f1693,f2298,f2359,f4101,f4102,f4104,f4105,f4106,f4107,f4178,f4179,f4180,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f4318,f4319,f4320,f4321,f4323,f4324,f4325,f5472,f5473,f5475,f5476,f5477,f5478,f5570,f5622,f5699,f5700,f5702,f5703,f5704,f5705,f5813,f5854,f5855,f5857,f5858,f5859,f5860,f5881,f6039,f6106,f6107,f6109,f6110,f6111,f6112,f6133,f6298,f6383,f6384,f6385,f6386,f6387,f6388,f6391,f6392,f6357,f6542,f6544,f6664,f6719,f6721,f6826,f6911,f6913,f7114,f7171,f7173,f7378,f7444,f7537,f7538,f7540,f7541,f7542,f7543,f7898,f7899,f7901,f7902,f7903,f7904,f8190,f8191,f8193,f8194,f8195,f8196,f8217,f8669,f8670,f8672,f8673,f8674,f8675,f8696,f9963,f10507,f10509,f10510,f10511,f8797,f8626,f10512,f10514,f10515,f10516,f8317,f8148,f10517,f10519,f10520,f10521,f7937,f7857,f10522,f10524,f10525,f10526,f7569,f7421,f10529,f10532,f10535,f10538,f4407,f4523,f10555,f10557,f4285,f10572,f3561,f3578,f2344,f2355,f5243,f10727])).
fof(f10727,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl34),
  inference(subsumption_resolution,[],[f10584,f5243])).
fof(f10584,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl34),
  inference(backward_demodulation,[],[f5243,f435])).
fof(f10726,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f10583,f424])).
fof(f10583,plain,(
  ~p(s(bool,'T_0')) | $spl34),
  inference(backward_demodulation,[],[f5243,f426])).
fof(f1697,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl22),
  inference(cnf_transformation,[],[f1697_D])).
fof(f1697_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2355,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2337,f424])).
fof(f2337,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1693,f693])).
fof(f2344,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f1693,f454])).
fof(f3578,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3550,f424])).
fof(f3550,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1693,f693])).
fof(f3561,plain,(
  ( ! [X39] : (~p(s(bool,X39)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f1693,f454])).
fof(f10541,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl53),
  inference(cnf_transformation,[],[f10541_D])).
fof(f10541_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f10553,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | $spl36),
  inference(condensation,[],[f10552])).
fof(f10552,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,X1) = s(bool,X2)) ) | $spl36),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1541,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f1962,f2004,f2007,f2008,f3070,f1787,f1893,f2024,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f1984,f1825,f1823,f1687,f3011,f3012,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f2976,f3071,f3207,f2991,f2977,f2959,f2958,f2957,f1693,f2298,f2359,f4101,f4102,f4104,f4105,f4106,f4107,f4178,f4179,f4180,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f4318,f4319,f4320,f4321,f4323,f4324,f4325,f5244,f5364,f5365,f5307,f5366,f5367,f5368,f5311,f5314,f5472,f5473,f5475,f5476,f5477,f5478,f5570,f5622,f5699,f5700,f5702,f5703,f5704,f5705,f5813,f5854,f5855,f5857,f5858,f5859,f5860,f5881,f6039,f6106,f6107,f6109,f6110,f6111,f6112,f6133,f6298,f6383,f6384,f6385,f6386,f6387,f6388,f6391,f6392,f6357,f6542,f6544,f6664,f6719,f6721,f6826,f6911,f6913,f7114,f7171,f7173,f7378,f7444,f7537,f7538,f7540,f7541,f7542,f7543,f7898,f7899,f7901,f7902,f7903,f7904,f8190,f8191,f8193,f8194,f8195,f8196,f8217,f8669,f8670,f8672,f8673,f8674,f8675,f8696,f9963,f10507,f10509,f10510,f10511,f8797,f8626,f10512,f10514,f10515,f10516,f8317,f8148,f10517,f10519,f10520,f10521,f7937,f7857,f10522,f10524,f10525,f10526,f7569,f7421,f10529,f10532,f10535,f10538,f5370,f5330,f5329,f5369,f5299,f5297,f10548,f5361,f5284,f10551,f5281])).
fof(f5281,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl36),
  inference(superposition,[],[f5244,f659])).
fof(f5284,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl36),
  inference(superposition,[],[f5244,f659])).
fof(f5297,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'F_0') = s(bool,X29)) ) | $spl36),
  inference(superposition,[],[f5244,f693])).
fof(f5314,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X6))) | s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | $spl36),
  inference(superposition,[],[f5244,f430])).
fof(f5311,plain,(
  ( ! [X57,X58] : (s(bool,X57) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,sK19(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) ) | $spl36),
  inference(superposition,[],[f5244,f994])).
fof(f5368,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') != s(bool,X55) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X54),'const.pair.,_2'(s(bool,X55),s(X54,X56)))))) ) | $spl36),
  inference(forward_demodulation,[],[f5310,f454])).
fof(f5310,plain,(
  ( ! [X54,X56,X55] : (s(bool,X55) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X54),'const.pair.,_2'(s(bool,X55),s(X54,X56)))))) ) | $spl36),
  inference(superposition,[],[f5244,f579])).
fof(f5367,plain,(
  ( ! [X52,X53] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),X53))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X52),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X52),X53)))) ) | $spl36),
  inference(forward_demodulation,[],[f5309,f454])).
fof(f5309,plain,(
  ( ! [X52,X53] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),X53))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X52),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X52),X53)))) ) | $spl36),
  inference(superposition,[],[f5244,f576])).
fof(f5366,plain,(
  ( ! [X50,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,'F_0') = s(bool,happ(s(fun(X49,bool),'_dst_x0_1'(s(bool,X50))),s(X49,X51)))) ) | $spl36),
  inference(forward_demodulation,[],[f5308,f454])).
fof(f5308,plain,(
  ( ! [X50,X51,X49] : (s(bool,X50) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X49,bool),'_dst_x0_1'(s(bool,X50))),s(X49,X51)))) ) | $spl36),
  inference(superposition,[],[f5244,f580])).
fof(f5307,plain,(
  ( ! [X48] : (s(bool,X48) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X48)))) ) | $spl36),
  inference(superposition,[],[f5244,f454])).
fof(f5365,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X46),s(fun('type.num.num',bool),X47),s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(forward_demodulation,[],[f5306,f454])).
fof(f5306,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X46),s(fun('type.num.num',bool),X47),s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(superposition,[],[f5244,f457])).
fof(f5364,plain,(
  ( ! [X45,X43,X44] : (s(bool,'F_0') != s(bool,X45) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,X45)))))) ) | $spl36),
  inference(forward_demodulation,[],[f5305,f454])).
fof(f5305,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,X45)))))) ) | $spl36),
  inference(superposition,[],[f5244,f578])).
fof(f10551,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | $spl36),
  inference(condensation,[],[f10550])).
fof(f10550,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X7)) ) | $spl36),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1541,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f1962,f2004,f2007,f2008,f3070,f1787,f1893,f2024,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f1984,f1825,f1823,f1687,f3011,f3012,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f2976,f3071,f3207,f2991,f2977,f2959,f2958,f2957,f1693,f2298,f2359,f4101,f4102,f4104,f4105,f4106,f4107,f4178,f4179,f4180,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f4318,f4319,f4320,f4321,f4323,f4324,f4325,f5244,f5364,f5365,f5307,f5366,f5367,f5368,f5311,f5314,f5472,f5473,f5475,f5476,f5477,f5478,f5570,f5622,f5699,f5700,f5702,f5703,f5704,f5705,f5813,f5854,f5855,f5857,f5858,f5859,f5860,f5881,f6039,f6106,f6107,f6109,f6110,f6111,f6112,f6133,f6298,f6383,f6384,f6385,f6386,f6387,f6388,f6391,f6392,f6357,f6542,f6544,f6664,f6719,f6721,f6826,f6911,f6913,f7114,f7171,f7173,f7378,f7444,f7537,f7538,f7540,f7541,f7542,f7543,f7898,f7899,f7901,f7902,f7903,f7904,f8190,f8191,f8193,f8194,f8195,f8196,f8217,f8669,f8670,f8672,f8673,f8674,f8675,f8696,f9963,f10507,f10509,f10510,f10511,f8797,f8626,f10512,f10514,f10515,f10516,f8317,f8148,f10517,f10519,f10520,f10521,f7937,f7857,f10522,f10524,f10525,f10526,f7569,f7421,f10529,f10532,f10535,f10538,f5370,f5330,f5329,f5369,f5299,f5297,f10548,f5361,f5284])).
fof(f5361,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X10)) ) | $spl36),
  inference(forward_demodulation,[],[f5352,f1846])).
fof(f5352,plain,(
  ( ! [X10,X11] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X10)) ) | $spl36),
  inference(duplicate_literal_removal,[],[f5286])).
fof(f5286,plain,(
  ( ! [X10,X11] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X10)) ) | $spl36),
  inference(superposition,[],[f5244,f1962])).
fof(f10549,plain,(
  ( ! [X29] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29)))) ) | $spl58),
  inference(cnf_transformation,[],[f10549_D])).
fof(f10549_D,plain,(
  ( ! [X29] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f10548,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29)))) ) | $spl36),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1541,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1694,f1698,f1686,f1805,f1847,f1846,f1913,f1961,f1962,f2004,f2007,f2008,f3070,f1787,f1893,f2024,f1543,f1542,f1540,f1535,f1533,f1532,f1531,f1530,f1529,f1822,f1984,f1825,f1823,f1687,f3011,f3012,f1826,f1585,f1547,f1516,f1515,f1514,f1513,f1512,f1511,f992,f991,f990,f989,f988,f528,f2976,f3071,f3207,f2991,f2977,f2959,f2958,f2957,f1693,f2298,f2359,f4101,f4102,f4104,f4105,f4106,f4107,f4178,f4179,f4180,f4235,f4236,f4237,f4238,f4240,f4241,f4242,f4318,f4319,f4320,f4321,f4323,f4324,f4325,f5244,f5364,f5365,f5307,f5366,f5367,f5368,f5311,f5314,f5472,f5473,f5475,f5476,f5477,f5478,f5570,f5622,f5699,f5700,f5702,f5703,f5704,f5705,f5813,f5854,f5855,f5857,f5858,f5859,f5860,f5881,f6039,f6106,f6107,f6109,f6110,f6111,f6112,f6133,f6298,f6383,f6384,f6385,f6386,f6387,f6388,f6391,f6392,f6357,f6542,f6544,f6664,f6719,f6721,f6826,f6911,f6913,f7114,f7171,f7173,f7378,f7444,f7537,f7538,f7540,f7541,f7542,f7543,f7898,f7899,f7901,f7902,f7903,f7904,f8190,f8191,f8193,f8194,f8195,f8196,f8217,f8669,f8670,f8672,f8673,f8674,f8675,f8696,f9963,f10507,f10509,f10510,f10511,f8797,f8626,f10512,f10514,f10515,f10516,f8317,f8148,f10517,f10519,f10520,f10521,f7937,f7857,f10522,f10524,f10525,f10526,f7569,f7421,f10529,f10532,f10535,f10538,f5370,f5330,f5329,f5369,f5299,f5297])).
fof(f10547,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X33) = s(X31,sK19(X33,X32,X31))) ) | $spl56),
  inference(cnf_transformation,[],[f10547_D])).
fof(f10547_D,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X33) = s(X31,sK19(X33,X32,X31))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5299,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X33) = s(X31,sK19(X33,X32,X31))) ) | $spl36),
  inference(superposition,[],[f5244,f994])).
fof(f10539,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X12)) ) | $spl48),
  inference(cnf_transformation,[],[f10539_D])).
fof(f10539_D,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X12)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5369,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl36),
  inference(subsumption_resolution,[],[f5328,f430])).
fof(f5328,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl36),
  inference(superposition,[],[f5244,f659])).
fof(f10542,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl52),
  inference(cnf_transformation,[],[f10542_D])).
fof(f10542_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f10543,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X10)) ) | $spl54),
  inference(cnf_transformation,[],[f10543_D])).
fof(f10543_D,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X10)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5329,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | $spl36),
  inference(superposition,[],[f5244,f659])).
fof(f5330,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl36),
  inference(superposition,[],[f5244,f659])).
fof(f5370,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,X11) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X12)) ) | $spl36),
  inference(subsumption_resolution,[],[f5331,f430])).
fof(f5331,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X12)) ) | $spl36),
  inference(superposition,[],[f5244,f659])).
fof(f10523,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X106,bool),'const.pair.,_2'(s(X106,X107),s(bool,'F_0')))))) ) | $spl46),
  inference(cnf_transformation,[],[f10523_D])).
fof(f10523_D,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X106,bool),'const.pair.,_2'(s(X106,X107),s(bool,'F_0')))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7569,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f7500])).
fof(f7500,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6664,f430])).
fof(f10526,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X69,bool),'const.pair.,_2'(s(X69,X70),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f7518,f6664])).
fof(f7518,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X69,bool),'const.pair.,_2'(s(X69,X70),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X66,bool),'const.pair.,_2'(s(X66,X67),s(bool,X68))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f6664,f1452])).
fof(f10525,plain,(
  ( ! [X78,X76,X79] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X78,bool),'const.pair.,_2'(s(X78,X79),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f7520,f6826])).
fof(f7520,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X78,bool),'const.pair.,_2'(s(X78,X79),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X76),s(fun('type.num.num',bool),X77),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f6664,f1453])).
fof(f10524,plain,(
  ( ! [X94,X97,X96] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X96,bool),'const.pair.,_2'(s(X96,X97),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f7525,f7114])).
fof(f7525,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X96,bool),'const.pair.,_2'(s(X96,X97),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(X93,bool),'_dst_x0_1'(s(bool,X94))),s(X93,X95))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f6664,f1454])).
fof(f10522,plain,(
  ( ! [X107,X106,X104] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X106,bool),'const.pair.,_2'(s(X106,X107),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f7527,f7378])).
fof(f7527,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X106,bool),'const.pair.,_2'(s(X106,X107),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105))))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f6664,f1455])).
fof(f10518,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.|48|_0')))) ) | $spl44),
  inference(cnf_transformation,[],[f10518_D])).
fof(f10518_D,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7937,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f7861])).
fof(f7861,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6826,f430])).
fof(f10521,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X51),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f7879,f6664])).
fof(f7879,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X51),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X48,bool),'const.pair.,_2'(s(X48,X49),s(bool,X50))))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f6826,f1452])).
fof(f10520,plain,(
  ( ! [X58,X56] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X58),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f7881,f6826])).
fof(f7881,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X58),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X56),s(fun('type.num.num',bool),X57),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f6826,f1453])).
fof(f10519,plain,(
  ( ! [X71,X69] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X71),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f7886,f7114])).
fof(f7886,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X71),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),'_dst_x0_1'(s(bool,X69))),s(X68,X70))) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f6826,f1454])).
fof(f10517,plain,(
  ( ! [X79,X77] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f7888,f7378])).
fof(f7888,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X76),'const.pair.,_2'(s(bool,X77),s(X76,X78))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f6826,f1455])).
fof(f10513,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,happ(s(fun(X106,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X106,X107)))) ) | $spl42),
  inference(cnf_transformation,[],[f10513_D])).
fof(f10513_D,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,happ(s(fun(X106,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X106,X107)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f8317,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f8153])).
fof(f8153,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f7114,f430])).
fof(f10516,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X69,X70))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f8171,f6664])).
fof(f8171,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') != s(bool,happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X69,X70))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X66,bool),'const.pair.,_2'(s(X66,X67),s(bool,X68))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f7114,f1452])).
fof(f10515,plain,(
  ( ! [X78,X76,X79] : (s(bool,'T_0') != s(bool,happ(s(fun(X78,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X78,X79))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f8173,f6826])).
fof(f8173,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,happ(s(fun(X78,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X78,X79))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X76),s(fun('type.num.num',bool),X77),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f7114,f1453])).
fof(f10514,plain,(
  ( ! [X94,X97,X96] : (s(bool,'T_0') != s(bool,happ(s(fun(X96,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X96,X97))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f8178,f7114])).
fof(f8178,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,'T_0') != s(bool,happ(s(fun(X96,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X96,X97))) | s(bool,'T_0') = s(bool,happ(s(fun(X93,bool),'_dst_x0_1'(s(bool,X94))),s(X93,X95))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f7114,f1454])).
fof(f10512,plain,(
  ( ! [X107,X106,X104] : (s(bool,'T_0') != s(bool,happ(s(fun(X106,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X106,X107))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f8180,f7378])).
fof(f8180,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,happ(s(fun(X106,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X106,X107))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105))))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f7114,f1455])).
fof(f10508,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,'F_0'),s(X106,X107)))))) ) | $spl40),
  inference(cnf_transformation,[],[f10508_D])).
fof(f10508_D,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,'F_0'),s(X106,X107)))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f8797,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f8632])).
fof(f8632,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X2))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f7378,f430])).
fof(f10511,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),'const.pair.,_2'(s(bool,'F_0'),s(X69,X70))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f8650,f6664])).
fof(f8650,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),'const.pair.,_2'(s(bool,'F_0'),s(X69,X70))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X66,bool),'const.pair.,_2'(s(X66,X67),s(bool,X68))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f7378,f1452])).
fof(f10510,plain,(
  ( ! [X78,X76,X79] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,'F_0'),s(X78,X79))))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f8652,f6826])).
fof(f8652,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,'F_0'),s(X78,X79))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X76),s(fun('type.num.num',bool),X77),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f7378,f1453])).
fof(f10509,plain,(
  ( ! [X94,X97,X96] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X96),'const.pair.,_2'(s(bool,'F_0'),s(X96,X97))))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f8657,f7114])).
fof(f8657,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X96),'const.pair.,_2'(s(bool,'F_0'),s(X96,X97))))) | s(bool,'T_0') = s(bool,happ(s(fun(X93,bool),'_dst_x0_1'(s(bool,X94))),s(X93,X95))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f7378,f1454])).
fof(f10507,plain,(
  ( ! [X107,X106,X104] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,'F_0'),s(X106,X107))))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f8659,f7378])).
fof(f8659,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,'F_0'),s(X106,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105))))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f7378,f1455])).
fof(f10463,plain,(
  ( ! [X175,X173,X174] : (s(bool,X175) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X173,bool),'const.pair.,_2'(s(X173,X174),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X175)) ) | $spl38),
  inference(cnf_transformation,[],[f10463_D])).
fof(f10463_D,plain,(
  ( ! [X175,X173,X174] : (s(bool,X175) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X173,bool),'const.pair.,_2'(s(X173,X174),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X175)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f10365,plain,(
  ( ! [X175,X173,X174,X172] : (s(bool,X175) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X173,bool),'const.pair.,_2'(s(X173,X174),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X172) | s(bool,'F_0') = s(bool,X175)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f10189])).
fof(f10189,plain,(
  ( ! [X175,X173,X174,X172] : (s(bool,X175) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X173,bool),'const.pair.,_2'(s(X173,X174),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X172) | s(bool,'F_0') = s(bool,X175) | s(bool,'T_0') = s(bool,X172)) ) | $spl22),
  inference(superposition,[],[f2221,f7767])).
fof(f7767,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X8)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f7753])).
fof(f7753,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X8) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X8)) ) | $spl22),
  inference(superposition,[],[f7587,f659])).
fof(f7587,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0')))))) ) | $spl22),
  inference(forward_demodulation,[],[f7422,f454])).
fof(f7422,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))))))) ) | $spl22),
  inference(unit_resulting_resolution,[],[f4810,f6664])).
fof(f4810,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl22),
  inference(unit_resulting_resolution,[],[f650,f4771])).
fof(f4771,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl22),
  inference(subsumption_resolution,[],[f4707,f426])).
fof(f4707,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl22),
  inference(superposition,[],[f431,f2206])).
fof(f2206,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'T_0') = s(bool,X24)) ) | $spl22),
  inference(forward_demodulation,[],[f2109,f1712])).
fof(f1712,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl22),
  inference(unit_resulting_resolution,[],[f1697,f432])).
fof(f2109,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24)))) ) | $spl22),
  inference(superposition,[],[f1771,f454])).
fof(f1771,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X8)) ) | $spl22),
  inference(subsumption_resolution,[],[f1750,f424])).
fof(f1750,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | s(bool,X8) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X8)) ) | $spl22),
  inference(superposition,[],[f1457,f1712])).
fof(f2221,plain,(
  ( ! [X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,X19)) ) | $spl22),
  inference(subsumption_resolution,[],[f2220,f424])).
fof(f2220,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,X19)) ) | $spl22),
  inference(forward_demodulation,[],[f2134,f1712])).
fof(f2134,plain,(
  ( ! [X19,X20] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,X19)) ) | $spl22),
  inference(superposition,[],[f661,f1771])).
fof(f5244,plain,(
  ( ! [X51] : (s(bool,X51) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X51)) ) | $spl36),
  inference(cnf_transformation,[],[f5244_D])).
fof(f5244_D,plain,(
  ( ! [X51] : (s(bool,X51) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X51)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5184,plain,(
  ( ! [X50,X51] : (s(bool,X51) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X51)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f5092])).
fof(f5092,plain,(
  ( ! [X50,X51] : (s(bool,X51) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) | $spl22),
  inference(superposition,[],[f2221,f4947])).
fof(f4947,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f4935])).
fof(f4935,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f4810,f659])).
fof(f3206,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f528,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f988,f989,f990,f991,f992,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1511,f1512,f1513,f1514,f1515,f1516,f1529,f1530,f1531,f1532,f1533,f1535,f1540,f1541,f1542,f1543,f1547,f1585,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1693,f1694,f1698,f1686,f1687,f1787,f1805,f1823,f1825,f1826,f1822,f1847,f1846,f1893,f1913,f1961,f1962,f1984,f2004,f2007,f2008,f2024,f2298,f2359,f2957,f2958,f2959,f2976,f2977,f2991,f3011,f3012,f3071,f3070,f3077,f3102,f3199,f3200,f3097,f3090,f3084])).
fof(f3084,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl32),
  inference(unit_resulting_resolution,[],[f3077,f3077,f572])).
fof(f3090,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl32),
  inference(unit_resulting_resolution,[],[f3077,f3077,f572])).
fof(f3097,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl32),
  inference(unit_resulting_resolution,[],[f3077,f431])).
fof(f3200,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl32),
  inference(subsumption_resolution,[],[f3196,f426])).
fof(f3196,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl32),
  inference(condensation,[],[f3180])).
fof(f3180,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(backward_demodulation,[],[f3102,f1823])).
fof(f3199,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK19(X2,X1,X0))) ) | $spl32),
  inference(subsumption_resolution,[],[f3162,f426])).
fof(f3162,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK19(X2,X1,X0))) ) | $spl32),
  inference(backward_demodulation,[],[f3102,f528])).
fof(f3102,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl32),
  inference(unit_resulting_resolution,[],[f3077,f1825])).
fof(f3205,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f3087,f426])).
fof(f3087,plain,(
  p(s(bool,'F_0')) | $spl32),
  inference(unit_resulting_resolution,[],[f650,f3077,f572])).
fof(f3204,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f528,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f988,f989,f990,f991,f992,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1511,f1512,f1513,f1514,f1515,f1516,f1529,f1530,f1531,f1532,f1533,f1535,f1540,f1541,f1542,f1543,f1547,f1585,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1693,f1694,f1698,f1686,f1687,f1787,f1805,f1823,f1825,f1826,f1822,f1847,f1846,f1893,f1913,f1961,f1962,f1984,f2004,f2007,f2008,f2024,f2298,f2359,f2957,f2958,f2959,f2976,f2977,f2991,f3011,f3012,f3071,f3070,f3077,f3102,f3199,f3200,f3097,f3090])).
fof(f3203,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f3093,f426])).
fof(f3093,plain,(
  p(s(bool,'F_0')) | $spl32),
  inference(unit_resulting_resolution,[],[f650,f3077,f572])).
fof(f3202,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f528,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f988,f989,f990,f991,f992,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1511,f1512,f1513,f1514,f1515,f1516,f1529,f1530,f1531,f1532,f1533,f1535,f1540,f1541,f1542,f1543,f1547,f1585,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1693,f1694,f1698,f1686,f1687,f1787,f1805,f1823,f1825,f1826,f1822,f1847,f1846,f1893,f1913,f1961,f1962,f1984,f2004,f2007,f2008,f2024,f2298,f2359,f2957,f2958,f2959,f2976,f2977,f2991,f3011,f3012,f3071,f3070,f3077,f3102,f3199,f3200,f3097])).
fof(f3201,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f423,f422,f421,f420,f419,f418,f417,f416,f415,f414,f424,f426,f427,f428,f429,f430,f432,f431,f436,f435,f454,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f486,f487,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f504,f505,f506,f507,f508,f509,f510,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f529,f528,f539,f538,f537,f536,f535,f534,f533,f532,f531,f530,f550,f551,f552,f553,f554,f555,f556,f557,f558,f559,f560,f561,f562,f563,f564,f565,f566,f567,f568,f569,f570,f571,f631,f575,f572,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f587,f588,f589,f590,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f601,f650,f659,f661,f693,f715,f752,f808,f868,f932,f988,f989,f990,f991,f992,f994,f1022,f1023,f1024,f1025,f1026,f1040,f1050,f1051,f1052,f1053,f1054,f1378,f1379,f1452,f1453,f1454,f1455,f1457,f1458,f1477,f1482,f1506,f1508,f1510,f1511,f1512,f1513,f1514,f1515,f1516,f1529,f1530,f1531,f1532,f1533,f1535,f1540,f1541,f1542,f1543,f1547,f1585,f1586,f1587,f1688,f1689,f1690,f1691,f1692,f1673,f1693,f1694,f1698,f1686,f1687,f1787,f1805,f1823,f1825,f1826,f1822,f1847,f1846,f1893,f1913,f1961,f1962,f1984,f2004,f2007,f2008,f2024,f2298,f2359,f2957,f2958,f2959,f2976,f2977,f2991,f3011,f3012,f3071,f3070,f3077,f3102,f3199,f3200])).
fof(f3105,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f424,f3077,f1687])).
fof(f3103,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f650,f3077,f1825])).
fof(f3098,plain,(
  $false | ($spl22 | $spl32)),
  inference(unit_resulting_resolution,[],[f1712,f3077,f431])).
fof(f3096,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f650,f3077,f436])).
fof(f3094,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f426,f650,f3077,f572])).
fof(f3092,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f3077,f650,f3077,f572])).
fof(f3088,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f426,f650,f3077,f572])).
fof(f3086,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f3077,f650,f3077,f572])).
fof(f3082,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f650,f650,f3077,f715])).
fof(f3078,plain,(
  $false | $spl32),
  inference(unit_resulting_resolution,[],[f424,f3077])).
fof(f3076,plain,(
  ( ! [X28] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))) ) | $spl30),
  inference(cnf_transformation,[],[f3076_D])).
fof(f3076_D,plain,(
  ( ! [X28] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3070,plain,(
  ( ! [X28,X29] : (~p(s(bool,X29)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f2977,f579])).
fof(f2005,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),'_dst_x0_1'(s(bool,X11)))))) ) | $spl26),
  inference(cnf_transformation,[],[f2005_D])).
fof(f2005_D,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),'_dst_x0_1'(s(bool,X11)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2007,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X15)) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X13,bool),'_dst_x0_1'(s(bool,X14)))))) )),
  inference(subsumption_resolution,[],[f1990,f431])).
fof(f1990,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X15)) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X13,bool),'_dst_x0_1'(s(bool,X14)))))) )),
  inference(superposition,[],[f1825,f659])).
fof(f2004,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),'_dst_x0_1'(s(bool,X11)))))) )),
  inference(subsumption_resolution,[],[f1989,f431])).
fof(f1989,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),'_dst_x0_1'(s(bool,X11))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1825,f659])).
fof(f1699,plain,(
  ( ! [X20] : (~p(s(bool,sK19(X20,'_i_0_0',bool)))) ) | $spl24),
  inference(cnf_transformation,[],[f1699_D])).
fof(f1699_D,plain,(
  ( ! [X20] : (~p(s(bool,sK19(X20,'_i_0_0',bool)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1686,plain,(
  ( ! [X20] : (~p(s(bool,sK19(X20,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f529,f454])).
fof(f1695,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK19(X0,'_i_0_0',bool))) ) | $spl20),
  inference(cnf_transformation,[],[f1695_D])).
fof(f1695_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK19(X0,'_i_0_0',bool))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1694,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK19(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f1680,f424])).
fof(f1680,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK19(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f529,f693])).
fof(f641,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(trivial_inequality_removal,[],[f640])).
fof(f640,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f639,f606])).
fof(f639,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl1 | $spl4 | $spl6 | $spl8 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f638,f608])).
fof(f638,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl1 | $spl6 | $spl8 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f637,f610])).
fof(f637,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl1 | $spl8 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f636,f612])).
fof(f636,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl1 | $spl10 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f635,f614])).
fof(f635,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl1 | $spl12 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f634,f616])).
fof(f634,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl1 | $spl14 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f633,f618])).
fof(f633,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) | ($spl1 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f632,f620])).
fof(f632,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK9))) | ($spl1 | $spl18)),
  inference(forward_demodulation,[],[f605,f622])).
fof(f604,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | $spl0),
  inference(cnf_transformation,[],[f604_D])).
fof(f604_D,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f606,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),sK10) | $spl2),
  inference(cnf_transformation,[],[f606_D])).
fof(f606_D,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),sK10) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f414,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),sK10) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f389,plain,(
  (s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | (s('type.sptree.spt'('type.bvlSem.v'),sK1) = s('type.sptree.spt'('type.bvlSem.v'),sK10) & s('type.list.list'('type.bvpSem.stack'),sK2) = s('type.list.list'('type.bvpSem.stack'),sK11) & s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK12) & s('type.num.num',sK4) = s('type.num.num',sK13) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) & s('type.num.num',sK6) = s('type.num.num',sK15) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) & s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),sK17) & s('type.num.num',sK9) = s('type.num.num',sK18))) & (s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | s('type.sptree.spt'('type.bvlSem.v'),sK1) != s('type.sptree.spt'('type.bvlSem.v'),sK10) | s('type.list.list'('type.bvpSem.stack'),sK2) != s('type.list.list'('type.bvpSem.stack'),sK11) | s('type.option.option'('type.num.num'),sK3) != s('type.option.option'('type.num.num'),sK12) | s('type.num.num',sK4) != s('type.num.num',sK13) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) | s('type.num.num',sK6) != s('type.num.num',sK15) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) | s('type.ffi.ffi_state'(sK0),sK8) != s('type.ffi.ffi_state'(sK0),sK17) | s('type.num.num',sK9) != s('type.num.num',sK18))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10,sK11,sK12,sK13,sK14,sK15,sK16,sK17,sK18])],[f388])).
fof(f388,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : ((s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X10),s('type.list.list'('type.bvpSem.stack'),X11),s('type.option.option'('type.num.num'),X12),s('type.num.num',X13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14),s('type.num.num',X15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16),s('type.ffi.ffi_state'(X0),X17),s('type.num.num',X18))) | (s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),X10) & s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),X11) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X12) & s('type.num.num',X4) = s('type.num.num',X13) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14) & s('type.num.num',X6) = s('type.num.num',X15) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16) & s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),X17) & s('type.num.num',X9) = s('type.num.num',X18))) & (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9))) != s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X10),s('type.list.list'('type.bvpSem.stack'),X11),s('type.option.option'('type.num.num'),X12),s('type.num.num',X13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14),s('type.num.num',X15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16),s('type.ffi.ffi_state'(X0),X17),s('type.num.num',X18))) | s('type.sptree.spt'('type.bvlSem.v'),X1) != s('type.sptree.spt'('type.bvlSem.v'),X10) | s('type.list.list'('type.bvpSem.stack'),X2) != s('type.list.list'('type.bvpSem.stack'),X11) | s('type.option.option'('type.num.num'),X3) != s('type.option.option'('type.num.num'),X12) | s('type.num.num',X4) != s('type.num.num',X13) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14) | s('type.num.num',X6) != s('type.num.num',X15) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16) | s('type.ffi.ffi_state'(X0),X8) != s('type.ffi.ffi_state'(X0),X17) | s('type.num.num',X9) != s('type.num.num',X18)))),
  inference(flattening,[],[f387])).
fof(f387,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : ((s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X10),s('type.list.list'('type.bvpSem.stack'),X11),s('type.option.option'('type.num.num'),X12),s('type.num.num',X13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14),s('type.num.num',X15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16),s('type.ffi.ffi_state'(X0),X17),s('type.num.num',X18))) | (s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),X10) & s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),X11) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X12) & s('type.num.num',X4) = s('type.num.num',X13) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14) & s('type.num.num',X6) = s('type.num.num',X15) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16) & s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),X17) & s('type.num.num',X9) = s('type.num.num',X18))) & (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9))) != s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X10),s('type.list.list'('type.bvpSem.stack'),X11),s('type.option.option'('type.num.num'),X12),s('type.num.num',X13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14),s('type.num.num',X15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16),s('type.ffi.ffi_state'(X0),X17),s('type.num.num',X18))) | (s('type.sptree.spt'('type.bvlSem.v'),X1) != s('type.sptree.spt'('type.bvlSem.v'),X10) | s('type.list.list'('type.bvpSem.stack'),X2) != s('type.list.list'('type.bvpSem.stack'),X11) | s('type.option.option'('type.num.num'),X3) != s('type.option.option'('type.num.num'),X12) | s('type.num.num',X4) != s('type.num.num',X13) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14) | s('type.num.num',X6) != s('type.num.num',X15) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16) | s('type.ffi.ffi_state'(X0),X8) != s('type.ffi.ffi_state'(X0),X17) | s('type.num.num',X9) != s('type.num.num',X18))))),
  inference(nnf_transformation,[],[f382])).
fof(f382,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X10),s('type.list.list'('type.bvpSem.stack'),X11),s('type.option.option'('type.num.num'),X12),s('type.num.num',X13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14),s('type.num.num',X15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16),s('type.ffi.ffi_state'(X0),X17),s('type.num.num',X18))) <~> (s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),X10) & s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),X11) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X12) & s('type.num.num',X4) = s('type.num.num',X13) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14) & s('type.num.num',X6) = s('type.num.num',X15) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16) & s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),X17) & s('type.num.num',X9) = s('type.num.num',X18)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X10),s('type.list.list'('type.bvpSem.stack'),X11),s('type.option.option'('type.num.num'),X12),s('type.num.num',X13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14),s('type.num.num',X15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16),s('type.ffi.ffi_state'(X0),X17),s('type.num.num',X18))) <=> (s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),X10) & s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),X11) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X12) & s('type.num.num',X4) = s('type.num.num',X13) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14) & s('type.num.num',X6) = s('type.num.num',X15) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16) & s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),X17) & s('type.num.num',X9) = s('type.num.num',X18)))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X10),s('type.list.list'('type.bvpSem.stack'),X11),s('type.option.option'('type.num.num'),X12),s('type.num.num',X13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14),s('type.num.num',X15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16),s('type.ffi.ffi_state'(X0),X17),s('type.num.num',X18))) <=> (s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),X10) & s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),X11) & s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),X12) & s('type.num.num',X4) = s('type.num.num',X13) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X14) & s('type.num.num',X6) = s('type.num.num',X15) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X16) & s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),X17) & s('type.num.num',X9) = s('type.num.num',X18)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24] : (s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X7),s('type.list.list'('type.bvpSem.stack'),X8),s('type.option.option'('type.num.num'),X9),s('type.num.num',X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X11),s('type.num.num',X12),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X13),s('type.ffi.ffi_state'(X6),X14),s('type.num.num',X15))) = s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X16),s('type.list.list'('type.bvpSem.stack'),X17),s('type.option.option'('type.num.num'),X18),s('type.num.num',X19),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20),s('type.num.num',X21),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X22),s('type.ffi.ffi_state'(X6),X23),s('type.num.num',X24))) <=> (s('type.sptree.spt'('type.bvlSem.v'),X7) = s('type.sptree.spt'('type.bvlSem.v'),X16) & s('type.list.list'('type.bvpSem.stack'),X8) = s('type.list.list'('type.bvpSem.stack'),X17) & s('type.option.option'('type.num.num'),X9) = s('type.option.option'('type.num.num'),X18) & s('type.num.num',X10) = s('type.num.num',X19) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X11) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20) & s('type.num.num',X12) = s('type.num.num',X21) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X13) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X22) & s('type.ffi.ffi_state'(X6),X14) = s('type.ffi.ffi_state'(X6),X23) & s('type.num.num',X15) = s('type.num.num',X24)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24] : (s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X7),s('type.list.list'('type.bvpSem.stack'),X8),s('type.option.option'('type.num.num'),X9),s('type.num.num',X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X11),s('type.num.num',X12),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X13),s('type.ffi.ffi_state'(X6),X14),s('type.num.num',X15))) = s('type.bvpSem.state'(X6),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X16),s('type.list.list'('type.bvpSem.stack'),X17),s('type.option.option'('type.num.num'),X18),s('type.num.num',X19),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20),s('type.num.num',X21),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X22),s('type.ffi.ffi_state'(X6),X23),s('type.num.num',X24))) <=> (s('type.sptree.spt'('type.bvlSem.v'),X7) = s('type.sptree.spt'('type.bvlSem.v'),X16) & s('type.list.list'('type.bvpSem.stack'),X8) = s('type.list.list'('type.bvpSem.stack'),X17) & s('type.option.option'('type.num.num'),X9) = s('type.option.option'('type.num.num'),X18) & s('type.num.num',X10) = s('type.num.num',X19) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X11) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20) & s('type.num.num',X12) = s('type.num.num',X21) & s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X13) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X22) & s('type.ffi.ffi_state'(X6),X14) = s('type.ffi.ffi_state'(X6),X23) & s('type.num.num',X15) = s('type.num.num',X24)))),
  file('thm.bvpSem.state_|49||49|','thm.bvpSem.state_|49||49|')).
fof(f608,plain,(
  s('type.list.list'('type.bvpSem.stack'),sK2) = s('type.list.list'('type.bvpSem.stack'),sK11) | $spl4),
  inference(cnf_transformation,[],[f608_D])).
fof(f608_D,plain,(
  s('type.list.list'('type.bvpSem.stack'),sK2) = s('type.list.list'('type.bvpSem.stack'),sK11) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f415,plain,(
  s('type.list.list'('type.bvpSem.stack'),sK2) = s('type.list.list'('type.bvpSem.stack'),sK11) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f610,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK12) | $spl6),
  inference(cnf_transformation,[],[f610_D])).
fof(f610_D,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK12) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f416,plain,(
  s('type.option.option'('type.num.num'),sK3) = s('type.option.option'('type.num.num'),sK12) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f612,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK13) | $spl8),
  inference(cnf_transformation,[],[f612_D])).
fof(f612_D,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK13) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f417,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK13) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f614,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) | $spl10),
  inference(cnf_transformation,[],[f614_D])).
fof(f614_D,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f418,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f616,plain,(
  s('type.num.num',sK6) = s('type.num.num',sK15) | $spl12),
  inference(cnf_transformation,[],[f616_D])).
fof(f616_D,plain,(
  s('type.num.num',sK6) = s('type.num.num',sK15) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f419,plain,(
  s('type.num.num',sK6) = s('type.num.num',sK15) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f618,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) | $spl14),
  inference(cnf_transformation,[],[f618_D])).
fof(f618_D,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f420,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f620,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),sK17) | $spl16),
  inference(cnf_transformation,[],[f620_D])).
fof(f620_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),sK17) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f421,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),sK17) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f622,plain,(
  s('type.num.num',sK9) = s('type.num.num',sK18) | $spl18),
  inference(cnf_transformation,[],[f622_D])).
fof(f622_D,plain,(
  s('type.num.num',sK9) = s('type.num.num',sK18) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f422,plain,(
  s('type.num.num',sK9) = s('type.num.num',sK18) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) = s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
fof(f605,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) | $spl1),
  inference(cnf_transformation,[],[f605_D])).
fof(f605_D,plain,(
  s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f607,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) != s('type.sptree.spt'('type.bvlSem.v'),sK10) | $spl3),
  inference(cnf_transformation,[],[f607_D])).
fof(f607_D,plain,(
  s('type.sptree.spt'('type.bvlSem.v'),sK1) != s('type.sptree.spt'('type.bvlSem.v'),sK10) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f609,plain,(
  s('type.list.list'('type.bvpSem.stack'),sK2) != s('type.list.list'('type.bvpSem.stack'),sK11) | $spl5),
  inference(cnf_transformation,[],[f609_D])).
fof(f609_D,plain,(
  s('type.list.list'('type.bvpSem.stack'),sK2) != s('type.list.list'('type.bvpSem.stack'),sK11) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f611,plain,(
  s('type.option.option'('type.num.num'),sK3) != s('type.option.option'('type.num.num'),sK12) | $spl7),
  inference(cnf_transformation,[],[f611_D])).
fof(f611_D,plain,(
  s('type.option.option'('type.num.num'),sK3) != s('type.option.option'('type.num.num'),sK12) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f613,plain,(
  s('type.num.num',sK4) != s('type.num.num',sK13) | $spl9),
  inference(cnf_transformation,[],[f613_D])).
fof(f613_D,plain,(
  s('type.num.num',sK4) != s('type.num.num',sK13) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f615,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) | $spl11),
  inference(cnf_transformation,[],[f615_D])).
fof(f615_D,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f617,plain,(
  s('type.num.num',sK6) != s('type.num.num',sK15) | $spl13),
  inference(cnf_transformation,[],[f617_D])).
fof(f617_D,plain,(
  s('type.num.num',sK6) != s('type.num.num',sK15) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f619,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) | $spl15),
  inference(cnf_transformation,[],[f619_D])).
fof(f619_D,plain,(
  s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f621,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) != s('type.ffi.ffi_state'(sK0),sK17) | $spl17),
  inference(cnf_transformation,[],[f621_D])).
fof(f621_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) != s('type.ffi.ffi_state'(sK0),sK17) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f623,plain,(
  s('type.num.num',sK9) != s('type.num.num',sK18) | $spl19),
  inference(cnf_transformation,[],[f623_D])).
fof(f623_D,plain,(
  s('type.num.num',sK9) != s('type.num.num',sK18) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f423,plain,(
  s('type.num.num',sK9) != s('type.num.num',sK18) | s('type.ffi.ffi_state'(sK0),sK8) != s('type.ffi.ffi_state'(sK0),sK17) | s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7) != s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16) | s('type.num.num',sK6) != s('type.num.num',sK15) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14) | s('type.num.num',sK4) != s('type.num.num',sK13) | s('type.option.option'('type.num.num'),sK3) != s('type.option.option'('type.num.num'),sK12) | s('type.list.list'('type.bvpSem.stack'),sK2) != s('type.list.list'('type.bvpSem.stack'),sK11) | s('type.sptree.spt'('type.bvlSem.v'),sK1) != s('type.sptree.spt'('type.bvlSem.v'),sK10) | s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK1),s('type.list.list'('type.bvpSem.stack'),sK2),s('type.option.option'('type.num.num'),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK5),s('type.num.num',sK6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9))) != s('type.bvpSem.state'(sK0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),sK10),s('type.list.list'('type.bvpSem.stack'),sK11),s('type.option.option'('type.num.num'),sK12),s('type.num.num',sK13),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK14),s('type.num.num',sK15),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),sK16),s('type.ffi.ffi_state'(sK0),sK17),s('type.num.num',sK18)))),
  inference(cnf_transformation,[],[f389])).
% SZS output end Proof for thm.bvpSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 338
% Passive clauses: 2493
% Generated clauses: 73026
% Final active clauses: 242
% Final passive clauses: 1399
% Input formulas: 133
% Initial clauses: 188
% 
% Duplicate literals: 2400
% Trivial inequalities: 708
% Fw subsumption resolutions: 6210
% Fw demodulations: 10542
% Bw demodulations: 520
% Condensations: 6658
% Global subsumptions: 45
% 
% Simple tautologies: 666
% Equational tautologies: 381
% Forward subsumptions: 42633
% Fw demodulations to eq. taut.: 152
% Bw demodulations to eq. taut.: 45
% 
% Unit resulting resolution: 3898
% Factoring: 1
% Forward superposition: 20833
% Backward superposition: 19689
% Self superposition: 419
% Equality factoring: 231
% Equality resolution: 57
% 
% Split clauses: 87
% Split components: 182
% SAT solver clauses: 2817
% SAT solver unit clauses: 174
% SAT solver binary clauses: 696
% 
% Sat splits: 87
% Sat splitting refutations: 103
% 
% Memory used [KB]: 36204
% Time elapsed: 2.879 s
% ------------------------------
% ------------------------------
% Success in time 4.231 s
