Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.inferProps
% (114133)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 112
% Passive clauses: 3399
% Generated clauses: 19135
% Final active clauses: 106
% Final passive clauses: 3106
% Input formulas: 44
% Initial clauses: 80
% 
% Duplicate literals: 420
% Trivial inequalities: 138
% Fw subsumption resolutions: 1865
% Fw demodulations: 2766
% Bw demodulations: 4
% 
% Simple tautologies: 189
% Equational tautologies: 67
% Forward subsumptions: 9705
% 
% Binary resolution: 19
% Forward superposition: 7334
% Backward superposition: 5973
% Self superposition: 284
% Equality factoring: 176
% Equality resolution: 25
% 
% Split clauses: 173
% Split components: 352
% SAT solver clauses: 375
% SAT solver unit clauses: 1
% SAT solver binary clauses: 39
% 
% Sat splits: 173
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 375
% TWLsolver calls for satisfiability: 86
% 
% Memory used [KB]: 19701
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.inferProps
% (114152)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 784
% Generated clauses: 10645
% Final active clauses: 85
% Final passive clauses: 682
% Input formulas: 44
% Initial clauses: 80
% 
% Duplicate literals: 330
% Trivial inequalities: 52
% Fw subsumption resolutions: 525
% Fw demodulations: 744
% Bw demodulations: 13
% Condensations: 1025
% Global subsumptions: 7
% 
% Simple tautologies: 109
% Equational tautologies: 42
% Forward subsumptions: 5388
% 
% Binary resolution: 12
% Forward superposition: 4086
% Backward superposition: 2997
% Self superposition: 193
% Equality factoring: 102
% Equality resolution: 15
% 
% Split clauses: 58
% Split components: 122
% SAT solver clauses: 1400
% SAT solver unit clauses: 123
% SAT solver binary clauses: 65
% SAT solver learnt clauses: 14
% SAT solver learnt literals: 10
% 
% Sat splits: 58
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1372
% TWLsolver calls for satisfiability: 629
% 
% Memory used [KB]: 14839
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.inferProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.inferProps
% SZS output start Proof for thm.inferProps
fof(f13480,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f161,f244,f242,f240,f160,f243,f159,f241,f239,f273,f276,f274,f975,f976,f992,f977,f978,f993,f979,f994,f991,f3421,f3422,f3423,f3424,f3425,f3443,f3444,f3445,f3446,f3447,f3451,f3453,f3455,f3456,f3457,f3458,f3459,f3460,f3500,f3549,f3502,f3550,f3504,f3505,f3506,f3507,f3508,f3552,f3548,f3509,f3551,f3547,f3604,f3618,f3605,f3606,f3619,f3607,f3622,f3620,f8719,f8720,f8721,f8722,f8723,f8724,f8725,f8726,f8727,f8728,f8729,f8730,f8731,f8732,f8733,f8734,f8735,f8736,f8737,f8738,f8739,f8740,f8741,f8742,f8743,f8744,f8745,f8746,f8747,f8748,f8749,f8750,f8751,f8752,f8753,f8754,f8755,f8756,f8757,f8758,f8759,f8760,f8761,f8762,f8763,f8764,f8768,f8773,f8769,f8772,f8767,f8770,f8777,f8920,f11396,f11128,f11783,f12357,f12370,f12918,f13467])).
fof(f13467,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12917,f242])).
fof(f12917,plain,(
  s(sK1,sK3) = s(sK1,sK5) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f12371,f209])).
fof(f209,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & ((s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6)) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7,X8] : ! [X14,X15,X11,X16] : (s('type.pair.prod'('type.infer.exc'(X7,X6),X8),'const.infer.st_ex_return_2'(s(X7,X14),s(X8,X15))) = s('type.pair.prod'('type.infer.exc'(X7,X6),X8),'const.pair.,_2'(s('type.infer.exc'(X7,X6),'const.infer.Success_1'(s(X7,X11))),s(X8,X16))) <=> (s(X7,X11) = s(X7,X14) & s(X8,X15) = s(X8,X16)))),
  file('thm.inferProps.success_eqns|split|0','thm.inferProps.inferProps_24_prove_ltb')).
fof(f12371,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK4))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f243,f239])).
fof(f12918,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f242,f12371,f209])).
fof(f12370,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f239,f11775])).
fof(f11775,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,X1),sK2),'const.infer.st_ex_return_2'(s(X0,X2),s(sK2,sK4))) != s('type.pair.prod'('type.infer.exc'(X0,X1),sK2),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X3))),s(sK2,sK6)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f210])).
fof(f210,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | s(X2,X4) = s(X2,X6)) )),
  inference(cnf_transformation,[],[f141])).
fof(f12357,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f11782,f244])).
fof(f11782,plain,(
  s(sK2,sK4) = s(sK2,sK6) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f11771,f210])).
fof(f11771,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK3))),s(sK2,sK6))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f239,f241])).
fof(f11783,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(unit_resulting_resolution,[],[f244,f11771,f210])).
fof(f11128,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f253,f5578])).
fof(f5578,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,X0),'const.infer.Success_1'(s(sK1,sK3))),s(sK2,sK4)))) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f5577,f241])).
fof(f5577,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK5),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,X0),'const.infer.Success_1'(s(sK1,sK3))),s(sK2,sK4)))) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f5494,f243])).
fof(f5494,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK5),s(sK2,sK6))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,X0),'const.infer.Success_1'(s(sK1,sK3))),s(sK2,sK4)))) ) | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f241,f243,f211])).
fof(f211,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f253,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) != s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK3))),s(sK2,sK4))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f252,f241])).
fof(f252,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) != s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK4))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f240,f243])).
fof(f11396,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f11357])).
fof(f11357,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) != s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) | ($spl1 | $spl2 | $spl4)),
  inference(superposition,[],[f253,f5578])).
fof(f8920,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f8807,f162])).
fof(f162,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.inferProps.success_eqns|split|0',aHL_TRUTH)).
fof(f8807,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(backward_demodulation,[],[f8770,f163])).
fof(f163,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.inferProps.success_eqns|split|0',aHL_FALSITY)).
fof(f8777,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) ) | $spl42),
  inference(cnf_transformation,[],[f8777_D])).
fof(f8777_D,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f8770,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl40),
  inference(cnf_transformation,[],[f8770_D])).
fof(f8770_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f8767,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) )),
  inference(condensation,[],[f8766])).
fof(f8766,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(X40,X41) != s(X40,X42) | s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.infer.guard_3'(s(bool,'T_0'),s(X43,X44),s(X40,X41))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X43),'const.infer.Success_1'(s('type.one.one',X45))),s(X40,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f8712,f355])).
fof(f355,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f162,f268])).
fof(f268,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f166,f166])).
fof(f166,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.inferProps.success_eqns|split|0',aHL_BOOL_CASES)).
fof(f8712,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (~p(s(bool,X39)) | s(X40,X41) != s(X40,X42) | s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.infer.guard_3'(s(bool,'T_0'),s(X43,X44),s(X40,X41))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X43),'const.infer.Success_1'(s('type.one.one',X45))),s(X40,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f247,f268])).
fof(f247,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X2)) | s(X1,X4) != s(X1,X6) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))) )),
  inference(subsumption_resolution,[],[f197,f164])).
fof(f164,plain,(
  ( ! [X0] : (s('type.one.one','const.one.one_0') = s('type.one.one',X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : s('type.one.one','const.one.one_0') = s('type.one.one',X0)),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X14] : s('type.one.one','const.one.one_0') = s('type.one.one',X14)),
  file('thm.inferProps.success_eqns|split|0','thm.one.one')).
fof(f197,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | s(X1,X4) != s(X1,X6) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & (~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & ((~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7,X8] : ! [X35,X31,X15,X14,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X7),X8),'const.infer.guard_3'(s(bool,X35),s(X7,X31),s(X8,X15))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X7),X8),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X7),'const.infer.Success_1'(s('type.one.one',X14))),s(X8,X16))) <=> (p(s(bool,X35)) & s('type.one.one','const.one.one_0') = s('type.one.one',X14) & s(X8,X15) = s(X8,X16)))),
  file('thm.inferProps.success_eqns|split|0','thm.inferProps.inferProps_38_prove_ltb')).
fof(f8772,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) )),
  inference(condensation,[],[f8771])).
fof(f8771,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (s(X48,X49) != s(X48,X50) | s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.infer.guard_3'(s(bool,'T_0'),s(X51,X52),s(X48,X49))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X51),'const.infer.Success_1'(s('type.one.one',X53))),s(X48,X50))) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f8713,f355])).
fof(f8713,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (~p(s(bool,X47)) | s(X48,X49) != s(X48,X50) | s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.infer.guard_3'(s(bool,'T_0'),s(X51,X52),s(X48,X49))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X51),'const.infer.Success_1'(s('type.one.one',X53))),s(X48,X50))) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f247,f268])).
fof(f8769,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) | $spl38),
  inference(cnf_transformation,[],[f8769_D])).
fof(f8769_D,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f8773,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3)) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8772,f3633])).
fof(f3633,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,'T_0'),s(X29,X31),s(X30,X32))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34)))) ) | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f3548,f3550])).
fof(f8768,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3)) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8767,f3633])).
fof(f8764,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8661,f3633])).
fof(f8661,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(X1,happ(s(fun('type.infer.infer_st'(X1),X1),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X1),X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f167,f247])).
fof(f167,plain,(
  ( ! [X0,X1] : (s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7] : ! [X38] : s(X7,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X7),X38))) = s(X7,happ(s(fun('type.infer.infer_st'(X7),X7),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X7),X38)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_0_1_const.infer.infer_st_subst')).
fof(f8763,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8662,f3633])).
fof(f8662,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,X7))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(X1,happ(s(fun(X5,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X5,X3),X6))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f162,f205,f247])).
fof(f205,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X40,X41,X42] : ! [X44,X45,X47] : s(X40,happ(s(fun(X41,X40),X44),s(X41,happ(s(fun(X42,X41),X45),s(X42,X47))))) = s(X40,happ(s(fun(X42,X40),'_dst_x0x1_2'(s(fun(X41,X40),X44),s(fun(X42,X41),X45))),s(X42,X47)))),
  file('thm.inferProps.success_eqns|split|0','_dst_x0x1_ax')).
fof(f8762,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8663,f3633])).
fof(f8663,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,'type.one.one'),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,'type.one.one'),X5)))) )),
  inference(unit_resulting_resolution,[],[f162,f557,f247])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,'type.one.one'),X1) = s(fun(X0,'type.one.one'),X2)) )),
  inference(unit_resulting_resolution,[],[f254,f181])).
fof(f181,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.inferProps.success_eqns|split|0',aHL_EXT)).
fof(f254,plain,(
  ( ! [X0,X1] : (s('type.one.one',X0) = s('type.one.one',X1)) )),
  inference(superposition,[],[f164,f164])).
fof(f8761,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8664,f3633])).
fof(f8664,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,'type.one.one')),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,'type.one.one')),X6)))) )),
  inference(unit_resulting_resolution,[],[f162,f672,f247])).
fof(f672,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X1,'type.one.one')),X2) = s(fun(X0,fun(X1,'type.one.one')),X3)) )),
  inference(unit_resulting_resolution,[],[f557,f181])).
fof(f8760,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8665,f3633])).
fof(f8665,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X7)))) )),
  inference(unit_resulting_resolution,[],[f162,f700,f247])).
fof(f700,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,fun(X2,'type.one.one'))),X3) = s(fun(X0,fun(X1,fun(X2,'type.one.one'))),X4)) )),
  inference(unit_resulting_resolution,[],[f672,f181])).
fof(f8759,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8666,f3633])).
fof(f8666,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X8)))) )),
  inference(unit_resulting_resolution,[],[f162,f735,f247])).
fof(f735,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,'type.one.one')))),X4) = s(fun(X0,fun(X1,fun(X2,fun(X3,'type.one.one')))),X5)) )),
  inference(unit_resulting_resolution,[],[f700,f181])).
fof(f8758,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8667,f3633])).
fof(f8667,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X7))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X9)))) )),
  inference(unit_resulting_resolution,[],[f162,f763,f247])).
fof(f763,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),X5) = s(fun(X0,fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),X6)) )),
  inference(unit_resulting_resolution,[],[f735,f181])).
fof(f8757,plain,(
  $false | ($spl10 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8668,f3633])).
fof(f8668,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',bool),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',bool),X4)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f162,f991,f247])).
fof(f8756,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8669,f3633])).
fof(f8669,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X2),'const.infer.Success_1'(s(X1,X4))),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f176,f247])).
fof(f176,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X7] : ! [X5] : s('type.pair.prod'('type.infer.exc'(X7,X8),X7),'const.infer.read_1'(s(X7,X5))) = s('type.pair.prod'('type.infer.exc'(X7,X8),X7),'const.pair.,_2'(s('type.infer.exc'(X7,X8),'const.infer.Success_1'(s(X7,X5))),s(X7,X5)))),
  file('thm.inferProps.success_eqns|split|0','thm.infer.read_def')).
fof(f8755,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8670,f3633])).
fof(f8670,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.read_0'),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f177,f247])).
fof(f177,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X8,X7] : ! [X38] : s('type.pair.prod'('type.infer.exc'(X7,X8),X7),'const.infer.read_1'(s(X7,X38))) = s('type.pair.prod'('type.infer.exc'(X7,X8),X7),happ(s(fun(X7,'type.pair.prod'('type.infer.exc'(X7,X8),X7)),'const.infer.read_0'),s(X7,X38)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_0_1_const.infer.read')).
fof(f8754,plain,(
  $false | ($spl4 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8671,f3633])).
fof(f8671,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK2,X1),sK2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.pair.prod'('type.infer.exc'(sK2,X1),sK2),'const.infer.read_1'(s(sK2,sK4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK2,X1),sK2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(sK2,X1),sK2),'const.infer.st_ex_return_2'(s(sK2,sK4),s(sK2,sK4)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f162,f5576,f247])).
fof(f5576,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK2,X0),sK2),'const.infer.read_1'(s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK2,X0),sK2),'const.infer.st_ex_return_2'(s(sK2,sK4),s(sK2,sK4)))) ) | $spl4),
  inference(forward_demodulation,[],[f5575,f243])).
fof(f5575,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK2,X0),sK2),'const.infer.read_1'(s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK2,X0),sK2),'const.infer.st_ex_return_2'(s(sK2,sK6),s(sK2,sK6)))) ) | $spl4),
  inference(forward_demodulation,[],[f5495,f176])).
fof(f5495,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK2,X0),sK2),'const.infer.st_ex_return_2'(s(sK2,sK6),s(sK2,sK6))) = s('type.pair.prod'('type.infer.exc'(sK2,X0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK2,X0),'const.infer.Success_1'(s(sK2,sK4))),s(sK2,sK4)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f243,f243,f211])).
fof(f8753,plain,(
  $false | ($spl2 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8672,f3633])).
fof(f8672,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK1,X1),sK1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.pair.prod'('type.infer.exc'(sK1,X1),sK1),'const.infer.read_1'(s(sK1,sK3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK1,X1),sK1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(sK1,X1),sK1),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK1,sK3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f162,f5654,f247])).
fof(f5654,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.read_1'(s(sK1,sK3))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK1,sK3)))) ) | $spl2),
  inference(forward_demodulation,[],[f5653,f241])).
fof(f5653,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.read_1'(s(sK1,sK3))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.st_ex_return_2'(s(sK1,sK5),s(sK1,sK5)))) ) | $spl2),
  inference(forward_demodulation,[],[f5452,f176])).
fof(f5452,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.infer.st_ex_return_2'(s(sK1,sK5),s(sK1,sK5))) = s('type.pair.prod'('type.infer.exc'(sK1,X0),sK1),'const.pair.,_2'(s('type.infer.exc'(sK1,X0),'const.infer.Success_1'(s(sK1,sK3))),s(sK1,sK3)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f241,f241,f211])).
fof(f8752,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8673,f3633])).
fof(f8673,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X2),X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),happ(s(fun('type.infer.infer_st'(X2),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f178,f247])).
fof(f178,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X7,X8] : ! [X38] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X7),'type.infer.infer_st'(X8)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X8),X38))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X7),'type.infer.infer_st'(X8)),happ(s(fun('type.infer.infer_st'(X8),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X7),'type.infer.infer_st'(X8))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X8),X38)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_0_1_const.infer.fresh_uvar')).
fof(f8751,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8674,f3633])).
fof(f8674,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.failwith_2'(s(X2,X5),s(X3,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.failwith_1'(s(X2,X5))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f162,f203,f247])).
fof(f203,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6,X7,X8] : ! [X38,X39] : s('type.pair.prod'('type.infer.exc'(X6,X7),X8),'const.infer.failwith_2'(s(X7,X38),s(X8,X39))) = s('type.pair.prod'('type.infer.exc'(X6,X7),X8),happ(s(fun(X8,'type.pair.prod'('type.infer.exc'(X6,X7),X8)),'const.infer.failwith_1'(s(X7,X38))),s(X8,X39)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_1_2_const.infer.failwith')).
fof(f8750,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8675,f3633])).
fof(f8675,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X4),s(X1,X5),s(X2,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_2'(s(bool,X4),s(X1,X5))),s(X2,X6)))))) )),
  inference(unit_resulting_resolution,[],[f162,f182,f247])).
fof(f182,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X8,X7] : ! [X38,X39,X48] : s('type.pair.prod'('type.infer.exc'('type.one.one',X8),X7),'const.infer.guard_3'(s(bool,X38),s(X8,X39),s(X7,X48))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X8),X7),happ(s(fun(X7,'type.pair.prod'('type.infer.exc'('type.one.one',X8),X7)),'const.infer.guard_2'(s(bool,X38),s(X8,X39))),s(X7,X48)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_2_3_const.infer.guard')).
fof(f8749,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8676,f3633])).
fof(f8676,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X2))),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f174,f247])).
fof(f174,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X38,X39] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X39))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X38))),s('type.num.num',X39)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_1_2_const.arithmetic.+')).
fof(f8748,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8677,f3633])).
fof(f8677,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X2))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f162,f165,f247])).
fof(f165,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X38] : s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X38))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X38)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_0_1_const.infer_t.Infer_Tuvar')).
fof(f8747,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8678,f3633])).
fof(f8678,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s('type.infer_t.infer_t',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f175,f247])).
fof(f175,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X38,X39] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X38),s('type.infer_t.infer_t',X39))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X38))),s('type.infer_t.infer_t',X39)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_0_2_const.unify.t_walkstar')).
fof(f8746,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8679,f3633])).
fof(f8679,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X3,X1),X4),s('type.list.list'(X3),X5))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s('type.list.list'(X1),happ(s(fun('type.list.list'(X3),'type.list.list'(X1)),happ(s(fun(fun(X3,X1),fun('type.list.list'(X3),'type.list.list'(X1))),'const.list.MAP_0'),s(fun(X3,X1),X4))),s('type.list.list'(X3),X5)))))) )),
  inference(unit_resulting_resolution,[],[f162,f179,f247])).
fof(f179,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X8,X7] : ! [X38,X39] : s('type.list.list'(X8),'const.list.MAP_2'(s(fun(X7,X8),X38),s('type.list.list'(X7),X39))) = s('type.list.list'(X8),happ(s(fun('type.list.list'(X7),'type.list.list'(X8)),happ(s(fun(fun(X7,X8),fun('type.list.list'(X7),'type.list.list'(X8))),'const.list.MAP_0'),s(fun(X7,X8),X38))),s('type.list.list'(X7),X39)))),
  file('thm.inferProps.success_eqns|split|0','ARITY_0_2_const.list.MAP')).
fof(f8745,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8680,f3633])).
fof(f8680,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one','const.one.one_0'))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.one.one',X3)))) )),
  inference(unit_resulting_resolution,[],[f162,f164,f247])).
fof(f8744,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8681,f3633])).
fof(f8681,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one',X4)))) )),
  inference(unit_resulting_resolution,[],[f162,f254,f247])).
fof(f8743,plain,(
  $false | ($spl2 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8682,f3633])).
fof(f8682,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(sK1,sK3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(sK1,sK5)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f162,f241,f247])).
fof(f8742,plain,(
  $false | ($spl4 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8683,f3633])).
fof(f8683,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK2),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK2),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(sK2,sK6)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f162,f243,f247])).
fof(f8741,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8684,f3633])).
fof(f8684,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(X1,happ(s(fun('type.infer.infer_st'(X1),X1),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X1),X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f167,f247])).
fof(f8740,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8685,f3633])).
fof(f8685,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(X1,happ(s(fun(X3,X1),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,X7))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X8),s(X1,happ(s(fun(X5,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X5,X3),X6))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f162,f205,f247])).
fof(f8739,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8686,f3633])).
fof(f8686,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,'type.one.one'),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,'type.one.one'),X5)))) )),
  inference(unit_resulting_resolution,[],[f162,f557,f247])).
fof(f8738,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8687,f3633])).
fof(f8687,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,'type.one.one')),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,'type.one.one')),X6)))) )),
  inference(unit_resulting_resolution,[],[f162,f672,f247])).
fof(f8737,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8688,f3633])).
fof(f8688,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X7)))) )),
  inference(unit_resulting_resolution,[],[f162,f700,f247])).
fof(f8736,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8689,f3633])).
fof(f8689,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X8)))) )),
  inference(unit_resulting_resolution,[],[f162,f735,f247])).
fof(f8735,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8690,f3633])).
fof(f8690,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X7))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X9)))) )),
  inference(unit_resulting_resolution,[],[f162,f763,f247])).
fof(f8734,plain,(
  $false | ($spl10 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8691,f3633])).
fof(f8691,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',bool),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',bool),X4)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f162,f991,f247])).
fof(f8733,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8692,f3633])).
fof(f8692,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X2),'const.infer.Success_1'(s(X1,X4))),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f176,f247])).
fof(f8732,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8693,f3633])).
fof(f8693,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.read_0'),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f177,f247])).
fof(f8731,plain,(
  $false | ($spl4 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8694,f3633])).
fof(f8694,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK2,X1),sK2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.pair.prod'('type.infer.exc'(sK2,X1),sK2),'const.infer.read_1'(s(sK2,sK4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK2,X1),sK2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(sK2,X1),sK2),'const.infer.st_ex_return_2'(s(sK2,sK4),s(sK2,sK4)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f162,f5576,f247])).
fof(f8730,plain,(
  $false | ($spl2 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8695,f3633])).
fof(f8695,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK1,X1),sK1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.pair.prod'('type.infer.exc'(sK1,X1),sK1),'const.infer.read_1'(s(sK1,sK3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK1,X1),sK1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(sK1,X1),sK1),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK1,sK3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f162,f5654,f247])).
fof(f8729,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8696,f3633])).
fof(f8696,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X2),X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),happ(s(fun('type.infer.infer_st'(X2),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f178,f247])).
fof(f8728,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8697,f3633])).
fof(f8697,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.failwith_2'(s(X2,X5),s(X3,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X7),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.failwith_1'(s(X2,X5))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f162,f203,f247])).
fof(f8727,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8698,f3633])).
fof(f8698,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X4),s(X1,X5),s(X2,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X7),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_2'(s(bool,X4),s(X1,X5))),s(X2,X6)))))) )),
  inference(unit_resulting_resolution,[],[f162,f182,f247])).
fof(f8726,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8699,f3633])).
fof(f8699,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X2))),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f174,f247])).
fof(f8725,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8700,f3633])).
fof(f8700,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X2))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f162,f165,f247])).
fof(f8724,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8701,f3633])).
fof(f8701,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s('type.infer_t.infer_t',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f175,f247])).
fof(f8723,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8702,f3633])).
fof(f8702,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X3,X1),X4),s('type.list.list'(X3),X5))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s('type.list.list'(X1),happ(s(fun('type.list.list'(X3),'type.list.list'(X1)),happ(s(fun(fun(X3,X1),fun('type.list.list'(X3),'type.list.list'(X1))),'const.list.MAP_0'),s(fun(X3,X1),X4))),s('type.list.list'(X3),X5)))))) )),
  inference(unit_resulting_resolution,[],[f162,f179,f247])).
fof(f8722,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8703,f3633])).
fof(f8703,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one','const.one.one_0')))) )),
  inference(unit_resulting_resolution,[],[f162,f164,f247])).
fof(f8721,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8704,f3633])).
fof(f8704,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one',X4)))) )),
  inference(unit_resulting_resolution,[],[f162,f254,f247])).
fof(f8720,plain,(
  $false | ($spl2 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8705,f3633])).
fof(f8705,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(sK1,sK5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(sK1,sK3)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f162,f241,f247])).
fof(f8719,plain,(
  $false | ($spl4 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f8706,f3633])).
fof(f8706,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK2),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(sK2,sK6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),sK2),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(sK2,sK4)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f162,f243,f247])).
fof(f3620,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) | $spl34),
  inference(cnf_transformation,[],[f3620_D])).
fof(f3620_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3622,plain,(
  s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | $spl37),
  inference(cnf_transformation,[],[f3622_D])).
fof(f3622_D,plain,(
  s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f3607,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X5)))) ) | $spl16),
  inference(superposition,[],[f994,f164])).
fof(f3619,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0)))) ) | $spl32),
  inference(cnf_transformation,[],[f3619_D])).
fof(f3619_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3606,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X4))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X3)))) ) | $spl16),
  inference(superposition,[],[f994,f254])).
fof(f3605,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X2))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) | $spl16),
  inference(superposition,[],[f994,f254])).
fof(f3618,plain,(
  s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0'))) | $spl30),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0'))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3604,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl16),
  inference(superposition,[],[f994,f164])).
fof(f3547,plain,(
  ( ! [X26,X29,X25] : (s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl18),
  inference(cnf_transformation,[],[f3547_D])).
fof(f3547_D,plain,(
  ( ! [X26,X29,X25] : (s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3551,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))))) ) | $spl26),
  inference(cnf_transformation,[],[f3551_D])).
fof(f3551_D,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3509,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,happ(s(fun('type.one.one',bool),X71),s('type.one.one',X72)))) | s(bool,X73) = s(bool,happ(s(fun('type.one.one',bool),X69),s('type.one.one',X70))) | s(bool,'T_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f355,f3295])).
fof(f3295,plain,(
  ( ! [X269,X271,X268,X270] : (s(bool,happ(s(fun('type.one.one',bool),X270),s('type.one.one',X268))) = s(bool,happ(s(fun('type.one.one',bool),X271),s('type.one.one',X269)))) ) | $spl10),
  inference(superposition,[],[f1938,f254])).
fof(f1938,plain,(
  ( ! [X125,X123,X126,X124,X122] : (s(bool,happ(s(fun('type.one.one',bool),X125),s('type.one.one',X126))) = s(bool,happ(s(fun(X122,bool),X123),s(X122,happ(s(fun('type.one.one',X122),X124),s('type.one.one',X126)))))) ) | $spl10),
  inference(superposition,[],[f205,f991])).
fof(f3548,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) ) | $spl20),
  inference(cnf_transformation,[],[f3548_D])).
fof(f3548_D,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3552,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68))))) ) | $spl28),
  inference(cnf_transformation,[],[f3552_D])).
fof(f3552_D,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3508,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X65),s('type.one.one',X66)))) ) | $spl10),
  inference(superposition,[],[f270,f3295])).
fof(f270,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f163,f166])).
fof(f3507,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X62),s('type.one.one',X63))) | s(bool,X64) = s(bool,happ(s(fun('type.one.one',bool),X60),s('type.one.one',X61))) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(superposition,[],[f268,f3295])).
fof(f3506,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X59) = s(bool,happ(s(fun('type.one.one',bool),X57),s('type.one.one',X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X55),s('type.one.one',X56))) | s(bool,'T_0') = s(bool,X59)) ) | $spl10),
  inference(superposition,[],[f268,f3295])).
fof(f3505,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun('type.one.one',bool),X52),s('type.one.one',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X50),s('type.one.one',X51)))) ) | $spl10),
  inference(superposition,[],[f268,f3295])).
fof(f3504,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,X49) = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46))) | s(bool,'T_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f268,f3295])).
fof(f3550,plain,(
  ( ! [X30,X26,X33,X31,X29,X25,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl24),
  inference(cnf_transformation,[],[f3550_D])).
fof(f3550_D,plain,(
  ( ! [X30,X26,X33,X31,X29,X25,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3502,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X34,X32] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) | s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl10),
  inference(superposition,[],[f194,f3295])).
fof(f194,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f3549,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48)))) ) | $spl22),
  inference(cnf_transformation,[],[f3549_D])).
fof(f3549_D,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3500,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X15),s('type.one.one',X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14)))) ) | $spl10),
  inference(superposition,[],[f166,f3295])).
fof(f3460,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,happ(s(fun('type.one.one',bool),X71),s('type.one.one',X72)))) | s(bool,X73) = s(bool,happ(s(fun('type.one.one',bool),X69),s('type.one.one',X70))) | s(bool,'T_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f355,f3295])).
fof(f3459,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X65),s('type.one.one',X66)))) ) | $spl10),
  inference(superposition,[],[f270,f3295])).
fof(f3458,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X62),s('type.one.one',X63))) | s(bool,X64) = s(bool,happ(s(fun('type.one.one',bool),X60),s('type.one.one',X61))) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(superposition,[],[f268,f3295])).
fof(f3457,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X59) = s(bool,happ(s(fun('type.one.one',bool),X57),s('type.one.one',X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X55),s('type.one.one',X56))) | s(bool,'T_0') = s(bool,X59)) ) | $spl10),
  inference(superposition,[],[f268,f3295])).
fof(f3456,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun('type.one.one',bool),X52),s('type.one.one',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X50),s('type.one.one',X51)))) ) | $spl10),
  inference(superposition,[],[f268,f3295])).
fof(f3455,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,X49) = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46))) | s(bool,'T_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f268,f3295])).
fof(f3453,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X34,X32] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) | s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl10),
  inference(superposition,[],[f194,f3295])).
fof(f3451,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X15),s('type.one.one',X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14)))) ) | $spl10),
  inference(superposition,[],[f166,f3295])).
fof(f3447,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46)))) ) | $spl10),
  inference(superposition,[],[f3295,f166])).
fof(f3446,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.one.one',bool),X43),s('type.one.one',X44))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X40),s('type.one.one',X41)))) ) | $spl10),
  inference(superposition,[],[f3295,f268])).
fof(f3445,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.one.one',bool),X38),s('type.one.one',X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X35),s('type.one.one',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl10),
  inference(superposition,[],[f3295,f268])).
fof(f3444,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X32),s('type.one.one',X33))) | s(bool,X34) = s(bool,happ(s(fun('type.one.one',bool),X30),s('type.one.one',X31))) | s(bool,'T_0') = s(bool,X34)) ) | $spl10),
  inference(superposition,[],[f3295,f268])).
fof(f3443,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))) | s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f3295,f268])).
fof(f3425,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46)))) ) | $spl10),
  inference(superposition,[],[f3295,f166])).
fof(f3424,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.one.one',bool),X43),s('type.one.one',X44))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X40),s('type.one.one',X41)))) ) | $spl10),
  inference(superposition,[],[f3295,f268])).
fof(f3423,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.one.one',bool),X38),s('type.one.one',X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X35),s('type.one.one',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl10),
  inference(superposition,[],[f3295,f268])).
fof(f3422,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X32),s('type.one.one',X33))) | s(bool,X34) = s(bool,happ(s(fun('type.one.one',bool),X30),s('type.one.one',X31))) | s(bool,'T_0') = s(bool,X34)) ) | $spl10),
  inference(superposition,[],[f3295,f268])).
fof(f3421,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))) | s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f3295,f268])).
fof(f991,plain,(
  ( ! [X2,X1] : (s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2)) ) | $spl10),
  inference(cnf_transformation,[],[f991_D])).
fof(f991_D,plain,(
  ( ! [X2,X1] : (s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f994,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) ) | $spl16),
  inference(cnf_transformation,[],[f994_D])).
fof(f994_D,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f979,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(fun('type.one.one',bool),X17) = s(fun('type.one.one',bool),X18) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) )),
  inference(superposition,[],[f563,f166])).
fof(f563,plain,(
  ( ! [X21,X19,X22,X20] : (s(X21,f(s('type.one.one',X22))) != s(X21,g(s('type.one.one',X22))) | s(fun('type.one.one',X21),X19) = s(fun('type.one.one',X21),X20)) )),
  inference(superposition,[],[f181,f254])).
fof(f993,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) | $spl14),
  inference(cnf_transformation,[],[f993_D])).
fof(f993_D,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f978,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X13) != s(bool,f(s('type.one.one',X12))) | s(fun('type.one.one',bool),X14) = s(fun('type.one.one',bool),X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,g(s('type.one.one',X12)))) )),
  inference(superposition,[],[f563,f268])).
fof(f977,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(fun('type.one.one',bool),X10) = s(fun('type.one.one',bool),X11) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f563,f268])).
fof(f992,plain,(
  ( ! [X0,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,g(s('type.one.one',X0)))) ) | $spl12),
  inference(cnf_transformation,[],[f992_D])).
fof(f992_D,plain,(
  ( ! [X0,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,g(s('type.one.one',X0)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f976,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X4))) | s(fun('type.one.one',bool),X5) = s(fun('type.one.one',bool),X6) | s(bool,X7) = s(bool,g(s('type.one.one',X4))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f563,f268])).
fof(f975,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2) | s(bool,X3) = s(bool,g(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f563,f268])).
fof(f274,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f274_D])).
fof(f274_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f276,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f276_D])).
fof(f276_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f273,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f166])).
fof(f239,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6))) | $spl0),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f241,plain,(
  s(sK1,sK3) = s(sK1,sK5) | $spl2),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  s(sK1,sK3) = s(sK1,sK5) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f159,plain,(
  s(sK1,sK3) = s(sK1,sK5) | s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6)))),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  (s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6))) | (s(sK1,sK3) = s(sK1,sK5) & s(sK2,sK4) = s(sK2,sK6))) & (s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) != s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6))) | s(sK1,sK3) != s(sK1,sK5) | s(sK2,sK4) != s(sK2,sK6))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f124])).
fof(f124,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6)))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <~> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11,X12] : (s('type.pair.prod'('type.infer.exc'(X7,X6),X8),'const.infer.st_ex_return_2'(s(X7,X9),s(X8,X10))) = s('type.pair.prod'('type.infer.exc'(X7,X6),X8),'const.pair.,_2'(s('type.infer.exc'(X7,X6),'const.infer.Success_1'(s(X7,X11))),s(X8,X12))) <=> (s(X7,X9) = s(X7,X11) & s(X8,X10) = s(X8,X12)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X12] : (s('type.pair.prod'('type.infer.exc'(X7,X6),X8),'const.infer.st_ex_return_2'(s(X7,X9),s(X8,X10))) = s('type.pair.prod'('type.infer.exc'(X7,X6),X8),'const.pair.,_2'(s('type.infer.exc'(X7,X6),'const.infer.Success_1'(s(X7,X11))),s(X8,X12))) <=> (s(X7,X9) = s(X7,X11) & s(X8,X10) = s(X8,X12)))),
  file('thm.inferProps.success_eqns|split|0','thm.inferProps.success_eqns|split|0')).
fof(f243,plain,(
  s(sK2,sK4) = s(sK2,sK6) | $spl4),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  s(sK2,sK4) = s(sK2,sK6) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f160,plain,(
  s(sK2,sK4) = s(sK2,sK6) | s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) = s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6)))),
  inference(cnf_transformation,[],[f125])).
fof(f240,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) != s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6))) | $spl1),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) != s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f242,plain,(
  s(sK1,sK3) != s(sK1,sK5) | $spl3),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  s(sK1,sK3) != s(sK1,sK5) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f244,plain,(
  s(sK2,sK4) != s(sK2,sK6) | $spl5),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  s(sK2,sK4) != s(sK2,sK6) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f161,plain,(
  s(sK2,sK4) != s(sK2,sK6) | s(sK1,sK3) != s(sK1,sK5) | s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.infer.st_ex_return_2'(s(sK1,sK3),s(sK2,sK4))) != s('type.pair.prod'('type.infer.exc'(sK1,sK0),sK2),'const.pair.,_2'(s('type.infer.exc'(sK1,sK0),'const.infer.Success_1'(s(sK1,sK5))),s(sK2,sK6)))),
  inference(cnf_transformation,[],[f125])).
% SZS output end Proof for thm.inferProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 77
% Passive clauses: 1306
% Generated clauses: 13456
% Final active clauses: 49
% Final passive clauses: 885
% Input formulas: 44
% Initial clauses: 80
% 
% Duplicate literals: 53
% Trivial inequalities: 8
% Fw subsumption resolutions: 75
% Fw demodulations: 5018
% Bw demodulations: 120
% Condensations: 162
% 
% Simple tautologies: 18
% Equational tautologies: 33
% Forward subsumptions: 6075
% Fw demodulations to eq. taut.: 611
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 4025
% Forward superposition: 2059
% Backward superposition: 1518
% Self superposition: 129
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 43
% Split components: 87
% SAT solver clauses: 1448
% SAT solver unit clauses: 419
% SAT solver binary clauses: 618
% 
% Sat splits: 43
% Sat splitting refutations: 54
% 
% Memory used [KB]: 14456
% Time elapsed: 0.503 s
% ------------------------------
% ------------------------------
% Success in time 1.83 s
