Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.clos_numberProof
% (105321)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 116
% Passive clauses: 3446
% Generated clauses: 20212
% Final active clauses: 111
% Final passive clauses: 2978
% Input formulas: 28
% Initial clauses: 62
% 
% Duplicate literals: 1027
% Trivial inequalities: 430
% Fw subsumption resolutions: 3102
% Fw demodulations: 1964
% Bw demodulations: 3
% 
% Simple tautologies: 450
% Equational tautologies: 72
% Forward subsumptions: 8139
% Fw demodulations to eq. taut.: 13
% 
% Binary resolution: 178
% Factoring: 2
% Forward superposition: 6843
% Backward superposition: 6159
% Self superposition: 295
% Equality factoring: 208
% Equality resolution: 41
% 
% Split clauses: 749
% Split components: 1513
% SAT solver clauses: 1782
% SAT solver unit clauses: 13
% SAT solver binary clauses: 197
% 
% Sat splits: 749
% 
% TWLsolver clauses: 1738
% TWLsolver calls for satisfiability: 185
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_numberProof
% (105350)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 116
% Passive clauses: 1116
% Generated clauses: 6264
% Final active clauses: 107
% Final passive clauses: 934
% Input formulas: 28
% Initial clauses: 62
% 
% Duplicate literals: 186
% Trivial inequalities: 199
% Fw subsumption resolutions: 252
% Fw demodulations: 586
% Bw demodulations: 69
% Condensations: 273
% Global subsumptions: 9
% 
% Simple tautologies: 74
% Equational tautologies: 32
% Forward subsumptions: 1737
% Fw demodulations to eq. taut.: 7
% 
% Binary resolution: 191
% Factoring: 3
% Forward superposition: 2115
% Backward superposition: 1007
% Self superposition: 70
% Equality factoring: 54
% Equality resolution: 54
% 
% Split clauses: 31
% Split components: 76
% SAT solver clauses: 1761
% SAT solver unit clauses: 62
% SAT solver binary clauses: 48
% SAT solver learnt clauses: 18
% 
% Sat splits: 31
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1725
% TWLsolver calls for satisfiability: 309
% 
% Memory used [KB]: 8955
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_numberProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_numberProof
% SZS output start Proof for thm.clos_numberProof
fof(f29073,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f110,f172,f109,f176,f174,f171,f108,f107,f178,f106,f175,f105,f180,f170,f183,f186,f197,f200,f218,f222,f220,f776,f778,f860,f863,f866,f1472,f1473,f1467,f1460,f1479,f1481,f1480,f1482,f1485,f1486,f1474,f1610,f1611,f1612,f2131,f2154,f2466,f2476,f2656,f2659,f2718,f2950,f2721,f2954,f2949,f3539,f3533,f3541,f3525,f3519,f3547,f3518,f3548,f3517,f3549,f3515,f3552,f3554,f3553,f3555,f3558,f3560,f3564,f3550,f3567,f3568,f3540,f3570,f3325,f3571,f3572,f3575,f3327,f3576,f3577,f3290,f3102,f3581,f3579,f3585,f3091,f3586,f3661,f3663,f3664,f3836,f3909,f2263,f2260,f3917,f3938,f3939,f4171,f4194,f4195,f4198,f4242,f4244,f4247,f4248,f4250,f5192,f5193,f6302,f6313,f6314,f6312,f7427,f7452,f7662,f7680,f8224,f8206,f8241,f9607,f9645,f10740,f10759,f11259,f11253,f11245,f11300,f11307,f11313,f11314,f11239,f11350,f11353,f11308,f11226,f11398,f11279,f11648,f11663,f11649,f11650,f11651,f11652,f11653,f11654,f11670,f11669,f11655,f11666,f11671,f11656,f11667,f11673,f11657,f11674,f11661,f11660,f12471,f12473,f12475,f12484,f12504,f12505,f12530,f12542,f12569,f12570,f12571,f12572,f12578,f12579,f12580,f12581,f12582,f12591,f12594,f12597,f12598,f12601,f14585,f14588,f14863,f14861,f14749,f14979,f11664,f16202,f16204,f16205,f16203,f17310,f17313,f17676,f17318,f17681,f17381,f17712,f17392,f17738,f17739,f17740,f17719,f17741,f17747,f17718,f17713,f17772,f17775,f17501,f17786,f17513,f17792,f17527,f17803,f17532,f17781,f17849,f17943,f17944,f17946,f17947,f17949,f17956,f18061,f18077,f4243,f18079,f18080,f4249,f18081,f18083,f198,f4245,f18414,f18415,f18418,f18419,f18449,f3578,f18451,f18452,f2475,f18455,f18457,f18456,f18493,f18695,f18670,f18696,f18748,f18749,f18750,f18751,f18752,f21099,f21102,f17678,f21367,f21455,f21457,f17773,f21259,f17787,f21480,f21601,f17782,f21838,f16206,f21839,f24998,f25024,f24999,f25953,f25956,f25954,f25964,f25965,f25967,f3569,f25981,f25968,f219,f28448,f28677,f28714,f173,f28897,f28899,f28900,f28901,f28912,f28913,f28914,f28915,f28916,f29058,f29060,f29061,f29062])).
fof(f29062,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f29035,f176])).
fof(f29035,plain,(
  s(sK0,sK2) = s(sK0,sK5) | ($spl0 | $spl4)),
  inference(superposition,[],[f163,f28917])).
fof(f28917,plain,(
  s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f173,f170])).
fof(f163,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : ! [X8,X19] : s(X7,X8) = s(X7,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.NONE_0'),s(X7,X8),s(fun(X6,X7),X19)))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','thm.option.option_case_def|split|0')).
fof(f29061,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f29033,f176])).
fof(f29033,plain,(
  s(sK0,sK2) = s(sK0,sK5) | ($spl0 | $spl4)),
  inference(superposition,[],[f28917,f163])).
fof(f29060,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f29059,f28994])).
fof(f28994,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5))) | $spl7),
  inference(unit_resulting_resolution,[],[f176,f154])).
fof(f154,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X12,X16] : (s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X16))) <=> s(X6,X12) = s(X6,X16))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','thm.option.SOME_|49||49|')).
fof(f29059,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f29031,f163])).
fof(f29031,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f28917,f155])).
fof(f155,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f100])).
fof(f29058,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f29057,f28994])).
fof(f29057,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f29032,f163])).
fof(f29032,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f28917,f155])).
fof(f28916,plain,(
  $false | ($spl2 | $spl5 | $spl188)),
  inference(subsumption_resolution,[],[f174,f28788])).
fof(f28788,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK4) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f28750,f146])).
fof(f146,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK7(X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK7(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f45])).
fof(f45,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X10] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X10) | ? [X12] : s('type.option.option'(X6),X10) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','thm.option.option_nchotomy')).
fof(f28750,plain,(
  s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK7(sK4,sK1)))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f28714,f172])).
fof(f28915,plain,(
  $false | ($spl0 | $spl2 | $spl188)),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f603,f606,f608,f611,f617,f619,f716,f739,f172,f170,f813,f929,f931,f946,f967,f1022,f1089,f1091,f1208,f1246,f1472,f1467,f1460,f1479,f1481,f1432,f1485,f1433,f1486,f2131,f2466,f3570,f3325,f3575,f3327,f3223,f3290,f3102,f3581,f3058,f3091,f2263,f113,f2260,f3078,f3920,f4196,f1200,f9311,f9374,f9837,f9838,f9839,f9840,f9841,f9775,f9776,f9777,f9778,f9780,f9781,f9817,f9827,f10090,f10100,f10102,f10103,f10104,f10105,f10144,f11648,f11649,f11650,f11651,f11652,f588,f10098,f9836,f9825,f8548,f2252,f2179,f1999,f1984,f1331,f1283,f1253,f1199,f1195,f1093,f933,f610,f239,f9474,f2212,f2205,f2191,f10152,f10121,f10108,f10107,f10069,f10067,f10054,f10053,f9460,f3242,f3186,f3099,f2214,f2206,f2189,f2183,f2178,f2175,f2167,f2161,f1998,f1995,f1983,f1980,f1308,f1201,f1176,f963,f961,f538,f537,f535,f533,f532,f497,f495,f493,f492,f3337,f9809,f9381,f2518,f2517,f2516,f2515,f2440,f2410,f2408,f2393,f2366,f2349,f2283,f2280,f1448,f1436,f1416,f1399,f1368,f1367,f1365,f1363,f1349,f1348,f1345,f1271,f1116,f928,f921,f869,f1045,f2282,f1351,f922,f2281,f1350,f900,f621,f10603,f2194,f2193,f2217,f2216,f3298,f10943,f8516,f3299,f10944,f8517,f1402,f10565,f8518,f8515,f575,f10911,f10894,f10566,f2521,f1335,f1204,f571,f10963,f10945,f10942,f10927,f10924,f10910,f10895,f10582,f10580,f10567,f3300,f3297,f2520,f1441,f1440,f1415,f1401,f1203,f898,f897,f895,f893,f892,f675,f673,f671,f576,f573,f570,f1048,f1070,f10559,f760,f788,f10558,f781,f1060,f10925,f10658,f10560,f10557,f1346,f1284,f1196,f1143,f1135,f1119,f912,f737,f735,f733,f683,f681,f585,f568,f158,f157,f156,f10926,f10620,f2479,f1347,f1285,f1197,f1062,f1061,f1032,f1020,f1018,f1016,f2478,f10124,f9797,f2477,f2359,f1409,f938,f936,f934,f2176,f1996,f1981,f3188,f2395,f2363,f2347,f1438,f1413,f1397,f2177,f1997,f1982,f10923,f10604,f10601,f10595,f10145,f10142,f10136,f10091,f10088,f10082,f10052,f9818,f9815,f9514,f9494,f9493,f9492,f9491,f9465,f9464,f9462,f9392,f9391,f9389,f9387,f9371,f9370,f9368,f9364,f9357,f9356,f9354,f9352,f3279,f3278,f3276,f3274,f3267,f3253,f3252,f3250,f3248,f3231,f3230,f3229,f3070,f3069,f3068,f3067,f2251,f2250,f2249,f2248,f965,f26859,f26860,f26861,f26874,f26876,f26904,f26905,f26906,f26907,f26909,f26910,f26927,f26997,f27046,f27202,f27058,f27060,f27886,f27887,f27888,f27814,f27815,f27828,f27924,f27925,f27860,f27861,f27862,f27863,f27937,f27880,f28336,f28688,f28338,f28339,f28340,f28353,f28700,f28701,f28359,f28702,f28361,f28703,f28447,f28448,f28451,f28453,f28455,f28456,f28462,f28463,f28464,f28465,f28466,f28467,f28469,f28470,f28471,f28472,f28474,f28476,f28481,f28483,f28484,f28485,f28486,f28487,f28488,f28489,f28490,f28491,f28492,f28544,f28545,f28546,f28547,f28548,f28549,f28550,f28552,f28553,f28554,f28559,f28563,f28567,f28569,f28571,f28572,f28578,f28579,f28580,f28581,f28582,f28583,f28589,f28591,f28596,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f28605,f28606,f28607,f28659,f28660,f28661,f28662,f28663,f28664,f28665,f28667,f28668,f28669,f28674,f28677,f28679,f28714,f28750,f28756,f28781,f28784,f28790,f28788,f28793,f28795,f28796,f28797,f28798,f28799,f28800,f28801,f28802,f28803,f28804,f28805,f28806,f28807,f28808,f28809,f28810,f28811,f28812,f28813,f28814,f28815,f28816,f28817,f28818,f28819,f28820,f28821,f28822,f28823,f28824,f28825,f28826,f28835,f28836,f28837,f28838,f28839,f28840,f28841,f28842,f28843,f28844,f28845,f28846,f28847,f28848,f28849,f28850,f28851,f28852,f28853,f28854,f28855,f28856,f28857,f28858,f28863,f28865,f28902,f28904,f28905,f28906,f28907,f28908,f28909,f28910,f28911,f28874,f28872,f28898,f28896])).
fof(f28896,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5))) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28871,f163])).
fof(f28871,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))))) | ($spl0 | $spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f28793,f155])).
fof(f28898,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5))) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28870,f163])).
fof(f28870,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))))) | ($spl0 | $spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f28793,f155])).
fof(f28872,plain,(
  s(sK0,sK2) = s(sK0,sK5) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f28793,f163])).
fof(f28874,plain,(
  s(sK0,sK2) = s(sK0,sK5) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f163,f28793])).
fof(f28911,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X35),s(sK0,sK5))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))),s(bool,X34),s(fun(sK0,bool),X35)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28894,f163])).
fof(f28894,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))),s(bool,X34),s(fun(sK0,bool),X35))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X35),s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f737,f28793])).
fof(f28910,plain,(
  ( ! [X30] : (s(sK0,sK5) = s(sK0,sK7(X30,sK0)) | s('type.option.option'(sK0),X30) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),X30)) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28892,f163])).
fof(f28892,plain,(
  ( ! [X30] : (s('type.option.option'(sK0),X30) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) | s(sK0,sK7(X30,sK0)) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),X30)) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f611,f28793])).
fof(f28909,plain,(
  ( ! [X28,X29,X27] : (s(sK0,sK5) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(X28),'const.option.NONE_0'),s(sK0,sK5),s(fun(X28,sK0),X29)),X27,sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X27),s(sK0,sK2))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28891,f163])).
fof(f28891,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun(sK0,bool),X27),s(sK0,sK2)))) | s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(X28),'const.option.NONE_0'),s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))),s(fun(X28,sK0),X29)),X27,sK0))) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f585,f28793])).
fof(f28908,plain,(
  ( ! [X26,X24,X25] : (s(sK0,sK5) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X25),'const.option.NONE_0'),s(fun(sK0,bool),X24),s(fun(X25,fun(sK0,bool)),X26)),sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X24),s(sK0,sK2))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28890,f163])).
fof(f28890,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(sK0,bool),X24),s(sK0,sK2)))) | s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X25),'const.option.NONE_0'),s(fun(sK0,bool),X24),s(fun(X25,fun(sK0,bool)),X26)),sK0))) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f568,f28793])).
fof(f28907,plain,(
  ( ! [X23,X21,X22] : (s(sK0,sK5) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(X22),'const.option.NONE_0'),s(sK0,X21),s(fun(X22,sK0),X23))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X21)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28889,f163])).
fof(f28889,plain,(
  ( ! [X23,X21,X22] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X21))) | s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(X22),'const.option.NONE_0'),s(sK0,X21),s(fun(X22,sK0),X23)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f491,f28793])).
fof(f28906,plain,(
  ( ! [X19,X20,X18] : (s(sK0,X18) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(X19),'const.option.NONE_0'),s(sK0,sK5),s(fun(X19,sK0),X20))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X18)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28888,f163])).
fof(f28888,plain,(
  ( ! [X19,X20,X18] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X18))) | s(sK0,X18) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(X19),'const.option.NONE_0'),s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))),s(fun(X19,sK0),X20)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f491,f28793])).
fof(f28905,plain,(
  ( ! [X6] : (s(sK0,sK5) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X6,sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK2))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28882,f163])).
fof(f28882,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK2)))) | s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X6,sK0))) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f157,f28793])).
fof(f28904,plain,(
  ( ! [X4,X5] : (s(sK0,sK5) != s(sK0,X5) | ~p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK2)))) | p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,X5))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28881,f163])).
fof(f28881,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK2)))) | p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,X5)))) | s(sK0,X5) != s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f156,f28793])).
fof(f28902,plain,(
  ( ! [X2] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X2))) | s(sK0,sK2) != s(sK0,X2)) ) | ($spl0 | $spl2 | $spl188)),
  inference(forward_demodulation,[],[f28879,f163])).
fof(f28879,plain,(
  ( ! [X2] : (s(sK0,sK2) != s(sK0,X2) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(superposition,[],[f155,f28793])).
fof(f28865,plain,(
  ( ! [X472,X471,X473,X470] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X470),X471),s(fun(sK0,bool),X472),s(fun(X470,fun(sK0,bool)),X473)),sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(X470,fun(sK0,bool)),X473),s(X470,sK7(X471,X470)))),s(sK0,sK2)))) | s('type.option.option'(X470),'const.option.NONE_0') = s('type.option.option'(X470),X471)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f28659])).
fof(f28863,plain,(
  ( ! [X494,X496,X495,X497] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),happ(s(fun(X494,fun(sK0,bool)),X495),s(X494,sK7(X496,X494))),sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X494),X496),s(fun(sK0,bool),X497),s(fun(X494,fun(sK0,bool)),X495))),s(sK0,sK2)))) | s('type.option.option'(X494),'const.option.NONE_0') = s('type.option.option'(X494),X496)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f28544])).
fof(f28858,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0)) | s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2)))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f3188])).
fof(f28857,plain,(
  ( ! [X76,X74,X75] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X76,sK0)) | s(bool,X74) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X74)),bool),X75,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X75),s(bool,'T_0')))) | s(bool,X74) = s(bool,happ(s(fun(sK0,bool),X76),s(sK0,sK2)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2521])).
fof(f28856,plain,(
  ( ! [X70,X72,X71,X73] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X73,sK0)) | s(bool,X70) != s(bool,X72) | ~p(s(bool,happ(s(fun(bool,bool),X71),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X71),s(bool,X72)))) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X73),s(sK0,sK2)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2520])).
fof(f28855,plain,(
  ( ! [X66,X64,X65] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X66,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65))) | s(bool,'T_0') != s(bool,X65) | s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK2)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2518])).
fof(f28854,plain,(
  ( ! [X54,X55] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X55,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) | s(bool,X54) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK2)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2517])).
fof(f28853,plain,(
  ( ! [X47,X45,X46] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X47,sK0)) | s(bool,X46) = s(bool,sK7('const.option.SOME_1'(s(bool,X45)),bool)) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK2)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2516])).
fof(f28852,plain,(
  ( ! [X10,X11,X9] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X11,sK0)) | s(bool,X9) = s(bool,X10) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK2)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2515])).
fof(f28851,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X56,sK0)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X52),X53),s(bool,X54),s(fun(X52,bool),X55))) | s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK2))) = s(bool,happ(s(fun(fun(X52,bool),bool),happ(s(fun(bool,fun(fun(X52,bool),bool)),happ(s(fun('type.option.option'(X52),fun(bool,fun(fun(X52,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X52),X53))),s(bool,X54))),s(fun(X52,bool),X55)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2479])).
fof(f28850,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X51,sK0)) | s(bool,'T_0') = s(bool,happ(s(fun(X47,bool),X50),s(X47,X48))) | s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),'const.option.SOME_1'(s(X47,X48))),s(bool,X49),s(fun(X47,bool),X50)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2478])).
fof(f28849,plain,(
  ( ! [X45,X43,X46,X44] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X46,sK0)) | s(bool,'T_0') = s(bool,X44) | s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.NONE_0'),s(bool,X44),s(fun(X43,bool),X45)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2477])).
fof(f28848,plain,(
  ( ! [X59,X60,X58] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X60,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X59))) | s(bool,X59) = s(bool,sK7('const.option.SOME_1'(s(bool,X58)),bool)) | s(bool,X58) = s(bool,happ(s(fun(sK0,bool),X60),s(sK0,sK2)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2440])).
fof(f28847,plain,(
  ( ! [X14,X15,X16] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X15))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X14,sK0)))) | s(sK0,sK2) != s(sK0,X15) | s(bool,X16) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X16)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2410])).
fof(f28846,plain,(
  ( ! [X10,X8,X9] : (s(sK0,X9) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X8,sK0)) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X9))) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X10)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2408])).
fof(f28845,plain,(
  ( ! [X78,X76,X77] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X76,sK0)) | s(bool,X77) != s(bool,X78) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X78))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X76),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X77)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2395])).
fof(f28844,plain,(
  ( ! [X70,X72,X71] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X70,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X71))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X72))) | s(bool,X72) = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X71)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2393])).
fof(f28843,plain,(
  ( ! [X167,X169,X166,X168] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X169,sK0)) | s(X168,f(s(bool,happ(s(fun(sK0,bool),X169),s(sK0,sK2))))) != s(X168,g(s(bool,happ(s(fun(sK0,bool),X169),s(sK0,sK2))))) | s(fun(bool,X168),X166) = s(fun(bool,X168),X167) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK9(X166,X167,X168,bool))),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2366])).
fof(f28842,plain,(
  ( ! [X159,X160] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X160,sK0)) | s('type.option.option'(bool),X159) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X160),s(sK0,sK2))))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X159) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X159,bool))),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2363])).
fof(f28841,plain,(
  ( ! [X144,X142,X143,X141] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X144,sK0)) | s(bool,f(s(X141,sK9(X142,X143,bool,X141)))) != s(bool,happ(s(fun(sK0,bool),X144),s(sK0,sK2))) | s(fun(X141,bool),X142) = s(fun(X141,bool),X143) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(X141,sK9(X142,X143,bool,X141))))),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2359])).
fof(f28840,plain,(
  ( ! [X103,X101,X102] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X102,sK0)) | s(bool,X103) != s(bool,happ(s(fun(sK0,bool),X102),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2349])).
fof(f28839,plain,(
  ( ! [X97,X95,X96] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X96,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X97))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X96),s(sK0,sK2))))) | s(bool,X95) = s(bool,X97) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X95)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2347])).
fof(f28838,plain,(
  ( ! [X14,X15] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),sK7('const.option.SOME_1'(s(fun(sK0,bool),X14)),fun(sK0,bool)),sK0)) | s(bool,X15) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X15)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2283])).
fof(f28837,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),happ(s(fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X9),fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(sK0,bool),X11))),s(fun(X9,fun(sK0,bool)),X12)),sK0)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(sK0,bool),X11),s(fun(X9,fun(sK0,bool)),X12))),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X13)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2282])).
fof(f28836,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(sK0,bool),X6),s(fun(X4,fun(sK0,bool)),X7)),sK0)) | s(bool,X8) = s(bool,happ(s(fun(sK0,bool),happ(s(fun(X4,fun(sK0,bool)),X7),s(X4,X5))),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X8)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2281])).
fof(f28835,plain,(
  ( ! [X2,X0,X3,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0)) | s(bool,X3) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X3)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f2280])).
fof(f28826,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0)) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1448])).
fof(f28825,plain,(
  ( ! [X68,X66,X67] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X66,sK0)) | ~p(s(bool,happ(s(fun(bool,bool),X68),s(bool,X67)))) | s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK2))) = s(bool,sK8(happ(s(fun(sK0,bool),X66),s(sK0,sK2)),X68,bool)) | s(bool,'T_0') = s(bool,X67)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1441])).
fof(f28824,plain,(
  ( ! [X64,X62,X65,X63] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X62,sK0)) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X63)))) | p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | s(bool,X65) != s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1440])).
fof(f28823,plain,(
  ( ! [X57,X58,X56] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X56,sK0)) | s(bool,X57) != s(bool,X58) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X57)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1438])).
fof(f28822,plain,(
  ( ! [X52,X50,X51] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X50,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) | s(bool,X52) = s(bool,happ(s(fun(sK0,bool),X50),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X51)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1436])).
fof(f28821,plain,(
  ( ! [X146,X144,X147,X145] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X147,sK0)) | s(X146,f(s(bool,happ(s(fun(sK0,bool),X147),s(sK0,sK2))))) != s(X146,g(s(bool,happ(s(fun(sK0,bool),X147),s(sK0,sK2))))) | s(fun(bool,X146),X144) = s(fun(bool,X146),X145) | s(bool,'T_0') = s(bool,sK9(X144,X145,X146,bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1416])).
fof(f28820,plain,(
  ( ! [X142,X143,X141] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X143,sK0)) | p(s(bool,happ(s(fun(bool,bool),X142),s(bool,happ(s(fun(sK0,bool),X143),s(sK0,sK2)))))) | ~p(s(bool,happ(s(fun(bool,bool),X142),s(bool,X141)))) | s(bool,'T_0') = s(bool,sK8(X141,X142,bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1415])).
fof(f28819,plain,(
  ( ! [X138,X137] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X138,sK0)) | s('type.option.option'(bool),X137) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X138),s(sK0,sK2))))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(bool,'T_0') = s(bool,sK7(X137,bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1413])).
fof(f28818,plain,(
  ( ! [X121,X122,X120,X119] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X122,sK0)) | s(bool,f(s(X119,sK9(X120,X121,bool,X119)))) != s(bool,happ(s(fun(sK0,bool),X122),s(sK0,sK2))) | s(fun(X119,bool),X120) = s(fun(X119,bool),X121) | s(bool,'T_0') = s(bool,g(s(X119,sK9(X120,X121,bool,X119))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1409])).
fof(f28817,plain,(
  ( ! [X90,X91,X89] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X90,sK0)) | ~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,happ(s(fun(sK0,bool),X90),s(sK0,sK2)))))) | s(bool,X89) = s(bool,sK8(X89,X91,bool)) | s(bool,'T_0') = s(bool,X89)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1402])).
fof(f28816,plain,(
  ( ! [X88,X87,X85,X86] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X86,sK0)) | ~p(s(bool,happ(s(fun(bool,bool),X87),s(bool,happ(s(fun(sK0,bool),X86),s(sK0,sK2)))))) | p(s(bool,happ(s(fun(bool,bool),X87),s(bool,X88)))) | s(bool,X85) != s(bool,X88) | s(bool,'T_0') = s(bool,X85)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1401])).
fof(f28815,plain,(
  ( ! [X80,X81,X79] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X80,sK0)) | s(bool,X81) != s(bool,happ(s(fun(sK0,bool),X80),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X79))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X81))) | s(bool,'T_0') = s(bool,X79)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1399])).
fof(f28814,plain,(
  ( ! [X74,X75,X73] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X74,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X74),s(sK0,sK2))))) | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,X73)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1397])).
fof(f28813,plain,(
  ( ! [X24,X23,X22] : (s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X22,sK0)) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X22,sK0),X23,sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X23),s(sK0,sK2)))) | s(bool,X24) = s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1368])).
fof(f28812,plain,(
  ( ! [X21,X19,X20,X18] : (s(sK0,X20) != s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK2)))) | p(s(bool,happ(s(fun(sK0,bool),X19),s(sK0,X20)))) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1367])).
fof(f28811,plain,(
  ( ! [X14,X12,X13] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X13))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X12,sK0)))) | s(sK0,sK2) != s(sK0,X13) | s(bool,X14) = s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1365])).
fof(f28810,plain,(
  ( ! [X6,X8,X7] : (s(sK0,X7) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X6,sK0)) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X7))) | s(bool,X8) = s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1363])).
fof(f28809,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),happ(s(fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X9),fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(sK0,bool),X11))),s(fun(X9,fun(sK0,bool)),X12)),sK0)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(sK0,bool),X11),s(fun(X9,fun(sK0,bool)),X12))),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1351])).
fof(f28808,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(sK0,bool),X6),s(fun(X4,fun(sK0,bool)),X7)),sK0)) | s(bool,X8) = s(bool,happ(s(fun(sK0,bool),happ(s(fun(X4,fun(sK0,bool)),X7),s(X4,X5))),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1350])).
fof(f28807,plain,(
  ( ! [X2,X0,X3,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0)) | s(bool,X3) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1349])).
fof(f28806,plain,(
  ( ! [X47,X46] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X47,sK0)) | s(bool,X46) = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X46)),bool))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1348])).
fof(f28805,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X45,sK0)) | s(bool,happ(s(fun(sK0,bool),X45),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),X42),s(bool,X43),s(fun(X41,bool),X44))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X41,bool),bool),happ(s(fun(bool,fun(fun(X41,bool),bool)),happ(s(fun('type.option.option'(X41),fun(bool,fun(fun(X41,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X41),X42))),s(bool,X43))),s(fun(X41,bool),X44)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1347])).
fof(f28804,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X40,sK0)) | s(bool,happ(s(fun(sK0,bool),X40),s(sK0,sK2))) = s(bool,happ(s(fun(X36,bool),X39),s(X36,X37))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.SOME_1'(s(X36,X37))),s(bool,X38),s(fun(X36,bool),X39)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1346])).
fof(f28803,plain,(
  ( ! [X35,X33,X34,X32] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X35,sK0)) | s(bool,X33) = s(bool,happ(s(fun(sK0,bool),X35),s(sK0,sK2))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X32),'const.option.NONE_0'),s(bool,X33),s(fun(X32,bool),X34)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1345])).
fof(f28802,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),sK7('const.option.SOME_1'(s(fun(sK0,bool),X0)),fun(sK0,bool)),sK0)) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1271])).
fof(f28801,plain,(
  ( ! [X61] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),sK7('const.option.SOME_1'(s(fun(sK0,bool),X61)),fun(sK0,bool)),sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK2))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1116])).
fof(f28800,plain,(
  ( ! [X132,X130,X133,X131] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),happ(s(fun(fun(X130,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X130,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X130),fun(fun(sK0,bool),fun(fun(X130,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X130),X131))),s(fun(sK0,bool),X132))),s(fun(X130,fun(sK0,bool)),X133)),sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X130),X131),s(fun(sK0,bool),X132),s(fun(X130,fun(sK0,bool)),X133))),s(sK0,sK2))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f1045])).
fof(f28799,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X0,sK0)) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f928])).
fof(f28798,plain,(
  ( ! [X6,X4,X5,X3] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.SOME_1'(s(X3,X4))),s(fun(sK0,bool),X5),s(fun(X3,fun(sK0,bool)),X6)),sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(X3,fun(sK0,bool)),X6),s(X3,X4))),s(sK0,sK2))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f922])).
fof(f28797,plain,(
  ( ! [X2,X0,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f921])).
fof(f28796,plain,(
  ( ! [X4] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X4,sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK2))))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f869])).
fof(f28795,plain,(
  ( ! [X0] : (s(sK0,X0) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X0)))) ) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f813])).
fof(f28793,plain,(
  s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK5),s(fun(sK1,sK0),sK3))) | ($spl0 | $spl2 | $spl188)),
  inference(backward_demodulation,[],[f28788,f170])).
fof(f28790,plain,(
  s('type.option.option'('type.option.option'(sK1)),'const.option.SOME_1'(s('type.option.option'(sK1),sK4))) != s('type.option.option'('type.option.option'(sK1)),'const.option.SOME_1'(s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK7(sK4,sK1)))))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f28750,f154])).
fof(f28784,plain,(
  ( ! [X26,X27,X25] : (s(sK0,sK2) = s(sK0,sK8(happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1))),'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(fun(sK0,bool),X25),s(fun(X26,fun(sK0,bool)),X27)),sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X25),s(sK0,sK2))))) ) | $spl188),
  inference(forward_demodulation,[],[f28769,f28714])).
fof(f28769,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(sK0,bool),X25),s(sK0,sK2)))) | s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1)))) = s(sK0,sK8(happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1))),'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(fun(sK0,bool),X25),s(fun(X26,fun(sK0,bool)),X27)),sK0))) ) | $spl188),
  inference(superposition,[],[f568,f28714])).
fof(f28781,plain,(
  ( ! [X7] : (s(sK0,sK2) = s(sK0,sK8(happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1))),X7,sK0)) | ~p(s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK2))))) ) | $spl188),
  inference(forward_demodulation,[],[f28761,f28714])).
fof(f28761,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK2)))) | s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1)))) = s(sK0,sK8(happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1))),X7,sK0))) ) | $spl188),
  inference(superposition,[],[f157,f28714])).
fof(f28756,plain,(
  ( ! [X1] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X1))) | s(sK0,X1) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1))))) ) | $spl188),
  inference(superposition,[],[f154,f28714])).
fof(f28679,plain,(
  ( ! [X588,X590,X589,X587] : (s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,happ(s(fun(X587,sK1),X590),s(X587,sK7(X588,X587)))))) | s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,'const.option.option_CASE_3'(s('type.option.option'(X587),X588),s(sK1,X589),s(fun(X587,sK1),X590))))) | s('type.option.option'(X587),'const.option.NONE_0') = s('type.option.option'(X587),X588)) ) | $spl2),
  inference(superposition,[],[f172,f716])).
fof(f28674,plain,(
  ( ! [X569,X565,X567,X568,X570,X566] : (s('type.option.option'(X565),'const.option.SOME_1'(s(X565,X570))) != s('type.option.option'(X565),happ(s(fun(X566,'type.option.option'(X565)),X569),s(X566,sK7(X567,X566)))) | s(X565,X570) = s(X565,sK7('const.option.option_CASE_3'(s('type.option.option'(X566),X567),s('type.option.option'(X565),X568),s(fun(X566,'type.option.option'(X565)),X569)),X565)) | s('type.option.option'(X565),'const.option.NONE_0') = s('type.option.option'(X565),'const.option.option_CASE_3'(s('type.option.option'(X566),X567),s('type.option.option'(X565),X568),s(fun(X566,'type.option.option'(X565)),X569))) | s('type.option.option'(X566),'const.option.NONE_0') = s('type.option.option'(X566),X567)) )),
  inference(superposition,[],[f611,f716])).
fof(f28669,plain,(
  ( ! [X536,X532,X534,X530,X533,X535,X531] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X530),'const.option.SOME_1'(s(X530,X535))),s(bool,X536),s(fun(X530,bool),happ(s(fun(X531,fun(X530,bool)),X534),s(X531,sK7(X532,X531)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X530,bool),'const.option.option_CASE_3'(s('type.option.option'(X531),X532),s(fun(X530,bool),X533),s(fun(X531,fun(X530,bool)),X534))),s(X530,X535))) | s('type.option.option'(X531),'const.option.NONE_0') = s('type.option.option'(X531),X532)) )),
  inference(superposition,[],[f737,f716])).
fof(f28668,plain,(
  ( ! [X528,X524,X526,X522,X529,X525,X527,X523] : (~p(s(bool,happ(s(fun(X522,bool),happ(s(fun(X523,fun(X522,bool)),X526),s(X523,sK7(X524,X523)))),s(X522,X527)))) | s(X522,X527) = s(X522,sK8('const.option.option_CASE_3'(s('type.option.option'(X528),'const.option.NONE_0'),s(X522,X527),s(fun(X528,X522),X529)),'const.option.option_CASE_3'(s('type.option.option'(X523),X524),s(fun(X522,bool),X525),s(fun(X523,fun(X522,bool)),X526)),X522)) | s('type.option.option'(X523),'const.option.NONE_0') = s('type.option.option'(X523),X524)) )),
  inference(superposition,[],[f585,f716])).
fof(f28667,plain,(
  ( ! [X520,X516,X518,X514,X521,X517,X519,X515] : (~p(s(bool,happ(s(fun(X514,bool),happ(s(fun(X515,fun(X514,bool)),X518),s(X515,sK7(X516,X515)))),s(X514,X519)))) | s(X514,X519) = s(X514,sK8(X519,'const.option.option_CASE_3'(s('type.option.option'(X520),'const.option.NONE_0'),s(fun(X514,bool),'const.option.option_CASE_3'(s('type.option.option'(X515),X516),s(fun(X514,bool),X517),s(fun(X515,fun(X514,bool)),X518))),s(fun(X520,fun(X514,bool)),X521)),X514)) | s('type.option.option'(X515),'const.option.NONE_0') = s('type.option.option'(X515),X516)) )),
  inference(superposition,[],[f568,f716])).
fof(f28665,plain,(
  ( ! [X505,X507,X502,X504,X506,X503] : (p(s(bool,happ(s(fun(X502,bool),happ(s(fun(X503,fun(X502,bool)),X506),s(X503,sK7(X504,X503)))),s(X502,sK8(X507,'const.option.option_CASE_3'(s('type.option.option'(X503),X504),s(fun(X502,bool),X505),s(fun(X503,fun(X502,bool)),X506)),X502))))) | ~p(s(bool,happ(s(fun(X502,bool),'const.option.option_CASE_3'(s('type.option.option'(X503),X504),s(fun(X502,bool),X505),s(fun(X503,fun(X502,bool)),X506))),s(X502,X507)))) | s('type.option.option'(X503),'const.option.NONE_0') = s('type.option.option'(X503),X504)) )),
  inference(superposition,[],[f158,f716])).
fof(f28664,plain,(
  ( ! [X499,X500,X496,X498,X501,X497] : (~p(s(bool,happ(s(fun(X496,bool),happ(s(fun(X497,fun(X496,bool)),X500),s(X497,sK7(X498,X497)))),s(X496,X501)))) | s(X496,X501) = s(X496,sK8(X501,'const.option.option_CASE_3'(s('type.option.option'(X497),X498),s(fun(X496,bool),X499),s(fun(X497,fun(X496,bool)),X500)),X496)) | s('type.option.option'(X497),'const.option.NONE_0') = s('type.option.option'(X497),X498)) )),
  inference(superposition,[],[f157,f716])).
fof(f28663,plain,(
  ( ! [X492,X494,X490,X493,X495,X489,X491] : (~p(s(bool,happ(s(fun(X489,bool),happ(s(fun(X490,fun(X489,bool)),X493),s(X490,sK7(X491,X490)))),s(X489,X494)))) | p(s(bool,happ(s(fun(X489,bool),'const.option.option_CASE_3'(s('type.option.option'(X490),X491),s(fun(X489,bool),X492),s(fun(X490,fun(X489,bool)),X493))),s(X489,X495)))) | s(X489,X494) != s(X489,X495) | s('type.option.option'(X490),'const.option.NONE_0') = s('type.option.option'(X490),X491)) )),
  inference(superposition,[],[f156,f716])).
fof(f28662,plain,(
  ( ! [X484,X486,X488,X485,X487] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X484,fun(bool,bool)),X487),s(X484,sK7(X485,X484)))),s(bool,X488)))) | s(bool,'F_0') = s(bool,sK8('F_0','const.option.option_CASE_3'(s('type.option.option'(X484),X485),s(fun(bool,bool),X486),s(fun(X484,fun(bool,bool)),X487)),bool)) | s(bool,'T_0') = s(bool,X488) | s('type.option.option'(X484),'const.option.NONE_0') = s('type.option.option'(X484),X485)) )),
  inference(superposition,[],[f570,f716])).
fof(f28661,plain,(
  ( ! [X480,X482,X479,X481,X483] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X479,fun(bool,bool)),X482),s(X479,sK7(X480,X479)))),s(bool,'F_0')))) | s(bool,X483) = s(bool,sK8(X483,'const.option.option_CASE_3'(s('type.option.option'(X479),X480),s(fun(bool,bool),X481),s(fun(X479,fun(bool,bool)),X482)),bool)) | s(bool,'T_0') = s(bool,X483) | s('type.option.option'(X479),'const.option.NONE_0') = s('type.option.option'(X479),X480)) )),
  inference(superposition,[],[f575,f716])).
fof(f28660,plain,(
  ( ! [X476,X478,X474,X475,X477] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X474,fun(bool,bool)),X477),s(X474,sK7(X475,X474)))),s(bool,X478)))) | p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X474),X475),s(fun(bool,bool),X476),s(fun(X474,fun(bool,bool)),X477))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X478,'const.option.option_CASE_3'(s('type.option.option'(X474),X475),s(fun(bool,bool),X476),s(fun(X474,fun(bool,bool)),X477)),bool)) | s('type.option.option'(X474),'const.option.NONE_0') = s('type.option.option'(X474),X475)) )),
  inference(superposition,[],[f675,f716])).
fof(f28659,plain,(
  ( ! [X472,X471,X473,X470] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(X470,fun(sK0,bool)),X473),s(X470,sK7(X471,X470)))),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X470),X471),s(fun(sK0,bool),X472),s(fun(X470,fun(sK0,bool)),X473)),sK0)) | s('type.option.option'(X470),'const.option.NONE_0') = s('type.option.option'(X470),X471)) ) | $spl0),
  inference(superposition,[],[f869,f716])).
fof(f28607,plain,(
  ( ! [X243,X241,X239,X242,X240] : (~p(s(bool,happ(s(fun(bool,bool),X243),s(bool,happ(s(fun(X239,bool),X242),s(X239,sK7(X240,X239))))))) | p(s(bool,happ(s(fun(bool,bool),X243),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X239),X240),s(bool,X241),s(fun(X239,bool),X242)),X243,bool)) | s('type.option.option'(X239),'const.option.NONE_0') = s('type.option.option'(X239),X240)) )),
  inference(superposition,[],[f675,f716])).
fof(f28606,plain,(
  ( ! [X237,X235,X238,X236,X234] : (~p(s(bool,happ(s(fun(bool,bool),X238),s(bool,happ(s(fun(X234,bool),X237),s(X234,sK7(X235,X234))))))) | s(bool,'F_0') = s(bool,sK8('F_0',X238,bool)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X234),X235),s(bool,X236),s(fun(X234,bool),X237))) | s('type.option.option'(X234),'const.option.NONE_0') = s('type.option.option'(X234),X235)) )),
  inference(superposition,[],[f570,f716])).
fof(f28605,plain,(
  ( ! [X231,X229,X233,X230,X232] : (s(bool,'F_0') != s(bool,happ(s(fun(X229,bool),X232),s(X229,sK7(X230,X229)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X233))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X229),X230),s(bool,X231),s(fun(X229,bool),X232))))) | s(bool,'T_0') = s(bool,X233) | s('type.option.option'(X229),'const.option.NONE_0') = s('type.option.option'(X229),X230)) )),
  inference(superposition,[],[f537,f716])).
fof(f28604,plain,(
  ( ! [X227,X225,X228,X226,X224] : (s(bool,X228) != s(bool,happ(s(fun(X224,bool),X227),s(X224,sK7(X225,X224)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X228))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X224),X225),s(bool,X226),s(fun(X224,bool),X227))) | s('type.option.option'(X224),'const.option.NONE_0') = s('type.option.option'(X224),X225)) )),
  inference(superposition,[],[f532,f716])).
fof(f28603,plain,(
  ( ! [X222,X220,X219,X223,X221] : (s(bool,X223) != s(bool,happ(s(fun(X219,bool),X222),s(X219,sK7(X220,X219)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X219),X220),s(bool,X221),s(fun(X219,bool),X222))))) | s(bool,'T_0') = s(bool,X223) | s('type.option.option'(X219),'const.option.NONE_0') = s('type.option.option'(X219),X220)) )),
  inference(superposition,[],[f532,f716])).
fof(f28602,plain,(
  ( ! [X218,X216,X215,X217,X214] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X214,bool),X217),s(X214,sK7(X215,X214)))))) | s(bool,X218) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X214),X215),s(bool,X216),s(fun(X214,bool),X217))) | s(bool,'T_0') = s(bool,X218) | s('type.option.option'(X214),'const.option.NONE_0') = s('type.option.option'(X214),X215)) )),
  inference(superposition,[],[f497,f716])).
fof(f28601,plain,(
  ( ! [X212,X210,X213,X211,X209] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X213))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X209,bool),X212),s(X209,sK7(X210,X209)))))) | s(bool,'F_0') = s(bool,X213) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X209),X210),s(bool,X211),s(fun(X209,bool),X212))) | s('type.option.option'(X209),'const.option.NONE_0') = s('type.option.option'(X209),X210)) )),
  inference(superposition,[],[f492,f716])).
fof(f28600,plain,(
  ( ! [X208,X206,X204,X207,X205] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X208))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X204,bool),X207),s(X204,sK7(X205,X204)))))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X204),X205),s(bool,X206),s(fun(X204,bool),X207))) | s(bool,'T_0') = s(bool,X208) | s('type.option.option'(X204),'const.option.NONE_0') = s('type.option.option'(X204),X205)) )),
  inference(superposition,[],[f492,f716])).
fof(f28599,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X202),'const.option.NONE_0'),s(bool,happ(s(fun(X198,bool),X201),s(X198,sK7(X199,X198)))),s(fun(X202,bool),X203))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X198),X199),s(bool,X200),s(fun(X198,bool),X201))) | s('type.option.option'(X198),'const.option.NONE_0') = s('type.option.option'(X198),X199)) )),
  inference(superposition,[],[f239,f716])).
fof(f28598,plain,(
  ( ! [X196,X194,X197,X195,X193] : (p(s(bool,happ(s(fun(X193,bool),X196),s(X193,sK7(X194,X193))))) | s(bool,X197) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X193),X194),s(bool,X195),s(fun(X193,bool),X196))) | s(bool,'T_0') = s(bool,X197) | s('type.option.option'(X193),'const.option.NONE_0') = s('type.option.option'(X193),X194)) )),
  inference(superposition,[],[f221,f716])).
fof(f28596,plain,(
  ( ! [X187,X185,X188,X186] : (~p(s(bool,happ(s(fun(X185,bool),X188),s(X185,sK7(X186,X185))))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X185),X186),s(bool,X187),s(fun(X185,bool),X188))) | s('type.option.option'(X185),'const.option.NONE_0') = s('type.option.option'(X185),X186)) )),
  inference(superposition,[],[f192,f716])).
fof(f28591,plain,(
  ( ! [X163,X161,X164,X162,X160] : (~p(s(bool,happ(s(fun(X160,bool),X163),s(X160,sK7(X161,X160))))) | ~p(s(bool,X164)) | s(bool,X164) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X160),X161),s(bool,X162),s(fun(X160,bool),X163))) | s('type.option.option'(X160),'const.option.NONE_0') = s('type.option.option'(X160),X161)) )),
  inference(superposition,[],[f162,f716])).
fof(f28589,plain,(
  ( ! [X154,X152,X151,X153,X150] : (p(s(bool,happ(s(fun(X150,bool),X153),s(X150,sK7(X151,X150))))) | p(s(bool,X154)) | s(bool,X154) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X150),X151),s(bool,X152),s(fun(X150,bool),X153))) | s('type.option.option'(X150),'const.option.NONE_0') = s('type.option.option'(X150),X151)) )),
  inference(superposition,[],[f159,f716])).
fof(f28583,plain,(
  ( ! [X127,X125,X130,X128,X126,X131,X129] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X125),'const.option.SOME_1'(s(X125,happ(s(fun(X126,X125),X129),s(X126,sK7(X127,X126)))))),s(bool,X130),s(fun(X125,bool),X131))) | s(bool,'F_0') = s(bool,happ(s(fun(X125,bool),X131),s(X125,'const.option.option_CASE_3'(s('type.option.option'(X126),X127),s(X125,X128),s(fun(X126,X125),X129))))) | s('type.option.option'(X126),'const.option.NONE_0') = s('type.option.option'(X126),X127)) )),
  inference(superposition,[],[f737,f716])).
fof(f28582,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s('type.option.option'(X119),X124) != s('type.option.option'(X119),'const.option.SOME_1'(s(X119,happ(s(fun(X120,X119),X123),s(X120,sK7(X121,X120)))))) | s(X119,sK7(X124,X119)) = s(X119,'const.option.option_CASE_3'(s('type.option.option'(X120),X121),s(X119,X122),s(fun(X120,X119),X123))) | s('type.option.option'(X119),'const.option.NONE_0') = s('type.option.option'(X119),X124) | s('type.option.option'(X120),'const.option.NONE_0') = s('type.option.option'(X120),X121)) )),
  inference(superposition,[],[f611,f716])).
fof(f28581,plain,(
  ( ! [X111,X118,X116,X114,X112,X117,X115,X113] : (~p(s(bool,happ(s(fun(X111,bool),X116),s(X111,happ(s(fun(X112,X111),X115),s(X112,sK7(X113,X112))))))) | s(X111,'const.option.option_CASE_3'(s('type.option.option'(X112),X113),s(X111,X114),s(fun(X112,X111),X115))) = s(X111,sK8('const.option.option_CASE_3'(s('type.option.option'(X117),'const.option.NONE_0'),s(X111,'const.option.option_CASE_3'(s('type.option.option'(X112),X113),s(X111,X114),s(fun(X112,X111),X115))),s(fun(X117,X111),X118)),X116,X111)) | s('type.option.option'(X112),'const.option.NONE_0') = s('type.option.option'(X112),X113)) )),
  inference(superposition,[],[f585,f716])).
fof(f28580,plain,(
  ( ! [X103,X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,happ(s(fun(X103,bool),X108),s(X103,happ(s(fun(X104,X103),X107),s(X104,sK7(X105,X104))))))) | s(X103,'const.option.option_CASE_3'(s('type.option.option'(X104),X105),s(X103,X106),s(fun(X104,X103),X107))) = s(X103,sK8('const.option.option_CASE_3'(s('type.option.option'(X104),X105),s(X103,X106),s(fun(X104,X103),X107)),'const.option.option_CASE_3'(s('type.option.option'(X109),'const.option.NONE_0'),s(fun(X103,bool),X108),s(fun(X109,fun(X103,bool)),X110)),X103)) | s('type.option.option'(X104),'const.option.NONE_0') = s('type.option.option'(X104),X105)) )),
  inference(superposition,[],[f568,f716])).
fof(f28579,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (s('type.option.option'(X95),'const.option.SOME_1'(s(X95,X100))) != s('type.option.option'(X95),'const.option.SOME_1'(s(X95,happ(s(fun(X96,X95),X99),s(X96,sK7(X97,X96)))))) | s(X95,'const.option.option_CASE_3'(s('type.option.option'(X101),'const.option.NONE_0'),s(X95,X100),s(fun(X101,X95),X102))) = s(X95,'const.option.option_CASE_3'(s('type.option.option'(X96),X97),s(X95,X98),s(fun(X96,X95),X99))) | s('type.option.option'(X96),'const.option.NONE_0') = s('type.option.option'(X96),X97)) )),
  inference(superposition,[],[f491,f716])).
fof(f28578,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (s('type.option.option'(X87),'const.option.SOME_1'(s(X87,X92))) != s('type.option.option'(X87),'const.option.SOME_1'(s(X87,happ(s(fun(X88,X87),X91),s(X88,sK7(X89,X88)))))) | s(X87,X92) = s(X87,'const.option.option_CASE_3'(s('type.option.option'(X93),'const.option.NONE_0'),s(X87,'const.option.option_CASE_3'(s('type.option.option'(X88),X89),s(X87,X90),s(fun(X88,X87),X91))),s(fun(X93,X87),X94))) | s('type.option.option'(X88),'const.option.NONE_0') = s('type.option.option'(X88),X89)) )),
  inference(superposition,[],[f491,f716])).
fof(f28572,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (~p(s(bool,happ(s(fun(X45,bool),X50),s(X45,happ(s(fun(X46,X45),X49),s(X46,sK7(X47,X46))))))) | s(X45,'const.option.option_CASE_3'(s('type.option.option'(X46),X47),s(X45,X48),s(fun(X46,X45),X49))) = s(X45,sK8('const.option.option_CASE_3'(s('type.option.option'(X46),X47),s(X45,X48),s(fun(X46,X45),X49)),X50,X45)) | s('type.option.option'(X46),'const.option.NONE_0') = s('type.option.option'(X46),X47)) )),
  inference(superposition,[],[f157,f716])).
fof(f28571,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (~p(s(bool,happ(s(fun(X38,bool),X43),s(X38,happ(s(fun(X39,X38),X42),s(X39,sK7(X40,X39))))))) | p(s(bool,happ(s(fun(X38,bool),X43),s(X38,X44)))) | s(X38,X44) != s(X38,'const.option.option_CASE_3'(s('type.option.option'(X39),X40),s(X38,X41),s(fun(X39,X38),X42))) | s('type.option.option'(X39),'const.option.NONE_0') = s('type.option.option'(X39),X40)) )),
  inference(superposition,[],[f156,f716])).
fof(f28569,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X26,X31) != s(X26,happ(s(fun(X27,X26),X30),s(X27,sK7(X28,X27)))) | s('type.option.option'(X26),'const.option.SOME_1'(s(X26,X31))) = s('type.option.option'(X26),'const.option.SOME_1'(s(X26,'const.option.option_CASE_3'(s('type.option.option'(X27),X28),s(X26,X29),s(fun(X27,X26),X30))))) | s('type.option.option'(X27),'const.option.NONE_0') = s('type.option.option'(X27),X28)) )),
  inference(superposition,[],[f155,f716])).
fof(f28567,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s('type.option.option'(X14),'const.option.SOME_1'(s(X14,X19))) != s('type.option.option'(X14),'const.option.SOME_1'(s(X14,happ(s(fun(X15,X14),X18),s(X15,sK7(X16,X15)))))) | s(X14,X19) = s(X14,'const.option.option_CASE_3'(s('type.option.option'(X15),X16),s(X14,X17),s(fun(X15,X14),X18))) | s('type.option.option'(X15),'const.option.NONE_0') = s('type.option.option'(X15),X16)) )),
  inference(superposition,[],[f154,f716])).
fof(f28563,plain,(
  ( ! [X613,X615,X616,X614] : (s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,'const.option.option_CASE_3'(s('type.option.option'(X613),X615),s(sK1,X616),s(fun(X613,sK1),X614))))) | s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,happ(s(fun(X613,sK1),X614),s(X613,sK7(X615,X613)))))) | s('type.option.option'(X613),'const.option.NONE_0') = s('type.option.option'(X613),X615)) ) | $spl2),
  inference(superposition,[],[f172,f716])).
fof(f28559,plain,(
  ( ! [X592,X594,X590,X593,X589,X591] : (s('type.option.option'(X589),'const.option.SOME_1'(s(X589,X594))) != s('type.option.option'(X589),'const.option.option_CASE_3'(s('type.option.option'(X590),X592),s('type.option.option'(X589),X593),s(fun(X590,'type.option.option'(X589)),X591))) | s(X589,X594) = s(X589,sK7(happ(s(fun(X590,'type.option.option'(X589)),X591),s(X590,sK7(X592,X590))),X589)) | s('type.option.option'(X589),'const.option.NONE_0') = s('type.option.option'(X589),happ(s(fun(X590,'type.option.option'(X589)),X591),s(X590,sK7(X592,X590)))) | s('type.option.option'(X590),'const.option.NONE_0') = s('type.option.option'(X590),X592)) )),
  inference(superposition,[],[f611,f716])).
fof(f28554,plain,(
  ( ! [X557,X559,X555,X560,X556,X558,X554] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X554),'const.option.SOME_1'(s(X554,X559))),s(bool,X560),s(fun(X554,bool),'const.option.option_CASE_3'(s('type.option.option'(X555),X557),s(fun(X554,bool),X558),s(fun(X555,fun(X554,bool)),X556))))) | s(bool,'F_0') = s(bool,happ(s(fun(X554,bool),happ(s(fun(X555,fun(X554,bool)),X556),s(X555,sK7(X557,X555)))),s(X554,X559))) | s('type.option.option'(X555),'const.option.NONE_0') = s('type.option.option'(X555),X557)) )),
  inference(superposition,[],[f737,f716])).
fof(f28553,plain,(
  ( ! [X546,X553,X549,X551,X547,X552,X548,X550] : (~p(s(bool,happ(s(fun(X546,bool),'const.option.option_CASE_3'(s('type.option.option'(X547),X549),s(fun(X546,bool),X550),s(fun(X547,fun(X546,bool)),X548))),s(X546,X551)))) | s(X546,X551) = s(X546,sK8('const.option.option_CASE_3'(s('type.option.option'(X552),'const.option.NONE_0'),s(X546,X551),s(fun(X552,X546),X553)),happ(s(fun(X547,fun(X546,bool)),X548),s(X547,sK7(X549,X547))),X546)) | s('type.option.option'(X547),'const.option.NONE_0') = s('type.option.option'(X547),X549)) )),
  inference(superposition,[],[f585,f716])).
fof(f28552,plain,(
  ( ! [X540,X542,X538,X545,X541,X543,X539,X544] : (~p(s(bool,happ(s(fun(X538,bool),'const.option.option_CASE_3'(s('type.option.option'(X539),X541),s(fun(X538,bool),X542),s(fun(X539,fun(X538,bool)),X540))),s(X538,X543)))) | s(X538,X543) = s(X538,sK8(X543,'const.option.option_CASE_3'(s('type.option.option'(X544),'const.option.NONE_0'),s(fun(X538,bool),happ(s(fun(X539,fun(X538,bool)),X540),s(X539,sK7(X541,X539)))),s(fun(X544,fun(X538,bool)),X545)),X538)) | s('type.option.option'(X539),'const.option.NONE_0') = s('type.option.option'(X539),X541)) )),
  inference(superposition,[],[f568,f716])).
fof(f28550,plain,(
  ( ! [X528,X530,X526,X529,X531,X527] : (p(s(bool,happ(s(fun(X526,bool),'const.option.option_CASE_3'(s('type.option.option'(X527),X529),s(fun(X526,bool),X530),s(fun(X527,fun(X526,bool)),X528))),s(X526,sK8(X531,happ(s(fun(X527,fun(X526,bool)),X528),s(X527,sK7(X529,X527))),X526))))) | ~p(s(bool,happ(s(fun(X526,bool),happ(s(fun(X527,fun(X526,bool)),X528),s(X527,sK7(X529,X527)))),s(X526,X531)))) | s('type.option.option'(X527),'const.option.NONE_0') = s('type.option.option'(X527),X529)) )),
  inference(superposition,[],[f158,f716])).
fof(f28549,plain,(
  ( ! [X524,X520,X522,X525,X521,X523] : (~p(s(bool,happ(s(fun(X520,bool),'const.option.option_CASE_3'(s('type.option.option'(X521),X523),s(fun(X520,bool),X524),s(fun(X521,fun(X520,bool)),X522))),s(X520,X525)))) | s(X520,X525) = s(X520,sK8(X525,happ(s(fun(X521,fun(X520,bool)),X522),s(X521,sK7(X523,X521))),X520)) | s('type.option.option'(X521),'const.option.NONE_0') = s('type.option.option'(X521),X523)) )),
  inference(superposition,[],[f157,f716])).
fof(f28548,plain,(
  ( ! [X516,X518,X514,X517,X519,X513,X515] : (~p(s(bool,happ(s(fun(X513,bool),'const.option.option_CASE_3'(s('type.option.option'(X514),X516),s(fun(X513,bool),X517),s(fun(X514,fun(X513,bool)),X515))),s(X513,X518)))) | p(s(bool,happ(s(fun(X513,bool),happ(s(fun(X514,fun(X513,bool)),X515),s(X514,sK7(X516,X514)))),s(X513,X519)))) | s(X513,X518) != s(X513,X519) | s('type.option.option'(X514),'const.option.NONE_0') = s('type.option.option'(X514),X516)) )),
  inference(superposition,[],[f156,f716])).
fof(f28547,plain,(
  ( ! [X509,X511,X512,X508,X510] : (~p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X508),X510),s(fun(bool,bool),X511),s(fun(X508,fun(bool,bool)),X509))),s(bool,X512)))) | s(bool,'F_0') = s(bool,sK8('F_0',happ(s(fun(X508,fun(bool,bool)),X509),s(X508,sK7(X510,X508))),bool)) | s(bool,'T_0') = s(bool,X512) | s('type.option.option'(X508),'const.option.NONE_0') = s('type.option.option'(X508),X510)) )),
  inference(superposition,[],[f570,f716])).
fof(f28546,plain,(
  ( ! [X505,X507,X504,X506,X503] : (~p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X503),X505),s(fun(bool,bool),X506),s(fun(X503,fun(bool,bool)),X504))),s(bool,'F_0')))) | s(bool,X507) = s(bool,sK8(X507,happ(s(fun(X503,fun(bool,bool)),X504),s(X503,sK7(X505,X503))),bool)) | s(bool,'T_0') = s(bool,X507) | s('type.option.option'(X503),'const.option.NONE_0') = s('type.option.option'(X503),X505)) )),
  inference(superposition,[],[f575,f716])).
fof(f28545,plain,(
  ( ! [X499,X500,X502,X498,X501] : (~p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X498),X500),s(fun(bool,bool),X501),s(fun(X498,fun(bool,bool)),X499))),s(bool,X502)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X498,fun(bool,bool)),X499),s(X498,sK7(X500,X498)))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X502,happ(s(fun(X498,fun(bool,bool)),X499),s(X498,sK7(X500,X498))),bool)) | s('type.option.option'(X498),'const.option.NONE_0') = s('type.option.option'(X498),X500)) )),
  inference(superposition,[],[f675,f716])).
fof(f28544,plain,(
  ( ! [X494,X496,X495,X497] : (~p(s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X494),X496),s(fun(sK0,bool),X497),s(fun(X494,fun(sK0,bool)),X495))),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),happ(s(fun(X494,fun(sK0,bool)),X495),s(X494,sK7(X496,X494))),sK0)) | s('type.option.option'(X494),'const.option.NONE_0') = s('type.option.option'(X494),X496)) ) | $spl0),
  inference(superposition,[],[f869,f716])).
fof(f28492,plain,(
  ( ! [X255,X257,X254,X256,X258] : (~p(s(bool,happ(s(fun(bool,bool),X258),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X254),X256),s(bool,X257),s(fun(X254,bool),X255)))))) | p(s(bool,happ(s(fun(bool,bool),X258),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(happ(s(fun(X254,bool),X255),s(X254,sK7(X256,X254))),X258,bool)) | s('type.option.option'(X254),'const.option.NONE_0') = s('type.option.option'(X254),X256)) )),
  inference(superposition,[],[f675,f716])).
fof(f28491,plain,(
  ( ! [X253,X251,X249,X252,X250] : (~p(s(bool,happ(s(fun(bool,bool),X253),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X249),X251),s(bool,X252),s(fun(X249,bool),X250)))))) | s(bool,'F_0') = s(bool,sK8('F_0',X253,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X249,bool),X250),s(X249,sK7(X251,X249)))) | s('type.option.option'(X249),'const.option.NONE_0') = s('type.option.option'(X249),X251)) )),
  inference(superposition,[],[f570,f716])).
fof(f28490,plain,(
  ( ! [X246,X244,X248,X247,X245] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X244),X246),s(bool,X247),s(fun(X244,bool),X245))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X248))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X244,bool),X245),s(X244,sK7(X246,X244)))))) | s(bool,'T_0') = s(bool,X248) | s('type.option.option'(X244),'const.option.NONE_0') = s('type.option.option'(X244),X246)) )),
  inference(superposition,[],[f537,f716])).
fof(f28489,plain,(
  ( ! [X243,X241,X239,X242,X240] : (s(bool,X243) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X239),X241),s(bool,X242),s(fun(X239,bool),X240))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X243))) | s(bool,'T_0') = s(bool,happ(s(fun(X239,bool),X240),s(X239,sK7(X241,X239)))) | s('type.option.option'(X239),'const.option.NONE_0') = s('type.option.option'(X239),X241)) )),
  inference(superposition,[],[f532,f716])).
fof(f28488,plain,(
  ( ! [X237,X235,X238,X236,X234] : (s(bool,X238) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X234),X236),s(bool,X237),s(fun(X234,bool),X235))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X234,bool),X235),s(X234,sK7(X236,X234)))))) | s(bool,'T_0') = s(bool,X238) | s('type.option.option'(X234),'const.option.NONE_0') = s('type.option.option'(X234),X236)) )),
  inference(superposition,[],[f532,f716])).
fof(f28487,plain,(
  ( ! [X231,X229,X233,X230,X232] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X229),X231),s(bool,X232),s(fun(X229,bool),X230))))) | s(bool,X233) = s(bool,happ(s(fun(X229,bool),X230),s(X229,sK7(X231,X229)))) | s(bool,'T_0') = s(bool,X233) | s('type.option.option'(X229),'const.option.NONE_0') = s('type.option.option'(X229),X231)) )),
  inference(superposition,[],[f497,f716])).
fof(f28486,plain,(
  ( ! [X227,X225,X228,X226,X224] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X228))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X224),X226),s(bool,X227),s(fun(X224,bool),X225))))) | s(bool,'F_0') = s(bool,X228) | s(bool,'T_0') = s(bool,happ(s(fun(X224,bool),X225),s(X224,sK7(X226,X224)))) | s('type.option.option'(X224),'const.option.NONE_0') = s('type.option.option'(X224),X226)) )),
  inference(superposition,[],[f492,f716])).
fof(f28485,plain,(
  ( ! [X222,X220,X219,X223,X221] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X223))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X219),X221),s(bool,X222),s(fun(X219,bool),X220))))) | s(bool,'F_0') = s(bool,happ(s(fun(X219,bool),X220),s(X219,sK7(X221,X219)))) | s(bool,'T_0') = s(bool,X223) | s('type.option.option'(X219),'const.option.NONE_0') = s('type.option.option'(X219),X221)) )),
  inference(superposition,[],[f492,f716])).
fof(f28484,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X217),'const.option.NONE_0'),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X213),X215),s(bool,X216),s(fun(X213,bool),X214))),s(fun(X217,bool),X218))) | s(bool,'F_0') = s(bool,happ(s(fun(X213,bool),X214),s(X213,sK7(X215,X213)))) | s('type.option.option'(X213),'const.option.NONE_0') = s('type.option.option'(X213),X215)) )),
  inference(superposition,[],[f239,f716])).
fof(f28483,plain,(
  ( ! [X212,X210,X208,X211,X209] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X208),X210),s(bool,X211),s(fun(X208,bool),X209)))) | s(bool,X212) = s(bool,happ(s(fun(X208,bool),X209),s(X208,sK7(X210,X208)))) | s(bool,'T_0') = s(bool,X212) | s('type.option.option'(X208),'const.option.NONE_0') = s('type.option.option'(X208),X210)) )),
  inference(superposition,[],[f221,f716])).
fof(f28481,plain,(
  ( ! [X202,X200,X203,X201] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X200),X202),s(bool,X203),s(fun(X200,bool),X201)))) | s(bool,'T_0') = s(bool,happ(s(fun(X200,bool),X201),s(X200,sK7(X202,X200)))) | s('type.option.option'(X200),'const.option.NONE_0') = s('type.option.option'(X200),X202)) )),
  inference(superposition,[],[f192,f716])).
fof(f28476,plain,(
  ( ! [X177,X175,X178,X176,X179] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X175),X177),s(bool,X178),s(fun(X175,bool),X176)))) | ~p(s(bool,X179)) | s(bool,X179) = s(bool,happ(s(fun(X175,bool),X176),s(X175,sK7(X177,X175)))) | s('type.option.option'(X175),'const.option.NONE_0') = s('type.option.option'(X175),X177)) )),
  inference(superposition,[],[f162,f716])).
fof(f28474,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X165),X167),s(bool,X168),s(fun(X165,bool),X166)))) | p(s(bool,X169)) | s(bool,X169) = s(bool,happ(s(fun(X165,bool),X166),s(X165,sK7(X167,X165)))) | s('type.option.option'(X165),'const.option.NONE_0') = s('type.option.option'(X165),X167)) )),
  inference(superposition,[],[f159,f716])).
fof(f28472,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X151),X153),s(bool,X154),s(fun(X151,bool),X152)))) | s(X151,sK7(X153,X151)) = s(X151,sK8(sK7(X153,X151),X152,X151)) | s('type.option.option'(X151),'const.option.NONE_0') = s('type.option.option'(X151),X153)) )),
  inference(superposition,[],[f157,f716])).
fof(f28471,plain,(
  ( ! [X146,X149,X147,X150,X148] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X146),X148),s(bool,X149),s(fun(X146,bool),X147)))) | p(s(bool,happ(s(fun(X146,bool),X147),s(X146,X150)))) | s(X146,X150) != s(X146,sK7(X148,X146)) | s('type.option.option'(X146),'const.option.NONE_0') = s('type.option.option'(X146),X148)) )),
  inference(superposition,[],[f156,f716])).
fof(f28470,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X140),X142),s(bool,X143),s(fun(X140,bool),X141)))) | s(X140,sK7(X142,X140)) = s(X140,sK8('const.option.option_CASE_3'(s('type.option.option'(X144),'const.option.NONE_0'),s(X140,sK7(X142,X140)),s(fun(X144,X140),X145)),X141,X140)) | s('type.option.option'(X140),'const.option.NONE_0') = s('type.option.option'(X140),X142)) )),
  inference(superposition,[],[f585,f716])).
fof(f28469,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X134),X136),s(bool,X137),s(fun(X134,bool),X135)))) | s(X134,sK7(X136,X134)) = s(X134,sK8(sK7(X136,X134),'const.option.option_CASE_3'(s('type.option.option'(X138),'const.option.NONE_0'),s(fun(X134,bool),X135),s(fun(X138,fun(X134,bool)),X139)),X134)) | s('type.option.option'(X134),'const.option.NONE_0') = s('type.option.option'(X134),X136)) )),
  inference(superposition,[],[f568,f716])).
fof(f28467,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X122),'const.option.SOME_1'(s(X122,'const.option.option_CASE_3'(s('type.option.option'(X123),X125),s(X122,X126),s(fun(X123,X122),X124))))),s(bool,X127),s(fun(X122,bool),X128))) | s(bool,'F_0') = s(bool,happ(s(fun(X122,bool),X128),s(X122,happ(s(fun(X123,X122),X124),s(X123,sK7(X125,X123)))))) | s('type.option.option'(X123),'const.option.NONE_0') = s('type.option.option'(X123),X125)) )),
  inference(superposition,[],[f737,f716])).
fof(f28466,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (s('type.option.option'(X116),X121) != s('type.option.option'(X116),'const.option.SOME_1'(s(X116,'const.option.option_CASE_3'(s('type.option.option'(X117),X119),s(X116,X120),s(fun(X117,X116),X118))))) | s(X116,sK7(X121,X116)) = s(X116,happ(s(fun(X117,X116),X118),s(X117,sK7(X119,X117)))) | s('type.option.option'(X116),'const.option.NONE_0') = s('type.option.option'(X116),X121) | s('type.option.option'(X117),'const.option.NONE_0') = s('type.option.option'(X117),X119)) )),
  inference(superposition,[],[f611,f716])).
fof(f28465,plain,(
  ( ! [X111,X109,X114,X112,X110,X108,X115,X113] : (~p(s(bool,happ(s(fun(X108,bool),X113),s(X108,'const.option.option_CASE_3'(s('type.option.option'(X109),X111),s(X108,X112),s(fun(X109,X108),X110)))))) | s(X108,happ(s(fun(X109,X108),X110),s(X109,sK7(X111,X109)))) = s(X108,sK8('const.option.option_CASE_3'(s('type.option.option'(X114),'const.option.NONE_0'),s(X108,happ(s(fun(X109,X108),X110),s(X109,sK7(X111,X109)))),s(fun(X114,X108),X115)),X113,X108)) | s('type.option.option'(X109),'const.option.NONE_0') = s('type.option.option'(X109),X111)) )),
  inference(superposition,[],[f585,f716])).
fof(f28464,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X106,X104] : (~p(s(bool,happ(s(fun(X100,bool),X105),s(X100,'const.option.option_CASE_3'(s('type.option.option'(X101),X103),s(X100,X104),s(fun(X101,X100),X102)))))) | s(X100,happ(s(fun(X101,X100),X102),s(X101,sK7(X103,X101)))) = s(X100,sK8(happ(s(fun(X101,X100),X102),s(X101,sK7(X103,X101))),'const.option.option_CASE_3'(s('type.option.option'(X106),'const.option.NONE_0'),s(fun(X100,bool),X105),s(fun(X106,fun(X100,bool)),X107)),X100)) | s('type.option.option'(X101),'const.option.NONE_0') = s('type.option.option'(X101),X103)) )),
  inference(superposition,[],[f568,f716])).
fof(f28463,plain,(
  ( ! [X94,X92,X99,X97,X95,X93,X98,X96] : (s('type.option.option'(X92),'const.option.SOME_1'(s(X92,X97))) != s('type.option.option'(X92),'const.option.SOME_1'(s(X92,'const.option.option_CASE_3'(s('type.option.option'(X93),X95),s(X92,X96),s(fun(X93,X92),X94))))) | s(X92,happ(s(fun(X93,X92),X94),s(X93,sK7(X95,X93)))) = s(X92,'const.option.option_CASE_3'(s('type.option.option'(X98),'const.option.NONE_0'),s(X92,X97),s(fun(X98,X92),X99))) | s('type.option.option'(X93),'const.option.NONE_0') = s('type.option.option'(X93),X95)) )),
  inference(superposition,[],[f491,f716])).
fof(f28462,plain,(
  ( ! [X90,X88,X87,X85,X91,X89,X86,X84] : (s('type.option.option'(X84),'const.option.SOME_1'(s(X84,X89))) != s('type.option.option'(X84),'const.option.SOME_1'(s(X84,'const.option.option_CASE_3'(s('type.option.option'(X85),X87),s(X84,X88),s(fun(X85,X84),X86))))) | s(X84,X89) = s(X84,'const.option.option_CASE_3'(s('type.option.option'(X90),'const.option.NONE_0'),s(X84,happ(s(fun(X85,X84),X86),s(X85,sK7(X87,X85)))),s(fun(X90,X84),X91))) | s('type.option.option'(X85),'const.option.NONE_0') = s('type.option.option'(X85),X87)) )),
  inference(superposition,[],[f491,f716])).
fof(f28456,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,'const.option.option_CASE_3'(s('type.option.option'(X43),X45),s(X42,X46),s(fun(X43,X42),X44)))))) | s(X42,happ(s(fun(X43,X42),X44),s(X43,sK7(X45,X43)))) = s(X42,sK8(happ(s(fun(X43,X42),X44),s(X43,sK7(X45,X43))),X47,X42)) | s('type.option.option'(X43),'const.option.NONE_0') = s('type.option.option'(X43),X45)) )),
  inference(superposition,[],[f157,f716])).
fof(f28455,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,'const.option.option_CASE_3'(s('type.option.option'(X36),X38),s(X35,X39),s(fun(X36,X35),X37)))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41)))) | s(X35,X41) != s(X35,happ(s(fun(X36,X35),X37),s(X36,sK7(X38,X36)))) | s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X38)) )),
  inference(superposition,[],[f156,f716])).
fof(f28453,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s(X23,X28) != s(X23,'const.option.option_CASE_3'(s('type.option.option'(X24),X26),s(X23,X27),s(fun(X24,X23),X25))) | s('type.option.option'(X23),'const.option.SOME_1'(s(X23,X28))) = s('type.option.option'(X23),'const.option.SOME_1'(s(X23,happ(s(fun(X24,X23),X25),s(X24,sK7(X26,X24)))))) | s('type.option.option'(X24),'const.option.NONE_0') = s('type.option.option'(X24),X26)) )),
  inference(superposition,[],[f155,f716])).
fof(f28451,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s('type.option.option'(X11),'const.option.SOME_1'(s(X11,X16))) != s('type.option.option'(X11),'const.option.SOME_1'(s(X11,'const.option.option_CASE_3'(s('type.option.option'(X12),X14),s(X11,X15),s(fun(X12,X11),X13))))) | s(X11,X16) = s(X11,happ(s(fun(X12,X11),X13),s(X12,sK7(X14,X12)))) | s('type.option.option'(X12),'const.option.NONE_0') = s('type.option.option'(X12),X14)) )),
  inference(superposition,[],[f154,f716])).
fof(f28447,plain,(
  ( ! [X90,X88,X91,X89] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),happ(s(fun(X88,'type.option.option'(bool)),X91),s(X88,sK7(X89,X88)))) | s('type.option.option'(X88),'const.option.NONE_0') = s('type.option.option'(X88),X89) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X88),X89),s('type.option.option'(bool),X90),s(fun(X88,'type.option.option'(bool)),X91)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X88),X89),s('type.option.option'(bool),X90),s(fun(X88,'type.option.option'(bool)),X91)))) )),
  inference(superposition,[],[f716,f610])).
fof(f28703,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s('type.option.option'(X52),'const.option.NONE_0') = s('type.option.option'(X52),X53) | s(X54,'const.option.option_CASE_3'(s('type.option.option'(X52),X53),s(X54,X56),s(fun(X52,X54),X55))) = s(X54,happ(s(fun(X52,X54),X55),s(X52,sK7(sK7('const.option.SOME_1'(s('type.option.option'(X52),X53)),'type.option.option'(X52)),X52))))) )),
  inference(forward_demodulation,[],[f28362,f946])).
fof(f28362,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(X54,'const.option.option_CASE_3'(s('type.option.option'(X52),X53),s(X54,X56),s(fun(X52,X54),X55))) = s(X54,happ(s(fun(X52,X54),X55),s(X52,sK7(sK7('const.option.SOME_1'(s('type.option.option'(X52),X53)),'type.option.option'(X52)),X52)))) | s('type.option.option'(X52),'const.option.NONE_0') = s('type.option.option'(X52),sK7('const.option.SOME_1'(s('type.option.option'(X52),X53)),'type.option.option'(X52)))) )),
  inference(superposition,[],[f716,f946])).
fof(f28361,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X51,X49] : (s(X49,happ(s(fun(X44,X49),X50),s(X44,sK7(happ(s(fun(X45,'type.option.option'(X44)),X46),s(X45,sK7(X47,X45))),X44)))) = s(X49,'const.option.option_CASE_3'(s('type.option.option'(X44),'const.option.option_CASE_3'(s('type.option.option'(X45),X47),s('type.option.option'(X44),X48),s(fun(X45,'type.option.option'(X44)),X46))),s(X49,X51),s(fun(X44,X49),X50))) | s('type.option.option'(X44),'const.option.NONE_0') = s('type.option.option'(X44),happ(s(fun(X45,'type.option.option'(X44)),X46),s(X45,sK7(X47,X45)))) | s('type.option.option'(X45),'const.option.NONE_0') = s('type.option.option'(X45),X47)) )),
  inference(superposition,[],[f716,f716])).
fof(f28702,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s('type.option.option'(X36),X39),s(fun(X37,'type.option.option'(X36)),X40))) | s(X41,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s('type.option.option'(X36),X39),s(fun(X37,'type.option.option'(X36)),X40))),s(X41,X43),s(fun(X36,X41),X42))) = s(X41,happ(s(fun(X36,X41),X42),s(X36,sK7(happ(s(fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)),happ(s(fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36))),happ(s(fun('type.option.option'(X37),fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s('type.option.option'(X36),X39))),s(fun(X37,'type.option.option'(X36)),X40)),X36))))) )),
  inference(forward_demodulation,[],[f28360,f166])).
fof(f28360,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s(X41,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s('type.option.option'(X36),X39),s(fun(X37,'type.option.option'(X36)),X40))),s(X41,X43),s(fun(X36,X41),X42))) = s(X41,happ(s(fun(X36,X41),X42),s(X36,sK7(happ(s(fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)),happ(s(fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36))),happ(s(fun('type.option.option'(X37),fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s('type.option.option'(X36),X39))),s(fun(X37,'type.option.option'(X36)),X40)),X36)))) | s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),happ(s(fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)),happ(s(fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36))),happ(s(fun('type.option.option'(X37),fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s('type.option.option'(X36),X39))),s(fun(X37,'type.option.option'(X36)),X40)))) )),
  inference(superposition,[],[f716,f166])).
fof(f28359,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s(X33,'const.option.option_CASE_3'(s('type.option.option'(X28),happ(s(fun(X29,'type.option.option'(X28)),X32),s(X29,sK7(X30,X29)))),s(X33,X35),s(fun(X28,X33),X34))) = s(X33,happ(s(fun(X28,X33),X34),s(X28,sK7('const.option.option_CASE_3'(s('type.option.option'(X29),X30),s('type.option.option'(X28),X31),s(fun(X29,'type.option.option'(X28)),X32)),X28)))) | s('type.option.option'(X28),'const.option.NONE_0') = s('type.option.option'(X28),'const.option.option_CASE_3'(s('type.option.option'(X29),X30),s('type.option.option'(X28),X31),s(fun(X29,'type.option.option'(X28)),X32))) | s('type.option.option'(X29),'const.option.NONE_0') = s('type.option.option'(X29),X30)) )),
  inference(superposition,[],[f716,f716])).
fof(f28701,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s('type.option.option'(X20),'const.option.NONE_0') = s('type.option.option'(X20),happ(s(fun(X21,'type.option.option'(X20)),X24),s(X21,X22))) | s(X25,'const.option.option_CASE_3'(s('type.option.option'(X20),happ(s(fun(X21,'type.option.option'(X20)),X24),s(X21,X22))),s(X25,X27),s(fun(X20,X25),X26))) = s(X25,happ(s(fun(X20,X25),X26),s(X20,sK7('const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X22))),s('type.option.option'(X20),X23),s(fun(X21,'type.option.option'(X20)),X24)),X20))))) )),
  inference(forward_demodulation,[],[f28358,f165])).
fof(f28358,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(X25,'const.option.option_CASE_3'(s('type.option.option'(X20),happ(s(fun(X21,'type.option.option'(X20)),X24),s(X21,X22))),s(X25,X27),s(fun(X20,X25),X26))) = s(X25,happ(s(fun(X20,X25),X26),s(X20,sK7('const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X22))),s('type.option.option'(X20),X23),s(fun(X21,'type.option.option'(X20)),X24)),X20)))) | s('type.option.option'(X20),'const.option.NONE_0') = s('type.option.option'(X20),'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X22))),s('type.option.option'(X20),X23),s(fun(X21,'type.option.option'(X20)),X24)))) )),
  inference(superposition,[],[f716,f165])).
fof(f28700,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s('type.option.option'(X13),'const.option.NONE_0') = s('type.option.option'(X13),X15) | s(X17,'const.option.option_CASE_3'(s('type.option.option'(X13),X15),s(X17,X19),s(fun(X13,X17),X18))) = s(X17,happ(s(fun(X13,X17),X18),s(X13,sK7('const.option.option_CASE_3'(s('type.option.option'(X14),'const.option.NONE_0'),s('type.option.option'(X13),X15),s(fun(X14,'type.option.option'(X13)),X16)),X13))))) )),
  inference(forward_demodulation,[],[f28357,f163])).
fof(f28357,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s(X17,'const.option.option_CASE_3'(s('type.option.option'(X13),X15),s(X17,X19),s(fun(X13,X17),X18))) = s(X17,happ(s(fun(X13,X17),X18),s(X13,sK7('const.option.option_CASE_3'(s('type.option.option'(X14),'const.option.NONE_0'),s('type.option.option'(X13),X15),s(fun(X14,'type.option.option'(X13)),X16)),X13)))) | s('type.option.option'(X13),'const.option.NONE_0') = s('type.option.option'(X13),'const.option.option_CASE_3'(s('type.option.option'(X14),'const.option.NONE_0'),s('type.option.option'(X13),X15),s(fun(X14,'type.option.option'(X13)),X16)))) )),
  inference(superposition,[],[f716,f163])).
fof(f28353,plain,(
  ( ! [X80,X81,X79,X82] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X79),X81),s('type.option.option'(bool),X82),s(fun(X79,'type.option.option'(bool)),X80))) | s('type.option.option'(X79),'const.option.NONE_0') = s('type.option.option'(X79),X81) | s(bool,'T_0') = s(bool,sK7(happ(s(fun(X79,'type.option.option'(bool)),X80),s(X79,sK7(X81,X79))),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),happ(s(fun(X79,'type.option.option'(bool)),X80),s(X79,sK7(X81,X79))))) )),
  inference(superposition,[],[f716,f610])).
fof(f28340,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,X24) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X21),X23),s(bool,X25),s(fun(X21,bool),X22))) | s('type.option.option'(X21),'const.option.NONE_0') = s('type.option.option'(X21),X23) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,happ(s(fun(X21,bool),X22),s(X21,sK7(X23,X21))))) )),
  inference(superposition,[],[f716,f190])).
fof(f28339,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X16),X18),s(bool,X19),s(fun(X16,bool),X17))) | s('type.option.option'(X16),'const.option.NONE_0') = s('type.option.option'(X16),X18) | s(bool,X20) = s(bool,happ(s(fun(X16,bool),X17),s(X16,sK7(X18,X16)))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f716,f190])).
fof(f28338,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X12),X14),s(bool,X15),s(fun(X12,bool),X13))) | s('type.option.option'(X12),'const.option.NONE_0') = s('type.option.option'(X12),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK7(X14,X12))))) )),
  inference(superposition,[],[f716,f114])).
fof(f28688,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X9),s(X6,X10),s(fun(X7,X6),X8))) = s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X9),s(X6,X11),s(fun(X7,X6),X8))) | s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X9)) )),
  inference(duplicate_literal_removal,[],[f28337])).
fof(f28337,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X9),s(X6,X10),s(fun(X7,X6),X8))) = s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X9),s(X6,X11),s(fun(X7,X6),X8))) | s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X9) | s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X9)) )),
  inference(superposition,[],[f716,f716])).
fof(f28336,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X2),s(X0,X3),s(fun(X1,X0),sK7(X4,fun(X1,X0))))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(fun(X1,X0)),X4),s(X0,X5),s(fun(fun(X1,X0),X0),happ(s(fun(X0,fun(fun(X1,X0),X0)),happ(s(fun('type.option.option'(X1),fun(X0,fun(fun(X1,X0),X0))),'const.option.option_CASE_0'),s('type.option.option'(X1),X2))),s(X0,X3))))) | s('type.option.option'(fun(X1,X0)),'const.option.NONE_0') = s('type.option.option'(fun(X1,X0)),X4)) )),
  inference(superposition,[],[f716,f166])).
fof(f27880,plain,(
  ( ! [X87,X85,X83,X86,X84,X82] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X82),X83),s(bool,X84),s(fun(X82,bool),X85)))) | s(fun(X82,bool),X85) = s(fun(X82,bool),sK8(X85,'const.option.option_CASE_3'(s('type.option.option'(X86),'const.option.NONE_0'),s(fun(fun(X82,bool),bool),happ(s(fun(bool,fun(fun(X82,bool),bool)),happ(s(fun('type.option.option'(X82),fun(bool,fun(fun(X82,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X82),X83))),s(bool,X84))),s(fun(X86,fun(fun(X82,bool),bool)),X87)),fun(X82,bool)))) )),
  inference(superposition,[],[f568,f166])).
fof(f27937,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(X5,X7) = s(X5,sK8(X7,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(fun(X5,bool),X6),s(fun(X8,fun(X5,bool)),X9)),X5)) | s(bool,X10) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f27866,f111])).
fof(f27866,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | s(X5,X7) = s(X5,sK8(X7,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(fun(X5,bool),X6),s(fun(X8,fun(X5,bool)),X9)),X5)) | s(bool,X10) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f568,f190])).
fof(f27863,plain,(
  ( ! [X246,X244,X247,X245] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X245),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),X244) = s('type.option.option'(bool),sK8(X244,'const.option.option_CASE_3'(s('type.option.option'(X246),'const.option.NONE_0'),s(fun('type.option.option'(bool),bool),X245),s(fun(X246,fun('type.option.option'(bool),bool)),X247)),'type.option.option'(bool))) | s(bool,'T_0') = s(bool,sK7(X244,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X244)) )),
  inference(superposition,[],[f568,f610])).
fof(f27862,plain,(
  ( ! [X243,X241,X239,X242,X240] : (~p(s(bool,happ(s(fun('type.option.option'(X239),bool),X241),s('type.option.option'(X239),X240)))) | s('type.option.option'(X239),'const.option.SOME_1'(s(X239,sK7(X240,X239)))) = s('type.option.option'(X239),sK8('const.option.SOME_1'(s(X239,sK7(X240,X239))),'const.option.option_CASE_3'(s('type.option.option'(X242),'const.option.NONE_0'),s(fun('type.option.option'(X239),bool),X241),s(fun(X242,fun('type.option.option'(X239),bool)),X243)),'type.option.option'(X239))) | s('type.option.option'(X239),'const.option.NONE_0') = s('type.option.option'(X239),X240)) )),
  inference(superposition,[],[f568,f146])).
fof(f27861,plain,(
  ( ! [X237,X235,X238,X236] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X236),s('type.option.option'(bool),X235)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,'F_0')),'const.option.option_CASE_3'(s('type.option.option'(X237),'const.option.NONE_0'),s(fun('type.option.option'(bool),bool),X236),s(fun(X237,fun('type.option.option'(bool),bool)),X238)),'type.option.option'(bool))) | s(bool,'T_0') = s(bool,sK7(X235,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X235)) )),
  inference(superposition,[],[f568,f610])).
fof(f27860,plain,(
  ( ! [X231,X233,X234,X232] : (~p(s(bool,happ(s(fun(bool,bool),X232),s(bool,'T_0')))) | s(bool,sK7(X231,bool)) = s(bool,sK8(sK7(X231,bool),'const.option.option_CASE_3'(s('type.option.option'(X233),'const.option.NONE_0'),s(fun(bool,bool),X232),s(fun(X233,fun(bool,bool)),X234)),bool)) | s('type.option.option'(bool),X231) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X231)) )),
  inference(superposition,[],[f568,f610])).
fof(f27925,plain,(
  ( ! [X198,X196,X202,X200,X199,X197,X201] : (s(bool,happ(s(fun(X196,bool),X199),s(X196,X197))) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X196),'const.option.SOME_1'(s(X196,X197))),s(bool,X198),s(fun(X196,bool),X199)),'const.option.option_CASE_3'(s('type.option.option'(X201),'const.option.NONE_0'),s(fun(bool,bool),X200),s(fun(X201,fun(bool,bool)),X202)),bool)) | ~p(s(bool,happ(s(fun(bool,bool),X200),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),X199),s(X196,X197)))) )),
  inference(forward_demodulation,[],[f27853,f165])).
fof(f27853,plain,(
  ( ! [X198,X196,X202,X200,X199,X197,X201] : (~p(s(bool,happ(s(fun(bool,bool),X200),s(bool,'T_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X196),'const.option.SOME_1'(s(X196,X197))),s(bool,X198),s(fun(X196,bool),X199))) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X196),'const.option.SOME_1'(s(X196,X197))),s(bool,X198),s(fun(X196,bool),X199)),'const.option.option_CASE_3'(s('type.option.option'(X201),'const.option.NONE_0'),s(fun(bool,bool),X200),s(fun(X201,fun(bool,bool)),X202)),bool)) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),X199),s(X196,X197)))) )),
  inference(superposition,[],[f568,f737])).
fof(f27924,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (s(bool,X191) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X190),'const.option.NONE_0'),s(bool,X191),s(fun(X190,bool),X192)),'const.option.option_CASE_3'(s('type.option.option'(X194),'const.option.NONE_0'),s(fun(bool,bool),X193),s(fun(X194,fun(bool,bool)),X195)),bool)) | ~p(s(bool,happ(s(fun(bool,bool),X193),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X191)) )),
  inference(forward_demodulation,[],[f27852,f163])).
fof(f27852,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (~p(s(bool,happ(s(fun(bool,bool),X193),s(bool,'T_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X190),'const.option.NONE_0'),s(bool,X191),s(fun(X190,bool),X192))) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X190),'const.option.NONE_0'),s(bool,X191),s(fun(X190,bool),X192)),'const.option.option_CASE_3'(s('type.option.option'(X194),'const.option.NONE_0'),s(fun(bool,bool),X193),s(fun(X194,fun(bool,bool)),X195)),bool)) | s(bool,'F_0') = s(bool,X191)) )),
  inference(superposition,[],[f568,f239])).
fof(f27828,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(bool,bool),X96),s(bool,X95)))) | s(bool,'T_0') = s(bool,sK8('T_0','const.option.option_CASE_3'(s('type.option.option'(X97),'const.option.NONE_0'),s(fun(bool,bool),X96),s(fun(X97,fun(bool,bool)),X98)),bool)) | s(bool,X95) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f568,f190])).
fof(f27815,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,X37) = s(bool,sK8(X37,'const.option.option_CASE_3'(s('type.option.option'(X40),'const.option.NONE_0'),s(fun(bool,bool),X39),s(fun(X40,fun(bool,bool)),X41)),bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f568,f190])).
fof(f27814,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'T_0')))) | s(bool,X32) = s(bool,sK8(X32,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.NONE_0'),s(fun(bool,bool),X33),s(fun(X34,fun(bool,bool)),X35)),bool)) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f568,f190])).
fof(f27888,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(X23,X24) = s(X23,sK8(sK7('const.option.SOME_1'(s(X23,X24)),X23),'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(fun(X23,bool),X25),s(fun(X26,fun(X23,bool)),X27)),X23)) | ~p(s(bool,happ(s(fun(X23,bool),X25),s(X23,X24))))) )),
  inference(forward_demodulation,[],[f27812,f946])).
fof(f27812,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(X23,bool),X25),s(X23,X24)))) | s(X23,sK7('const.option.SOME_1'(s(X23,X24)),X23)) = s(X23,sK8(sK7('const.option.SOME_1'(s(X23,X24)),X23),'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(fun(X23,bool),X25),s(fun(X26,fun(X23,bool)),X27)),X23))) )),
  inference(superposition,[],[f568,f946])).
fof(f27887,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(X15,X18),s(fun(X16,X15),X19))) = s(X15,sK8(happ(s(fun(fun(X16,X15),X15),happ(s(fun(X15,fun(fun(X16,X15),X15)),happ(s(fun('type.option.option'(X16),fun(X15,fun(fun(X16,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(X15,X18))),s(fun(X16,X15),X19)),'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.NONE_0'),s(fun(X15,bool),X20),s(fun(X21,fun(X15,bool)),X22)),X15)) | ~p(s(bool,happ(s(fun(X15,bool),X20),s(X15,'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(X15,X18),s(fun(X16,X15),X19))))))) )),
  inference(forward_demodulation,[],[f27811,f166])).
fof(f27811,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun(X15,bool),X20),s(X15,'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(X15,X18),s(fun(X16,X15),X19)))))) | s(X15,happ(s(fun(fun(X16,X15),X15),happ(s(fun(X15,fun(fun(X16,X15),X15)),happ(s(fun('type.option.option'(X16),fun(X15,fun(fun(X16,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(X15,X18))),s(fun(X16,X15),X19))) = s(X15,sK8(happ(s(fun(fun(X16,X15),X15),happ(s(fun(X15,fun(fun(X16,X15),X15)),happ(s(fun('type.option.option'(X16),fun(X15,fun(fun(X16,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(X15,X18))),s(fun(X16,X15),X19)),'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.NONE_0'),s(fun(X15,bool),X20),s(fun(X21,fun(X15,bool)),X22)),X15))) )),
  inference(superposition,[],[f568,f166])).
fof(f27886,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s(X7,happ(s(fun(X8,X7),X11),s(X8,X9))) = s(X7,sK8('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(X7,X10),s(fun(X8,X7),X11)),'const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.NONE_0'),s(fun(X7,bool),X12),s(fun(X13,fun(X7,bool)),X14)),X7)) | ~p(s(bool,happ(s(fun(X7,bool),X12),s(X7,happ(s(fun(X8,X7),X11),s(X8,X9))))))) )),
  inference(forward_demodulation,[],[f27810,f165])).
fof(f27810,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,happ(s(fun(X7,bool),X12),s(X7,happ(s(fun(X8,X7),X11),s(X8,X9)))))) | s(X7,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(X7,X10),s(fun(X8,X7),X11))) = s(X7,sK8('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(X7,X10),s(fun(X8,X7),X11)),'const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.NONE_0'),s(fun(X7,bool),X12),s(fun(X13,fun(X7,bool)),X14)),X7))) )),
  inference(superposition,[],[f568,f165])).
fof(f27060,plain,(
  ( ! [X389,X391,X393,X390,X392] : (s('type.option.option'(bool),X393) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK7(X393,bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X389),'const.option.SOME_1'(s(X389,X390))),s(bool,X391),s(fun(X389,bool),X392))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X393) | s(bool,'F_0') = s(bool,happ(s(fun(X389,bool),X392),s(X389,X390)))) )),
  inference(superposition,[],[f611,f737])).
fof(f27058,plain,(
  ( ! [X377,X379,X380,X376,X378,X375,X381] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X379))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X380),'const.option.NONE_0'),s(bool,X379),s(fun(X380,bool),X381))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X375),'const.option.SOME_1'(s(X375,X376))),s(bool,X377),s(fun(X375,bool),X378))) | s(bool,'F_0') = s(bool,happ(s(fun(X375,bool),X378),s(X375,X376)))) )),
  inference(superposition,[],[f491,f737])).
fof(f27202,plain,(
  ( ! [X333,X335,X332,X334,X336] : (s(bool,happ(s(fun(X332,bool),X335),s(X332,X333))) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X332),'const.option.SOME_1'(s(X332,X333))),s(bool,X334),s(fun(X332,bool),X335)),X336,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X336),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,happ(s(fun(X332,bool),X335),s(X332,X333)))) )),
  inference(forward_demodulation,[],[f27051,f165])).
fof(f27051,plain,(
  ( ! [X333,X335,X332,X334,X336] : (~p(s(bool,happ(s(fun(bool,bool),X336),s(bool,'T_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X332),'const.option.SOME_1'(s(X332,X333))),s(bool,X334),s(fun(X332,bool),X335))) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X332),'const.option.SOME_1'(s(X332,X333))),s(bool,X334),s(fun(X332,bool),X335)),X336,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(X332,bool),X335),s(X332,X333)))) )),
  inference(superposition,[],[f157,f737])).
fof(f27046,plain,(
  ( ! [X308,X310,X306,X309,X307] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X310))) | s(bool,X310) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X306),'const.option.SOME_1'(s(X306,X307))),s(bool,X308),s(fun(X306,bool),X309))) | s(bool,'F_0') = s(bool,happ(s(fun(X306,bool),X309),s(X306,X307)))) )),
  inference(superposition,[],[f154,f737])).
fof(f26997,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X95),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X91),'const.option.SOME_1'(s(X91,X92))),s(bool,X93),s(fun(X91,bool),X94)),X95,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(X91,bool),X94),s(X91,X92)))) )),
  inference(superposition,[],[f675,f737])).
fof(f26927,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,X67) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X68),'const.option.SOME_1'(s(X68,X69))),s(bool,X67),s(fun(X68,bool),X70))) | s(bool,'F_0') = s(bool,happ(s(fun(X68,bool),X70),s(X68,X69))) | s(bool,X67) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f737,f190])).
fof(f26910,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X3),X4),s(bool,X5),s(fun(X3,bool),X6))) | s(bool,'F_0') = s(bool,happ(s(fun(X3,bool),X6),s(X3,sK7(X4,X3)))) | s('type.option.option'(X3),'const.option.NONE_0') = s('type.option.option'(X3),X4)) )),
  inference(superposition,[],[f737,f146])).
fof(f26909,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(bool),X0),s(bool,X1),s(fun(bool,bool),X2))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X0,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f737,f610])).
fof(f26907,plain,(
  ( ! [X185,X186,X184] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))),s(bool,X185),s(fun('type.option.option'(bool),bool),X186))) | s(bool,'F_0') = s(bool,happ(s(fun('type.option.option'(bool),bool),X186),s('type.option.option'(bool),X184))) | s(bool,'T_0') = s(bool,sK7(X184,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X184)) )),
  inference(superposition,[],[f737,f610])).
fof(f26906,plain,(
  ( ! [X182,X180,X183,X181] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.option.option'(X180)),'const.option.SOME_1'(s('type.option.option'(X180),X181))),s(bool,X182),s(fun('type.option.option'(X180),bool),X183))) | s(bool,'F_0') = s(bool,happ(s(fun('type.option.option'(X180),bool),X183),s('type.option.option'(X180),'const.option.SOME_1'(s(X180,sK7(X181,X180)))))) | s('type.option.option'(X180),'const.option.NONE_0') = s('type.option.option'(X180),X181)) )),
  inference(superposition,[],[f737,f146])).
fof(f26905,plain,(
  ( ! [X177,X178,X179] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X177))),s(bool,X178),s(fun('type.option.option'(bool),bool),X179))) | s(bool,'F_0') = s(bool,happ(s(fun('type.option.option'(bool),bool),X179),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK7(X177,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X177)) )),
  inference(superposition,[],[f737,f610])).
fof(f26904,plain,(
  ( ! [X175,X176,X174] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(bool,X175),s(fun(bool,bool),X176))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X176),s(bool,sK7(X174,bool)))) | s('type.option.option'(bool),X174) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X174)) )),
  inference(superposition,[],[f737,f610])).
fof(f26876,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X84))),s(bool,X85),s(fun(bool,bool),X86))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X86),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f737,f114])).
fof(f26874,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X76))),s(bool,X77),s(fun(bool,bool),X78))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X78),s(bool,'T_0'))) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f737,f190])).
fof(f26861,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))),s(bool,X33),s(fun(bool,bool),X34))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X34),s(bool,X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f737,f190])).
fof(f26860,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(bool,X28),s(fun(bool,bool),X29))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X29),s(bool,X27))) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f737,f190])).
fof(f26859,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(bool,X25),s(fun(bool,bool),X26))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X26),s(bool,X24))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f737,f114])).
fof(f965,plain,(
  ( ! [X8] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X8)),bool))) )),
  inference(superposition,[],[f603,f114])).
fof(f2248,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,X33) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.NONE_0'),s(bool,X31),s(fun(X30,bool),X32))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f497,f163])).
fof(f2249,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X34,bool),X37),s(X34,X35))))) | s(bool,X38) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.SOME_1'(s(X34,X35))),s(bool,X36),s(fun(X34,bool),X37))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f497,f165])).
fof(f2250,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X39),X40),s(bool,X41),s(fun(X39,bool),X42))))) | s(bool,X43) = s(bool,happ(s(fun(fun(X39,bool),bool),happ(s(fun(bool,fun(fun(X39,bool),bool)),happ(s(fun('type.option.option'(X39),fun(bool,fun(fun(X39,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X39),X40))),s(bool,X41))),s(fun(X39,bool),X42))) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f497,f166])).
fof(f2251,plain,(
  ( ! [X45,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,X45) = s(bool,sK7('const.option.SOME_1'(s(bool,X44)),bool)) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f497,f946])).
fof(f3067,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X42) != s(bool,X44) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.NONE_0'),s(bool,X42),s(fun(X41,bool),X43)))) )),
  inference(superposition,[],[f532,f163])).
fof(f3068,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,X49) != s(bool,happ(s(fun(X45,bool),X48),s(X45,X46))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X46))),s(bool,X47),s(fun(X45,bool),X48)))) )),
  inference(superposition,[],[f532,f165])).
fof(f3069,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),X51),s(bool,X52),s(fun(X50,bool),X53))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53)))) )),
  inference(superposition,[],[f532,f166])).
fof(f3070,plain,(
  ( ! [X56,X55] : (s(bool,X55) != s(bool,X56) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X55)),bool))) )),
  inference(superposition,[],[f532,f946])).
fof(f3229,plain,(
  ( ! [X28,X26,X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(bool,X27),s(fun(X26,bool),X28)))) )),
  inference(superposition,[],[f3099,f163])).
fof(f3230,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X29,bool),X32),s(X29,X30))))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X29),'const.option.SOME_1'(s(X29,X30))),s(bool,X31),s(fun(X29,bool),X32)))) )),
  inference(superposition,[],[f3099,f165])).
fof(f3231,plain,(
  ( ! [X35,X33,X36,X34] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X33),X34),s(bool,X35),s(fun(X33,bool),X36))))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X33,bool),bool),happ(s(fun(bool,fun(fun(X33,bool),bool)),happ(s(fun('type.option.option'(X33),fun(bool,fun(fun(X33,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X33),X34))),s(bool,X35))),s(fun(X33,bool),X36)))) )),
  inference(superposition,[],[f3099,f166])).
fof(f3248,plain,(
  ( ! [X26,X27] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X27))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))))) | s('type.option.option'(bool),X27) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f154,f3099])).
fof(f3250,plain,(
  ( ! [X30,X31] : (s('type.option.option'(bool),X31) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X31))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f155,f3099])).
fof(f3252,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X35),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34)))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X35),s('type.option.option'(bool),X36)))) | s('type.option.option'(bool),X36) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f156,f3099])).
fof(f3253,plain,(
  ( ! [X37,X38] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X38),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,'F_0')),X38,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f157,f3099])).
fof(f3267,plain,(
  ( ! [X12,X10,X13,X11] : (s(X11,happ(s(fun(bool,X11),X12),s(bool,X10))) = s(X11,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(X11,X13),s(fun(bool,X11),X12))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f165,f3099])).
fof(f3274,plain,(
  ( ! [X26,X25] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X26))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X26) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f154,f3099])).
fof(f3276,plain,(
  ( ! [X30,X29] : (s('type.option.option'(bool),X30) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X30))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f155,f3099])).
fof(f3278,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X34),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X34),s('type.option.option'(bool),X35)))) | s('type.option.option'(bool),X35) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f156,f3099])).
fof(f3279,plain,(
  ( ! [X37,X36] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X37),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,X36)),X37,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f157,f3099])).
fof(f9352,plain,(
  ( ! [X30,X29] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X30))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X29) = s('type.option.option'(bool),X30) | s(bool,'T_0') = s(bool,sK7(X29,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X29)) )),
  inference(superposition,[],[f154,f610])).
fof(f9354,plain,(
  ( ! [X33,X34] : (s('type.option.option'(bool),X34) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X33))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X34))) | s(bool,'T_0') = s(bool,sK7(X33,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X33)) )),
  inference(superposition,[],[f155,f610])).
fof(f9356,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X38),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X38),s('type.option.option'(bool),X39)))) | s('type.option.option'(bool),X37) != s('type.option.option'(bool),X39) | s(bool,'T_0') = s(bool,sK7(X37,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X37)) )),
  inference(superposition,[],[f156,f610])).
fof(f9357,plain,(
  ( ! [X41,X40] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X41),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),X40) = s('type.option.option'(bool),sK8(X40,X41,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,sK7(X40,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X40)) )),
  inference(superposition,[],[f157,f610])).
fof(f9364,plain,(
  ( ! [X64,X62,X63] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),f(s(X62,sK9(X63,X64,'type.option.option'(bool),X62)))) | s(fun(X62,'type.option.option'(bool)),X63) = s(fun(X62,'type.option.option'(bool)),X64) | s(bool,'T_0') = s(bool,sK7(g(s(X62,sK9(X63,X64,'type.option.option'(bool),X62))),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),g(s(X62,sK9(X63,X64,'type.option.option'(bool),X62))))) )),
  inference(superposition,[],[f164,f610])).
fof(f9368,plain,(
  ( ! [X76] : (s('type.option.option'('type.option.option'(bool)),X76) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X76) | s(bool,'T_0') = s(bool,sK7(sK7(X76,'type.option.option'(bool)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),sK7(X76,'type.option.option'(bool)))) )),
  inference(superposition,[],[f146,f610])).
fof(f9370,plain,(
  ( ! [X78,X79] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X79),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X79),s('type.option.option'(bool),X78)))) | s(bool,'T_0') = s(bool,sK7(sK8(X78,X79,'type.option.option'(bool)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),sK8(X78,X79,'type.option.option'(bool)))) )),
  inference(superposition,[],[f158,f610])).
fof(f9371,plain,(
  ( ! [X80,X81,X82] : (s(X82,f(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) != s(X82,g(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s(fun('type.option.option'(bool),X82),X80) = s(fun('type.option.option'(bool),X82),X81) | s(bool,'T_0') = s(bool,sK7(sK9(X80,X81,X82,'type.option.option'(bool)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),sK9(X80,X81,X82,'type.option.option'(bool)))) )),
  inference(superposition,[],[f164,f610])).
fof(f9387,plain,(
  ( ! [X31,X32] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X31))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X32))) | s('type.option.option'(bool),X32) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X31,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X31)) )),
  inference(superposition,[],[f154,f610])).
fof(f9389,plain,(
  ( ! [X35,X36] : (s('type.option.option'(bool),X35) != s('type.option.option'(bool),X36) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X36))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK7(X35,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X35)) )),
  inference(superposition,[],[f155,f610])).
fof(f9391,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X40),s('type.option.option'(bool),X39)))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X40),s('type.option.option'(bool),X41)))) | s('type.option.option'(bool),X41) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X39,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X39)) )),
  inference(superposition,[],[f156,f610])).
fof(f9392,plain,(
  ( ! [X43,X42] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X43),s('type.option.option'(bool),X42)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,'F_0')),X43,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,sK7(X42,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X42)) )),
  inference(superposition,[],[f157,f610])).
fof(f9462,plain,(
  ( ! [X105,X106] : (s(bool,'T_0') != s(bool,X106) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X106))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7(X105,bool)))) | s('type.option.option'(bool),X105) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X105)) )),
  inference(superposition,[],[f155,f610])).
fof(f9464,plain,(
  ( ! [X111,X109,X110] : (~p(s(bool,happ(s(fun(bool,bool),X110),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X110),s(bool,X111)))) | s(bool,X111) != s(bool,sK7(X109,bool)) | s('type.option.option'(bool),X109) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X109)) )),
  inference(superposition,[],[f156,f610])).
fof(f9465,plain,(
  ( ! [X112,X113] : (~p(s(bool,happ(s(fun(bool,bool),X113),s(bool,'T_0')))) | s(bool,sK7(X112,bool)) = s(bool,sK8(sK7(X112,bool),X113,bool)) | s('type.option.option'(bool),X112) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X112)) )),
  inference(superposition,[],[f157,f610])).
fof(f9491,plain,(
  ( ! [X12,X10,X11] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X11) | s('type.option.option'(bool),X11) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X10),'const.option.NONE_0'),s('type.option.option'(bool),X11),s(fun(X10,'type.option.option'(bool)),X12)),bool))) )),
  inference(forward_demodulation,[],[f9296,f163])).
fof(f9296,plain,(
  ( ! [X12,X10,X11] : (s('type.option.option'(bool),X11) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X10),'const.option.NONE_0'),s('type.option.option'(bool),X11),s(fun(X10,'type.option.option'(bool)),X12)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X10),'const.option.NONE_0'),s('type.option.option'(bool),X11),s(fun(X10,'type.option.option'(bool)),X12)))) )),
  inference(superposition,[],[f610,f163])).
fof(f9492,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),happ(s(fun(X13,'type.option.option'(bool)),X16),s(X13,X14))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),happ(s(fun(X13,'type.option.option'(bool)),X16),s(X13,X14))) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X14))),s('type.option.option'(bool),X15),s(fun(X13,'type.option.option'(bool)),X16)),bool))) )),
  inference(forward_demodulation,[],[f9297,f165])).
fof(f9297,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),happ(s(fun(X13,'type.option.option'(bool)),X16),s(X13,X14))) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X14))),s('type.option.option'(bool),X15),s(fun(X13,'type.option.option'(bool)),X16)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X14))),s('type.option.option'(bool),X15),s(fun(X13,'type.option.option'(bool)),X16)))) )),
  inference(superposition,[],[f610,f165])).
fof(f9493,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X17),X18),s('type.option.option'(bool),X19),s(fun(X17,'type.option.option'(bool)),X20))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X17),X18),s('type.option.option'(bool),X19),s(fun(X17,'type.option.option'(bool)),X20))) | s(bool,'T_0') = s(bool,sK7(happ(s(fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)),happ(s(fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool))),happ(s(fun('type.option.option'(X17),fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)))),'const.option.option_CASE_0'),s('type.option.option'(X17),X18))),s('type.option.option'(bool),X19))),s(fun(X17,'type.option.option'(bool)),X20)),bool))) )),
  inference(forward_demodulation,[],[f9298,f166])).
fof(f9298,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X17),X18),s('type.option.option'(bool),X19),s(fun(X17,'type.option.option'(bool)),X20))) | s(bool,'T_0') = s(bool,sK7(happ(s(fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)),happ(s(fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool))),happ(s(fun('type.option.option'(X17),fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)))),'const.option.option_CASE_0'),s('type.option.option'(X17),X18))),s('type.option.option'(bool),X19))),s(fun(X17,'type.option.option'(bool)),X20)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),happ(s(fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)),happ(s(fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool))),happ(s(fun('type.option.option'(X17),fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)))),'const.option.option_CASE_0'),s('type.option.option'(X17),X18))),s('type.option.option'(bool),X19))),s(fun(X17,'type.option.option'(bool)),X20)))) )),
  inference(superposition,[],[f610,f166])).
fof(f9494,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X21) | s('type.option.option'(bool),X21) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(sK7('const.option.SOME_1'(s('type.option.option'(bool),X21)),'type.option.option'(bool)),bool))) )),
  inference(forward_demodulation,[],[f9299,f946])).
fof(f9299,plain,(
  ( ! [X21] : (s('type.option.option'(bool),X21) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(sK7('const.option.SOME_1'(s('type.option.option'(bool),X21)),'type.option.option'(bool)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),sK7('const.option.SOME_1'(s('type.option.option'(bool),X21)),'type.option.option'(bool)))) )),
  inference(superposition,[],[f610,f946])).
fof(f9514,plain,(
  ( ! [X26,X25] : (~p(s(bool,X26)) | s(bool,X26) = s(bool,sK7(X25,bool)) | s('type.option.option'(bool),X25) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X25)) )),
  inference(subsumption_resolution,[],[f9416,f111])).
fof(f9416,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,sK7(X25,bool)) | s('type.option.option'(bool),X25) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X25)) )),
  inference(superposition,[],[f162,f610])).
fof(f9815,plain,(
  ( ! [X116,X115] : (s('type.option.option'('type.option.option'(bool)),X116) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X115))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK7(X116,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X116) | s(bool,'T_0') = s(bool,sK7(X115,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X115)) )),
  inference(superposition,[],[f611,f610])).
fof(f9818,plain,(
  ( ! [X123,X122] : (s('type.option.option'('type.option.option'(bool)),X123) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X122) = s('type.option.option'(bool),sK7(X123,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X123) | s(bool,'T_0') = s(bool,sK7(X122,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X122)) )),
  inference(superposition,[],[f611,f610])).
fof(f10052,plain,(
  ( ! [X30,X28,X31,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,X29) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.NONE_0'),s(bool,X28),s(fun(X30,bool),X31))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f491,f114])).
fof(f10082,plain,(
  ( ! [X171,X169,X172,X170] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X170))) | s(bool,X170) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X171),'const.option.NONE_0'),s(bool,sK7(X169,bool)),s(fun(X171,bool),X172))) | s('type.option.option'(bool),X169) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X169)) )),
  inference(superposition,[],[f491,f610])).
fof(f10088,plain,(
  ( ! [X198,X196,X197,X195] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X195))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X196))) | s('type.option.option'(bool),X196) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X197),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(fun(X197,'type.option.option'(bool)),X198))) | s(bool,'T_0') = s(bool,sK7(X195,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X195)) )),
  inference(superposition,[],[f491,f610])).
fof(f10091,plain,(
  ( ! [X210,X208,X211,X209] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X209))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X209) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X210),'const.option.NONE_0'),s('type.option.option'(bool),X208),s(fun(X210,'type.option.option'(bool)),X211))) | s(bool,'T_0') = s(bool,sK7(X208,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X208)) )),
  inference(superposition,[],[f491,f610])).
fof(f10136,plain,(
  ( ! [X171,X169,X172,X170] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X170))) | s(bool,sK7(X169,bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X171),'const.option.NONE_0'),s(bool,X170),s(fun(X171,bool),X172))) | s('type.option.option'(bool),X169) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X169)) )),
  inference(superposition,[],[f491,f610])).
fof(f10142,plain,(
  ( ! [X198,X196,X197,X195] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X195))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X196))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X197),'const.option.NONE_0'),s('type.option.option'(bool),X196),s(fun(X197,'type.option.option'(bool)),X198))) | s(bool,'T_0') = s(bool,sK7(X195,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X195)) )),
  inference(superposition,[],[f491,f610])).
fof(f10145,plain,(
  ( ! [X210,X208,X211,X209] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X209))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X208) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X210),'const.option.NONE_0'),s('type.option.option'(bool),X209),s(fun(X210,'type.option.option'(bool)),X211))) | s(bool,'T_0') = s(bool,sK7(X208,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X208)) )),
  inference(superposition,[],[f491,f610])).
fof(f10595,plain,(
  ( ! [X171,X169,X172,X170] : (~p(s(bool,happ(s(fun(bool,bool),X170),s(bool,'T_0')))) | s(bool,sK7(X169,bool)) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X171),'const.option.NONE_0'),s(bool,sK7(X169,bool)),s(fun(X171,bool),X172)),X170,bool)) | s('type.option.option'(bool),X169) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X169)) )),
  inference(superposition,[],[f585,f610])).
fof(f10601,plain,(
  ( ! [X198,X196,X197,X195] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X196),s('type.option.option'(bool),X195)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X197),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(fun(X197,'type.option.option'(bool)),X198)),X196,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,sK7(X195,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X195)) )),
  inference(superposition,[],[f585,f610])).
fof(f10604,plain,(
  ( ! [X210,X208,X211,X209] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X209),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),X208) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X210),'const.option.NONE_0'),s('type.option.option'(bool),X208),s(fun(X210,'type.option.option'(bool)),X211)),X209,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,sK7(X208,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X208)) )),
  inference(superposition,[],[f585,f610])).
fof(f10923,plain,(
  ( ! [X81,X82] : (~p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7(X81,bool),X82,bool)) | s('type.option.option'(bool),X81) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X81)) )),
  inference(superposition,[],[f675,f610])).
fof(f1982,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X35),X36),s(bool,X37),s(fun(X35,bool),X38))))) | s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X35,bool),bool),happ(s(fun(bool,fun(fun(X35,bool),bool)),happ(s(fun('type.option.option'(X35),fun(bool,fun(fun(X35,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X35),X36))),s(bool,X37))),s(fun(X35,bool),X38)))) )),
  inference(superposition,[],[f492,f166])).
fof(f1997,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X41))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s(bool,X39),s(fun(X37,bool),X40))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X37,bool),bool),happ(s(fun(bool,fun(fun(X37,bool),bool)),happ(s(fun('type.option.option'(X37),fun(bool,fun(fun(X37,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s(bool,X39))),s(fun(X37,bool),X40))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f492,f166])).
fof(f2177,plain,(
  ( ! [X28,X26,X27,X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X25),X26),s(bool,X27),s(fun(X25,bool),X28))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X25,bool),bool),happ(s(fun(bool,fun(fun(X25,bool),bool)),happ(s(fun('type.option.option'(X25),fun(bool,fun(fun(X25,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X25),X26))),s(bool,X27))),s(fun(X25,bool),X28)))) )),
  inference(superposition,[],[f1200,f166])).
fof(f1397,plain,(
  ( ! [X74,X75,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X74),s(sK0,sK2))))) | s(bool,X73) = s(bool,X75) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X74,sK0)) | s(bool,'T_0') = s(bool,X73)) ) | $spl0),
  inference(superposition,[],[f154,f928])).
fof(f1413,plain,(
  ( ! [X138,X137] : (s('type.option.option'(bool),X137) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X138),s(sK0,sK2))))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X138,sK0)) | s(bool,'T_0') = s(bool,sK7(X137,bool))) ) | $spl0),
  inference(superposition,[],[f146,f928])).
fof(f1438,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) != s(bool,X58) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X56,sK0)) | s(bool,'T_0') = s(bool,X57)) ) | $spl0),
  inference(superposition,[],[f155,f928])).
fof(f2347,plain,(
  ( ! [X97,X95,X96] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X97))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X96),s(sK0,sK2))))) | s(bool,X95) = s(bool,X97) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X96,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X95)),bool))) ) | $spl0),
  inference(superposition,[],[f154,f1348])).
fof(f2363,plain,(
  ( ! [X159,X160] : (s('type.option.option'(bool),X159) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X160),s(sK0,sK2))))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X159) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X160,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X159,bool))),bool))) ) | $spl0),
  inference(superposition,[],[f146,f1348])).
fof(f2395,plain,(
  ( ! [X78,X76,X77] : (s(bool,X77) != s(bool,X78) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X78))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X76),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X76,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X77)),bool))) ) | $spl0),
  inference(superposition,[],[f155,f1348])).
fof(f3188,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0))) ) | $spl0),
  inference(condensation,[],[f3187])).
fof(f3187,plain,(
  ( ! [X33,X31,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X33) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X31,sK0)) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f3176,f114])).
fof(f3176,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') != s(bool,X32) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X33) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X31,sK0)) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f537,f928])).
fof(f1981,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X30,bool),X33),s(X30,X31))))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(bool,X32),s(fun(X30,bool),X33)))) )),
  inference(superposition,[],[f492,f165])).
fof(f1996,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X32,bool),X35),s(X32,X33))))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X32),'const.option.SOME_1'(s(X32,X33))),s(bool,X34),s(fun(X32,bool),X35))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f492,f165])).
fof(f2176,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X21,bool),X24),s(X21,X22))))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X22))),s(bool,X23),s(fun(X21,bool),X24)))) )),
  inference(superposition,[],[f1200,f165])).
fof(f934,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK9(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK9(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f164,f190])).
fof(f936,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK9(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK9(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f164,f190])).
fof(f938,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK9(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK9(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f164,f114])).
fof(f1409,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,f(s(X119,sK9(X120,X121,bool,X119)))) != s(bool,happ(s(fun(sK0,bool),X122),s(sK0,sK2))) | s(fun(X119,bool),X120) = s(fun(X119,bool),X121) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X122,sK0)) | s(bool,'T_0') = s(bool,g(s(X119,sK9(X120,X121,bool,X119))))) ) | $spl0),
  inference(superposition,[],[f164,f928])).
fof(f2359,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,f(s(X141,sK9(X142,X143,bool,X141)))) != s(bool,happ(s(fun(sK0,bool),X144),s(sK0,sK2))) | s(fun(X141,bool),X142) = s(fun(X141,bool),X143) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X144,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(X141,sK9(X142,X143,bool,X141))))),bool))) ) | $spl0),
  inference(superposition,[],[f164,f1348])).
fof(f2477,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'T_0') = s(bool,X44) | s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.NONE_0'),s(bool,X44),s(fun(X43,bool),X45))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X46,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f2308,f946])).
fof(f2308,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X44)),bool)) | s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.NONE_0'),s(bool,X44),s(fun(X43,bool),X45))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X46,sK0))) ) | $spl0),
  inference(superposition,[],[f1348,f163])).
fof(f9797,plain,(
  ( ! [X70,X72,X71,X69] : (s('type.option.option'(bool),X72) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK7(X72,bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X69),'const.option.NONE_0'),s(bool,X70),s(fun(X69,bool),X71))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X72) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f611,f239])).
fof(f10124,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X116))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X113),'const.option.NONE_0'),s(bool,X114),s(fun(X113,bool),X115))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X117),'const.option.NONE_0'),s(bool,X116),s(fun(X117,bool),X118))) | s(bool,'F_0') = s(bool,X114)) )),
  inference(superposition,[],[f491,f239])).
fof(f2478,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(X47,bool),X50),s(X47,X48))) | s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),'const.option.SOME_1'(s(X47,X48))),s(bool,X49),s(fun(X47,bool),X50))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X51,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f2309,f946])).
fof(f2309,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,happ(s(fun(X47,bool),X50),s(X47,X48)))),bool)) | s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),'const.option.SOME_1'(s(X47,X48))),s(bool,X49),s(fun(X47,bool),X50))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X51,sK0))) ) | $spl0),
  inference(superposition,[],[f1348,f165])).
fof(f1016,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X5),X6),s(bool,X7),s(fun(X5,bool),X8))) | s(bool,X9) = s(bool,happ(s(fun(fun(X5,bool),bool),happ(s(fun(bool,fun(fun(X5,bool),bool)),happ(s(fun('type.option.option'(X5),fun(bool,fun(fun(X5,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X5),X6))),s(bool,X7))),s(fun(X5,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f166,f190])).
fof(f1018,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X19) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X15),X16),s(bool,X17),s(fun(X15,bool),X18))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X15,bool),bool),happ(s(fun(bool,fun(fun(X15,bool),bool)),happ(s(fun('type.option.option'(X15),fun(bool,fun(fun(X15,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X15),X16))),s(bool,X17))),s(fun(X15,bool),X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f166,f190])).
fof(f1020,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X25),X26),s(bool,X27),s(fun(X25,bool),X28))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X25,bool),bool),happ(s(fun(bool,fun(fun(X25,bool),bool)),happ(s(fun('type.option.option'(X25),fun(bool,fun(fun(X25,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X25),X26))),s(bool,X27))),s(fun(X25,bool),X28)))) )),
  inference(superposition,[],[f166,f114])).
fof(f1032,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X70),X71),s(bool,X72),s(fun(X70,bool),X73)))) | s(fun(X70,bool),X73) = s(fun(X70,bool),sK8(X73,happ(s(fun(bool,fun(fun(X70,bool),bool)),happ(s(fun('type.option.option'(X70),fun(bool,fun(fun(X70,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X70),X71))),s(bool,X72)),fun(X70,bool)))) )),
  inference(superposition,[],[f157,f166])).
fof(f1061,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X65),X66),s(bool,X67),s(fun(X65,bool),X69)))) | ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X65),X66),s(bool,X67),s(fun(X65,bool),X68)))) | s(fun(X65,bool),X68) != s(fun(X65,bool),X69)) )),
  inference(forward_demodulation,[],[f1031,f166])).
fof(f1031,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X65),X66),s(bool,X67),s(fun(X65,bool),X68)))) | p(s(bool,happ(s(fun(fun(X65,bool),bool),happ(s(fun(bool,fun(fun(X65,bool),bool)),happ(s(fun('type.option.option'(X65),fun(bool,fun(fun(X65,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X65),X66))),s(bool,X67))),s(fun(X65,bool),X69)))) | s(fun(X65,bool),X68) != s(fun(X65,bool),X69)) )),
  inference(superposition,[],[f156,f166])).
fof(f1062,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X74),X75),s(bool,X76),s(fun(X74,bool),X77)))) | p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X74),X75),s(bool,X76),s(fun(X74,bool),sK8(X77,happ(s(fun(bool,fun(fun(X74,bool),bool)),happ(s(fun('type.option.option'(X74),fun(bool,fun(fun(X74,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X74),X75))),s(bool,X76)),fun(X74,bool))))))) )),
  inference(forward_demodulation,[],[f1033,f166])).
fof(f1033,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X74),X75),s(bool,X76),s(fun(X74,bool),sK8(X77,happ(s(fun(bool,fun(fun(X74,bool),bool)),happ(s(fun('type.option.option'(X74),fun(bool,fun(fun(X74,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X74),X75))),s(bool,X76)),fun(X74,bool)))))) | ~p(s(bool,happ(s(fun(fun(X74,bool),bool),happ(s(fun(bool,fun(fun(X74,bool),bool)),happ(s(fun('type.option.option'(X74),fun(bool,fun(fun(X74,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X74),X75))),s(bool,X76))),s(fun(X74,bool),X77))))) )),
  inference(superposition,[],[f158,f166])).
fof(f1197,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X22),X23),s(bool,X24),s(fun(X22,bool),X25))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X22,bool),bool),happ(s(fun(bool,fun(fun(X22,bool),bool)),happ(s(fun('type.option.option'(X22),fun(bool,fun(fun(X22,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X22),X23))),s(bool,X24))),s(fun(X22,bool),X25)))) )),
  inference(forward_demodulation,[],[f1158,f946])).
fof(f1158,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X22),X23),s(bool,X24),s(fun(X22,bool),X25)))),bool)) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X22,bool),bool),happ(s(fun(bool,fun(fun(X22,bool),bool)),happ(s(fun('type.option.option'(X22),fun(bool,fun(fun(X22,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X22),X23))),s(bool,X24))),s(fun(X22,bool),X25)))) )),
  inference(superposition,[],[f1093,f166])).
fof(f1285,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X53),'const.option.NONE_0'),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X49),X50),s(bool,X51),s(fun(X49,bool),X52))),s(fun(X53,bool),X54))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X49,bool),bool),happ(s(fun(bool,fun(fun(X49,bool),bool)),happ(s(fun('type.option.option'(X49),fun(bool,fun(fun(X49,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X49),X50))),s(bool,X51))),s(fun(X49,bool),X52)))) )),
  inference(superposition,[],[f239,f166])).
fof(f1347,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,happ(s(fun(sK0,bool),X45),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),X42),s(bool,X43),s(fun(X41,bool),X44))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X45,sK0)) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X41,bool),bool),happ(s(fun(bool,fun(fun(X41,bool),bool)),happ(s(fun('type.option.option'(X41),fun(bool,fun(fun(X41,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X41),X42))),s(bool,X43))),s(fun(X41,bool),X44)))) ) | $spl0),
  inference(superposition,[],[f928,f166])).
fof(f2479,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X52),X53),s(bool,X54),s(fun(X52,bool),X55))) | s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK2))) = s(bool,happ(s(fun(fun(X52,bool),bool),happ(s(fun(bool,fun(fun(X52,bool),bool)),happ(s(fun('type.option.option'(X52),fun(bool,fun(fun(X52,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X52),X53))),s(bool,X54))),s(fun(X52,bool),X55))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X56,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f2310,f946])).
fof(f2310,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X52),X53),s(bool,X54),s(fun(X52,bool),X55)))),bool)) | s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK2))) = s(bool,happ(s(fun(fun(X52,bool),bool),happ(s(fun(bool,fun(fun(X52,bool),bool)),happ(s(fun('type.option.option'(X52),fun(bool,fun(fun(X52,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X52),X53))),s(bool,X54))),s(fun(X52,bool),X55))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X56,sK0))) ) | $spl0),
  inference(superposition,[],[f1348,f166])).
fof(f10620,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X80),X81),s(bool,X82),s(fun(X80,bool),X83)))) | s(fun(X80,bool),X83) = s(fun(X80,bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X84),'const.option.NONE_0'),s(fun(X80,bool),X83),s(fun(X84,fun(X80,bool)),X85)),happ(s(fun(bool,fun(fun(X80,bool),bool)),happ(s(fun('type.option.option'(X80),fun(bool,fun(fun(X80,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X80),X81))),s(bool,X82)),fun(X80,bool)))) )),
  inference(superposition,[],[f585,f166])).
fof(f10926,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,happ(s(fun(bool,bool),X96),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X92),X93),s(bool,X94),s(fun(X92,bool),X95)))))) | p(s(bool,happ(s(fun(bool,bool),X96),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(happ(s(fun(fun(X92,bool),bool),happ(s(fun(bool,fun(fun(X92,bool),bool)),happ(s(fun('type.option.option'(X92),fun(bool,fun(fun(X92,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X92),X93))),s(bool,X94))),s(fun(X92,bool),X95)),X96,bool))) )),
  inference(superposition,[],[f675,f166])).
fof(f156,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK8(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X17,X18] : (? [X12] : (s(X6,X12) = s(X6,X18) & p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X12))))) <=> p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X18)))))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','thm.bool.UNWIND_THM2')).
fof(f157,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK8(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f103])).
fof(f158,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f568,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | s(X0,X4) = s(X0,sK8(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) )),
  inference(superposition,[],[f157,f163])).
fof(f585,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X2)))) | s(X0,X2) = s(X0,sK8('const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)),X4,X0))) )),
  inference(forward_demodulation,[],[f569,f163])).
fof(f569,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X2)))) | s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3))) = s(X0,sK8('const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)),X4,X0))) )),
  inference(superposition,[],[f157,f163])).
fof(f681,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK8(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0)))))) )),
  inference(forward_demodulation,[],[f670,f163])).
fof(f670,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK8(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3))),s(X0,X4))))) )),
  inference(superposition,[],[f158,f163])).
fof(f683,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(condensation,[],[f682])).
fof(f682,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK8(X10,X9,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f678,f192])).
fof(f678,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK8(X10,X9,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f158,f190])).
fof(f733,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X5,bool),X8),s(X5,X6))) | s(bool,X9) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.SOME_1'(s(X5,X6))),s(bool,X7),s(fun(X5,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f165,f190])).
fof(f735,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun(X15,bool),X18),s(X15,X16))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.SOME_1'(s(X15,X16))),s(bool,X17),s(fun(X15,bool),X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f165,f190])).
fof(f737,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X25),'const.option.SOME_1'(s(X25,X26))),s(bool,X27),s(fun(X25,bool),X28))) | s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),X28),s(X25,X26)))) )),
  inference(superposition,[],[f165,f114])).
fof(f912,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) != s(X0,X3) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f902,f111])).
fof(f902,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) != s(X0,X3) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f156,f190])).
fof(f1119,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,happ(s(fun(X69,bool),X70),s(X69,X71)))) | s(X69,X71) = s(X69,sK8(X71,sK7('const.option.SOME_1'(s(fun(X69,bool),X70)),fun(X69,bool)),X69))) )),
  inference(superposition,[],[f157,f946])).
fof(f1135,plain,(
  ( ! [X21,X22,X20] : (s(X20,X21) = s(X20,sK8(sK7('const.option.SOME_1'(s(X20,X21)),X20),X22,X20)) | ~p(s(bool,happ(s(fun(X20,bool),X22),s(X20,X21))))) )),
  inference(forward_demodulation,[],[f1100,f946])).
fof(f1100,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,happ(s(fun(X20,bool),X22),s(X20,X21)))) | s(X20,sK7('const.option.SOME_1'(s(X20,X21)),X20)) = s(X20,sK8(sK7('const.option.SOME_1'(s(X20,X21)),X20),X22,X20))) )),
  inference(superposition,[],[f157,f946])).
fof(f1143,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,happ(s(fun(X66,bool),X67),s(X66,X68)))) | p(s(bool,happ(s(fun(X66,bool),X67),s(X66,sK8(X68,sK7('const.option.SOME_1'(s(fun(X66,bool),X67)),fun(X66,bool)),X66)))))) )),
  inference(forward_demodulation,[],[f1118,f946])).
fof(f1118,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,happ(s(fun(X66,bool),X67),s(X66,sK8(X68,sK7('const.option.SOME_1'(s(fun(X66,bool),X67)),fun(X66,bool)),X66))))) | ~p(s(bool,happ(s(fun(X66,bool),sK7('const.option.SOME_1'(s(fun(X66,bool),X67)),fun(X66,bool))),s(X66,X68))))) )),
  inference(superposition,[],[f158,f946])).
fof(f1196,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X21),s(X18,X19))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X18),'const.option.SOME_1'(s(X18,X19))),s(bool,X20),s(fun(X18,bool),X21)))) )),
  inference(forward_demodulation,[],[f1157,f946])).
fof(f1157,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,happ(s(fun(X18,bool),X21),s(X18,X19)))),bool)) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X18),'const.option.SOME_1'(s(X18,X19))),s(bool,X20),s(fun(X18,bool),X21)))) )),
  inference(superposition,[],[f1093,f165])).
fof(f1284,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),'const.option.NONE_0'),s(bool,happ(s(fun(X43,bool),X46),s(X43,X44))),s(fun(X47,bool),X48))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.SOME_1'(s(X43,X44))),s(bool,X45),s(fun(X43,bool),X46)))) )),
  inference(superposition,[],[f239,f165])).
fof(f1346,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,happ(s(fun(sK0,bool),X40),s(sK0,sK2))) = s(bool,happ(s(fun(X36,bool),X39),s(X36,X37))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X40,sK0)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.SOME_1'(s(X36,X37))),s(bool,X38),s(fun(X36,bool),X39)))) ) | $spl0),
  inference(superposition,[],[f928,f165])).
fof(f10557,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | s(X0,X4) = s(X0,sK8('const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.NONE_0'),s(X0,X4),s(fun(X5,X0),X6)),'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) )),
  inference(superposition,[],[f585,f163])).
fof(f10560,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(X23,bool),X24),s(X23,X25)))) | s(X23,X25) = s(X23,sK8('const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(X23,X25),s(fun(X26,X23),X27)),sK7('const.option.SOME_1'(s(fun(X23,bool),X24)),fun(X23,bool)),X23))) )),
  inference(superposition,[],[f585,f946])).
fof(f10658,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(X5,X7) = s(X5,sK8('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(X5,X7),s(fun(X8,X5),X9)),X6,X5)) | s(bool,X10) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f10607,f111])).
fof(f10607,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | s(X5,X7) = s(X5,sK8('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(X5,X7),s(fun(X8,X5),X9)),X6,X5)) | s(bool,X10) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f585,f190])).
fof(f10925,plain,(
  ( ! [X90,X88,X87,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,happ(s(fun(X87,bool),X90),s(X87,X88)))))) | p(s(bool,happ(s(fun(bool,bool),X91),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X87),'const.option.SOME_1'(s(X87,X88))),s(bool,X89),s(fun(X87,bool),X90)),X91,bool))) )),
  inference(superposition,[],[f675,f165])).
fof(f1060,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(X36,'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s(X36,X39),s(fun(X37,X36),X40))) = s(X36,sK8(happ(s(fun(fun(X37,X36),X36),happ(s(fun(X36,fun(fun(X37,X36),X36)),happ(s(fun('type.option.option'(X37),fun(X36,fun(fun(X37,X36),X36))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s(X36,X39))),s(fun(X37,X36),X40)),X41,X36)) | ~p(s(bool,happ(s(fun(X36,bool),X41),s(X36,'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s(X36,X39),s(fun(X37,X36),X40))))))) )),
  inference(forward_demodulation,[],[f1027,f166])).
fof(f1027,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,happ(s(fun(X36,bool),X41),s(X36,'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s(X36,X39),s(fun(X37,X36),X40)))))) | s(X36,happ(s(fun(fun(X37,X36),X36),happ(s(fun(X36,fun(fun(X37,X36),X36)),happ(s(fun('type.option.option'(X37),fun(X36,fun(fun(X37,X36),X36))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s(X36,X39))),s(fun(X37,X36),X40))) = s(X36,sK8(happ(s(fun(fun(X37,X36),X36),happ(s(fun(X36,fun(fun(X37,X36),X36)),happ(s(fun('type.option.option'(X37),fun(X36,fun(fun(X37,X36),X36))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s(X36,X39))),s(fun(X37,X36),X40)),X41,X36))) )),
  inference(superposition,[],[f157,f166])).
fof(f781,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(X29,happ(s(fun(X30,X29),X33),s(X30,X31))) = s(X29,sK8('const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(X29,X32),s(fun(X30,X29),X33)),X34,X29)) | ~p(s(bool,happ(s(fun(X29,bool),X34),s(X29,happ(s(fun(X30,X29),X33),s(X30,X31))))))) )),
  inference(forward_demodulation,[],[f743,f165])).
fof(f743,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,happ(s(fun(X29,bool),X34),s(X29,happ(s(fun(X30,X29),X33),s(X30,X31)))))) | s(X29,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(X29,X32),s(fun(X30,X29),X33))) = s(X29,sK8('const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(X29,X32),s(fun(X30,X29),X33)),X34,X29))) )),
  inference(superposition,[],[f157,f165])).
fof(f10558,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),X11),s(X8,X9))),s(X7,X12)))) | s(X7,X12) = s(X7,sK8('const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.NONE_0'),s(X7,X12),s(fun(X13,X7),X14)),'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(fun(X7,bool),X10),s(fun(X8,fun(X7,bool)),X11)),X7))) )),
  inference(superposition,[],[f585,f165])).
fof(f788,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (~p(s(bool,happ(s(fun(X106,bool),happ(s(fun(X107,fun(X106,bool)),X110),s(X107,X108))),s(X106,X111)))) | p(s(bool,happ(s(fun(X106,bool),happ(s(fun(X107,fun(X106,bool)),X110),s(X107,X108))),s(X106,sK8(X111,'const.option.option_CASE_3'(s('type.option.option'(X107),'const.option.SOME_1'(s(X107,X108))),s(fun(X106,bool),X109),s(fun(X107,fun(X106,bool)),X110)),X106)))))) )),
  inference(forward_demodulation,[],[f759,f165])).
fof(f759,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (p(s(bool,happ(s(fun(X106,bool),happ(s(fun(X107,fun(X106,bool)),X110),s(X107,X108))),s(X106,sK8(X111,'const.option.option_CASE_3'(s('type.option.option'(X107),'const.option.SOME_1'(s(X107,X108))),s(fun(X106,bool),X109),s(fun(X107,fun(X106,bool)),X110)),X106))))) | ~p(s(bool,happ(s(fun(X106,bool),'const.option.option_CASE_3'(s('type.option.option'(X107),'const.option.SOME_1'(s(X107,X108))),s(fun(X106,bool),X109),s(fun(X107,fun(X106,bool)),X110))),s(X106,X111))))) )),
  inference(superposition,[],[f158,f165])).
fof(f760,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (~p(s(bool,happ(s(fun(X112,bool),happ(s(fun(X113,fun(X112,bool)),X116),s(X113,X114))),s(X112,X117)))) | s(X112,X117) = s(X112,sK8(X117,'const.option.option_CASE_3'(s('type.option.option'(X113),'const.option.SOME_1'(s(X113,X114))),s(fun(X112,bool),X115),s(fun(X113,fun(X112,bool)),X116)),X112))) )),
  inference(superposition,[],[f157,f165])).
fof(f10559,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun(X15,bool),'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(fun(X15,bool),X18),s(fun(X16,fun(X15,bool)),X19))),s(X15,X20)))) | s(X15,X20) = s(X15,sK8('const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.NONE_0'),s(X15,X20),s(fun(X21,X15),X22)),happ(s(fun(fun(X16,fun(X15,bool)),fun(X15,bool)),happ(s(fun(fun(X15,bool),fun(fun(X16,fun(X15,bool)),fun(X15,bool))),happ(s(fun('type.option.option'(X16),fun(fun(X15,bool),fun(fun(X16,fun(X15,bool)),fun(X15,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(fun(X15,bool),X18))),s(fun(X16,fun(X15,bool)),X19)),X15))) )),
  inference(superposition,[],[f585,f166])).
fof(f1070,plain,(
  ( ! [X146,X144,X142,X145,X143,X141] : (~p(s(bool,happ(s(fun(X141,bool),'const.option.option_CASE_3'(s('type.option.option'(X142),X143),s(fun(X141,bool),X144),s(fun(X142,fun(X141,bool)),X145))),s(X141,X146)))) | p(s(bool,happ(s(fun(X141,bool),'const.option.option_CASE_3'(s('type.option.option'(X142),X143),s(fun(X141,bool),X144),s(fun(X142,fun(X141,bool)),X145))),s(X141,sK8(X146,happ(s(fun(fun(X142,fun(X141,bool)),fun(X141,bool)),happ(s(fun(fun(X141,bool),fun(fun(X142,fun(X141,bool)),fun(X141,bool))),happ(s(fun('type.option.option'(X142),fun(fun(X141,bool),fun(fun(X142,fun(X141,bool)),fun(X141,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X142),X143))),s(fun(X141,bool),X144))),s(fun(X142,fun(X141,bool)),X145)),X141)))))) )),
  inference(forward_demodulation,[],[f1047,f166])).
fof(f1047,plain,(
  ( ! [X146,X144,X142,X145,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),'const.option.option_CASE_3'(s('type.option.option'(X142),X143),s(fun(X141,bool),X144),s(fun(X142,fun(X141,bool)),X145))),s(X141,sK8(X146,happ(s(fun(fun(X142,fun(X141,bool)),fun(X141,bool)),happ(s(fun(fun(X141,bool),fun(fun(X142,fun(X141,bool)),fun(X141,bool))),happ(s(fun('type.option.option'(X142),fun(fun(X141,bool),fun(fun(X142,fun(X141,bool)),fun(X141,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X142),X143))),s(fun(X141,bool),X144))),s(fun(X142,fun(X141,bool)),X145)),X141))))) | ~p(s(bool,happ(s(fun(X141,bool),happ(s(fun(fun(X142,fun(X141,bool)),fun(X141,bool)),happ(s(fun(fun(X141,bool),fun(fun(X142,fun(X141,bool)),fun(X141,bool))),happ(s(fun('type.option.option'(X142),fun(fun(X141,bool),fun(fun(X142,fun(X141,bool)),fun(X141,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X142),X143))),s(fun(X141,bool),X144))),s(fun(X142,fun(X141,bool)),X145))),s(X141,X146))))) )),
  inference(superposition,[],[f158,f166])).
fof(f1048,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X147,bool),'const.option.option_CASE_3'(s('type.option.option'(X148),X149),s(fun(X147,bool),X150),s(fun(X148,fun(X147,bool)),X151))),s(X147,X152)))) | s(X147,X152) = s(X147,sK8(X152,happ(s(fun(fun(X148,fun(X147,bool)),fun(X147,bool)),happ(s(fun(fun(X147,bool),fun(fun(X148,fun(X147,bool)),fun(X147,bool))),happ(s(fun('type.option.option'(X148),fun(fun(X147,bool),fun(fun(X148,fun(X147,bool)),fun(X147,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X148),X149))),s(fun(X147,bool),X150))),s(fun(X148,fun(X147,bool)),X151)),X147))) )),
  inference(superposition,[],[f157,f166])).
fof(f570,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X5)))) | s(bool,'F_0') = s(bool,sK8('F_0',X6,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f157,f114])).
fof(f573,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(bool,bool),X15),s(bool,X14)))) | s(bool,X13) = s(bool,sK8(X13,X15,bool)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f157,f190])).
fof(f576,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,happ(s(fun(bool,bool),X22),s(bool,X21)))) | s(bool,'T_0') = s(bool,sK8('T_0',X22,bool)) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f157,f190])).
fof(f671,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,X2) = s(bool,sK8(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f158,f190])).
fof(f673,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,'T_0') = s(bool,sK8(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f158,f190])).
fof(f675,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X12,X13,bool))) )),
  inference(superposition,[],[f158,f114])).
fof(f892,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X13)))) | p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | s(bool,'F_0') != s(bool,X15) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f156,f114])).
fof(f893,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | s(bool,X16) != s(bool,X18) | s(bool,X16) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f156,f190])).
fof(f895,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X27)))) | s(bool,X24) != s(bool,X27) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f156,f190])).
fof(f897,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X34)))) | s(bool,X32) != s(bool,X34) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f156,f114])).
fof(f898,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X37)))) | s(bool,'T_0') != s(bool,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f156,f190])).
fof(f1203,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,X30) | ~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(forward_demodulation,[],[f1180,f946])).
fof(f1180,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) | s(bool,X30) != s(bool,sK7('const.option.SOME_1'(s(bool,X28)),bool)) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f156,f1093])).
fof(f1401,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun(bool,bool),X87),s(bool,happ(s(fun(sK0,bool),X86),s(sK0,sK2)))))) | p(s(bool,happ(s(fun(bool,bool),X87),s(bool,X88)))) | s(bool,X85) != s(bool,X88) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X86,sK0)) | s(bool,'T_0') = s(bool,X85)) ) | $spl0),
  inference(superposition,[],[f156,f928])).
fof(f1415,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(bool,bool),X142),s(bool,happ(s(fun(sK0,bool),X143),s(sK0,sK2)))))) | ~p(s(bool,happ(s(fun(bool,bool),X142),s(bool,X141)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X143,sK0)) | s(bool,'T_0') = s(bool,sK8(X141,X142,bool))) ) | $spl0),
  inference(superposition,[],[f158,f928])).
fof(f1440,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X63)))) | p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | s(bool,X65) != s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X62,sK0)) | s(bool,'T_0') = s(bool,X63)) ) | $spl0),
  inference(superposition,[],[f156,f928])).
fof(f1441,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,happ(s(fun(bool,bool),X68),s(bool,X67)))) | s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK2))) = s(bool,sK8(happ(s(fun(sK0,bool),X66),s(sK0,sK2)),X68,bool)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X66,sK0)) | s(bool,'T_0') = s(bool,X67)) ) | $spl0),
  inference(superposition,[],[f157,f928])).
fof(f2520,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,X70) != s(bool,X72) | ~p(s(bool,happ(s(fun(bool,bool),X71),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X71),s(bool,X72)))) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X73),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X73,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f2444,f946])).
fof(f2444,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,happ(s(fun(bool,bool),X71),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X71),s(bool,X72)))) | s(bool,X72) != s(bool,sK7('const.option.SOME_1'(s(bool,X70)),bool)) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X73),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X73,sK0))) ) | $spl0),
  inference(superposition,[],[f156,f1348])).
fof(f3297,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X3)))) | s(bool,'F_0') = s(bool,sK8('F_0','const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(bool,bool),X1),s(fun(X0,fun(bool,bool)),X2)),bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f570,f163])).
fof(f3300,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | s(bool,'F_0') = s(bool,sK8('F_0',sK7('const.option.SOME_1'(s(fun(bool,bool),X14)),fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f570,f946])).
fof(f10567,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,X37) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X40),'const.option.NONE_0'),s(bool,X37),s(fun(X40,bool),X41)),X39,bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f585,f190])).
fof(f10580,plain,(
  ( ! [X103,X101,X99,X102,X100] : (~p(s(bool,happ(s(fun(bool,bool),X100),s(bool,X99)))) | s(bool,'T_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X101),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X101,bool),X102)),X100,bool)) | s(bool,X99) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f585,f190])).
fof(f10582,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,happ(s(fun(bool,bool),X110),s(bool,X109)))) | s(bool,'F_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X111),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X111,bool),X112)),X110,bool)) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f585,f114])).
fof(f10895,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X5,X7,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f675,f190])).
fof(f10910,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(bool,bool),X48),s(bool,X47)))) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8('F_0',X48,bool)) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f675,f114])).
fof(f10924,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,happ(s(fun(bool,bool),X86),s(bool,X84)))) | p(s(bool,happ(s(fun(bool,bool),X86),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X83),'const.option.NONE_0'),s(bool,X84),s(fun(X83,bool),X85)),X86,bool))) )),
  inference(superposition,[],[f675,f163])).
fof(f10927,plain,(
  ( ! [X97,X98] : (~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X97)))) | p(s(bool,happ(s(fun(bool,bool),X98),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X97)),bool),X98,bool))) )),
  inference(superposition,[],[f675,f946])).
fof(f10942,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X3)))) | s(bool,'T_0') = s(bool,sK8(X3,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(bool,bool),X1),s(fun(X0,fun(bool,bool)),X2)),bool))) )),
  inference(forward_demodulation,[],[f10889,f163])).
fof(f10889,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X3)))) | p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(bool,bool),X1),s(fun(X0,fun(bool,bool)),X2))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X3,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(bool,bool),X1),s(fun(X0,fun(bool,bool)),X2)),bool))) )),
  inference(superposition,[],[f675,f163])).
fof(f10945,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | s(bool,'T_0') = s(bool,sK8(X15,sK7('const.option.SOME_1'(s(fun(bool,bool),X14)),fun(bool,bool)),bool))) )),
  inference(forward_demodulation,[],[f10892,f946])).
fof(f10892,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | p(s(bool,happ(s(fun(bool,bool),sK7('const.option.SOME_1'(s(fun(bool,bool),X14)),fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X15,sK7('const.option.SOME_1'(s(fun(bool,bool),X14)),fun(bool,bool)),bool))) )),
  inference(superposition,[],[f675,f946])).
fof(f10963,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X3,X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f10929,f111])).
fof(f10929,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X3,X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f675,f190])).
fof(f571,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) | s(bool,X7) = s(bool,sK8(X7,X8,bool)) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f157,f190])).
fof(f1204,plain,(
  ( ! [X31,X32] : (s(bool,X31) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X31)),bool),X32,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X31)) )),
  inference(forward_demodulation,[],[f1181,f946])).
fof(f1181,plain,(
  ( ! [X31,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,sK7('const.option.SOME_1'(s(bool,X31)),bool)) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X31)),bool),X32,bool)) | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f157,f1093])).
fof(f1335,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X74) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,X74),s(fun(X73,bool),X75)),X76,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X76),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X74)) )),
  inference(forward_demodulation,[],[f1313,f163])).
fof(f1313,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,happ(s(fun(bool,bool),X76),s(bool,'T_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,X74),s(fun(X73,bool),X75))) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,X74),s(fun(X73,bool),X75)),X76,bool)) | s(bool,'F_0') = s(bool,X74)) )),
  inference(superposition,[],[f157,f239])).
fof(f2521,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X74)),bool),X75,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X75),s(bool,'T_0')))) | s(bool,X74) = s(bool,happ(s(fun(sK0,bool),X76),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X76,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f2445,f946])).
fof(f2445,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,happ(s(fun(bool,bool),X75),s(bool,'T_0')))) | s(bool,sK7('const.option.SOME_1'(s(bool,X74)),bool)) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X74)),bool),X75,bool)) | s(bool,X74) = s(bool,happ(s(fun(sK0,bool),X76),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X76,sK0))) ) | $spl0),
  inference(superposition,[],[f157,f1348])).
fof(f10566,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'T_0')))) | s(bool,X32) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.NONE_0'),s(bool,X32),s(fun(X34,bool),X35)),X33,bool)) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f585,f190])).
fof(f10894,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X2,X3,bool)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f675,f190])).
fof(f10911,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,happ(s(fun(bool,bool),X52),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X52),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X49),'const.option.NONE_0'),s(bool,X50),s(fun(X49,bool),X51)),X52,bool)) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f675,f239])).
fof(f575,plain,(
  ( ! [X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'F_0')))) | s(bool,X19) = s(bool,sK8(X19,X20,bool)) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f157,f114])).
fof(f8515,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,X3) = s(bool,sK8(X3,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(bool,bool),X1),s(fun(X0,fun(bool,bool)),X2)),bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f575,f163])).
fof(f8518,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'F_0')))) | s(bool,X15) = s(bool,sK8(X15,sK7('const.option.SOME_1'(s(fun(bool,bool),X14)),fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f575,f946])).
fof(f10565,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0')))) | s(bool,X28) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.NONE_0'),s(bool,X28),s(fun(X30,bool),X31)),X29,bool)) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f585,f114])).
fof(f1402,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,happ(s(fun(sK0,bool),X90),s(sK0,sK2)))))) | s(bool,X89) = s(bool,sK8(X89,X91,bool)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X90,sK0)) | s(bool,'T_0') = s(bool,X89)) ) | $spl0),
  inference(superposition,[],[f157,f928])).
fof(f8517,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(bool,bool),X11),s(fun(X9,fun(bool,bool)),X12))),s(bool,'F_0')))) | s(bool,X13) = s(bool,sK8(X13,happ(s(fun(fun(X9,fun(bool,bool)),fun(bool,bool)),happ(s(fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool))),happ(s(fun('type.option.option'(X9),fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(bool,bool),X11))),s(fun(X9,fun(bool,bool)),X12)),bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f575,f166])).
fof(f10944,plain,(
  ( ! [X12,X10,X13,X11,X9] : (p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(bool,bool),X11),s(fun(X9,fun(bool,bool)),X12))),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(bool,bool),X11),s(fun(X9,fun(bool,bool)),X12))),s(bool,X13)))) | s(bool,'T_0') = s(bool,sK8(X13,happ(s(fun(fun(X9,fun(bool,bool)),fun(bool,bool)),happ(s(fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool))),happ(s(fun('type.option.option'(X9),fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(bool,bool),X11))),s(fun(X9,fun(bool,bool)),X12)),bool))) )),
  inference(forward_demodulation,[],[f10891,f166])).
fof(f10891,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(bool,bool),X11),s(fun(X9,fun(bool,bool)),X12))),s(bool,X13)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(fun(X9,fun(bool,bool)),fun(bool,bool)),happ(s(fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool))),happ(s(fun('type.option.option'(X9),fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(bool,bool),X11))),s(fun(X9,fun(bool,bool)),X12))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X13,happ(s(fun(fun(X9,fun(bool,bool)),fun(bool,bool)),happ(s(fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool))),happ(s(fun('type.option.option'(X9),fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(bool,bool),X11))),s(fun(X9,fun(bool,bool)),X12)),bool))) )),
  inference(superposition,[],[f675,f166])).
fof(f3299,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(bool,bool),X11),s(fun(X9,fun(bool,bool)),X12))),s(bool,X13)))) | s(bool,'F_0') = s(bool,sK8('F_0',happ(s(fun(fun(X9,fun(bool,bool)),fun(bool,bool)),happ(s(fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool))),happ(s(fun('type.option.option'(X9),fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(bool,bool),X11))),s(fun(X9,fun(bool,bool)),X12)),bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f570,f166])).
fof(f8516,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X4,fun(bool,bool)),X7),s(X4,X5))),s(bool,'F_0')))) | s(bool,X8) = s(bool,sK8(X8,'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(bool,bool),X6),s(fun(X4,fun(bool,bool)),X7)),bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f575,f165])).
fof(f10943,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X4,fun(bool,bool)),X7),s(X4,X5))),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X4,fun(bool,bool)),X7),s(X4,X5))),s(bool,X8)))) | s(bool,'T_0') = s(bool,sK8(X8,'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(bool,bool),X6),s(fun(X4,fun(bool,bool)),X7)),bool))) )),
  inference(forward_demodulation,[],[f10890,f165])).
fof(f10890,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X4,fun(bool,bool)),X7),s(X4,X5))),s(bool,X8)))) | p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(bool,bool),X6),s(fun(X4,fun(bool,bool)),X7))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X8,'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(bool,bool),X6),s(fun(X4,fun(bool,bool)),X7)),bool))) )),
  inference(superposition,[],[f675,f165])).
fof(f3298,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X4,fun(bool,bool)),X7),s(X4,X5))),s(bool,X8)))) | s(bool,'F_0') = s(bool,sK8('F_0','const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(bool,bool),X6),s(fun(X4,fun(bool,bool)),X7)),bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f570,f165])).
fof(f2216,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X29),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X29),s('type.option.option'(bool),X30)))) | s('type.option.option'(bool),X30) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f156,f1200])).
fof(f2217,plain,(
  ( ! [X31,X32] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X32),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,X31)),X32,'type.option.option'(bool))) | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f157,f1200])).
fof(f2193,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X29),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28)))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X29),s('type.option.option'(bool),X30)))) | s('type.option.option'(bool),X30) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f156,f1200])).
fof(f2194,plain,(
  ( ! [X31,X32] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X32),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,'T_0')),X32,'type.option.option'(bool))) | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f157,f1200])).
fof(f10603,plain,(
  ( ! [X206,X204,X207,X205,X203] : (~p(s(bool,happ(s(fun('type.option.option'(X203),bool),X205),s('type.option.option'(X203),X204)))) | s('type.option.option'(X203),'const.option.SOME_1'(s(X203,sK7(X204,X203)))) = s('type.option.option'(X203),sK8('const.option.option_CASE_3'(s('type.option.option'(X206),'const.option.NONE_0'),s('type.option.option'(X203),'const.option.SOME_1'(s(X203,sK7(X204,X203)))),s(fun(X206,'type.option.option'(X203)),X207)),X205,'type.option.option'(X203))) | s('type.option.option'(X203),'const.option.NONE_0') = s('type.option.option'(X203),X204)) )),
  inference(superposition,[],[f585,f146])).
fof(f621,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun('type.option.option'(X24),bool),X26),s('type.option.option'(X24),X25)))) | s('type.option.option'(X24),'const.option.SOME_1'(s(X24,sK7(X25,X24)))) = s('type.option.option'(X24),sK8('const.option.SOME_1'(s(X24,sK7(X25,X24))),X26,'type.option.option'(X24))) | s('type.option.option'(X24),'const.option.NONE_0') = s('type.option.option'(X24),X25)) )),
  inference(superposition,[],[f157,f146])).
fof(f900,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun('type.option.option'(X43),bool),X45),s('type.option.option'(X43),X44)))) | p(s(bool,happ(s(fun('type.option.option'(X43),bool),X45),s('type.option.option'(X43),X46)))) | s('type.option.option'(X43),X46) != s('type.option.option'(X43),'const.option.SOME_1'(s(X43,sK7(X44,X43)))) | s('type.option.option'(X43),'const.option.NONE_0') = s('type.option.option'(X43),X44)) )),
  inference(superposition,[],[f156,f146])).
fof(f1350,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,X8) = s(bool,happ(s(fun(sK0,bool),happ(s(fun(X4,fun(sK0,bool)),X7),s(X4,X5))),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(sK0,bool),X6),s(fun(X4,fun(sK0,bool)),X7)),sK0)) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f928,f165])).
fof(f2281,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,X8) = s(bool,happ(s(fun(sK0,bool),happ(s(fun(X4,fun(sK0,bool)),X7),s(X4,X5))),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(sK0,bool),X6),s(fun(X4,fun(sK0,bool)),X7)),sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X8)),bool))) ) | $spl0),
  inference(superposition,[],[f1348,f165])).
fof(f922,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(X3,fun(sK0,bool)),X6),s(X3,X4))),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.SOME_1'(s(X3,X4))),s(fun(sK0,bool),X5),s(fun(X3,fun(sK0,bool)),X6)),sK0))) ) | $spl0),
  inference(superposition,[],[f869,f165])).
fof(f1351,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(bool,X13) = s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(sK0,bool),X11),s(fun(X9,fun(sK0,bool)),X12))),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),happ(s(fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X9),fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(sK0,bool),X11))),s(fun(X9,fun(sK0,bool)),X12)),sK0)) | s(bool,'T_0') = s(bool,X13)) ) | $spl0),
  inference(superposition,[],[f928,f166])).
fof(f2282,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(bool,X13) = s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(sK0,bool),X11),s(fun(X9,fun(sK0,bool)),X12))),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),happ(s(fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X9),fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(sK0,bool),X11))),s(fun(X9,fun(sK0,bool)),X12)),sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X13)),bool))) ) | $spl0),
  inference(superposition,[],[f1348,f166])).
fof(f1045,plain,(
  ( ! [X132,X130,X133,X131] : (~p(s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X130),X131),s(fun(sK0,bool),X132),s(fun(X130,fun(sK0,bool)),X133))),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),happ(s(fun(fun(X130,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X130,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X130),fun(fun(sK0,bool),fun(fun(X130,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X130),X131))),s(fun(sK0,bool),X132))),s(fun(X130,fun(sK0,bool)),X133)),sK0))) ) | $spl0),
  inference(superposition,[],[f869,f166])).
fof(f869,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X4,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f817,f170])).
fof(f817,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK2)))) | s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X4,sK0))) ) | $spl0),
  inference(superposition,[],[f157,f170])).
fof(f921,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0))) ) | $spl0),
  inference(superposition,[],[f869,f163])).
fof(f928,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X0,sK0)) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(subsumption_resolution,[],[f923,f111])).
fof(f923,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X0,sK0)) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(superposition,[],[f869,f190])).
fof(f1116,plain,(
  ( ! [X61] : (~p(s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),sK7('const.option.SOME_1'(s(fun(sK0,bool),X61)),fun(sK0,bool)),sK0))) ) | $spl0),
  inference(superposition,[],[f869,f946])).
fof(f1271,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),sK7('const.option.SOME_1'(s(fun(sK0,bool),X0)),fun(sK0,bool)),sK0)) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(subsumption_resolution,[],[f1262,f111])).
fof(f1262,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),sK7('const.option.SOME_1'(s(fun(sK0,bool),X0)),fun(sK0,bool)),sK0)) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(superposition,[],[f1116,f190])).
fof(f1345,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) = s(bool,happ(s(fun(sK0,bool),X35),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X35,sK0)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X32),'const.option.NONE_0'),s(bool,X33),s(fun(X32,bool),X34)))) ) | $spl0),
  inference(superposition,[],[f928,f163])).
fof(f1348,plain,(
  ( ! [X47,X46] : (s(bool,X46) = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X47,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X46)),bool))) ) | $spl0),
  inference(superposition,[],[f928,f946])).
fof(f1349,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X3) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0)) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f928,f163])).
fof(f1363,plain,(
  ( ! [X6,X8,X7] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X7))) | s(sK0,X7) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X6,sK0)) | s(bool,X8) = s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f154,f928])).
fof(f1365,plain,(
  ( ! [X14,X12,X13] : (s(sK0,sK2) != s(sK0,X13) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X13))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X12,sK0)))) | s(bool,X14) = s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X14)) ) | $spl0),
  inference(superposition,[],[f155,f928])).
fof(f1367,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK2)))) | p(s(bool,happ(s(fun(sK0,bool),X19),s(sK0,X20)))) | s(sK0,X20) != s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X21)) ) | $spl0),
  inference(superposition,[],[f156,f928])).
fof(f1368,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,happ(s(fun(sK0,bool),X23),s(sK0,sK2)))) | s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X22,sK0)) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X22,sK0),X23,sK0)) | s(bool,X24) = s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X24)) ) | $spl0),
  inference(superposition,[],[f157,f928])).
fof(f1399,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,happ(s(fun(sK0,bool),X80),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X79))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X81))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X80,sK0)) | s(bool,'T_0') = s(bool,X79)) ) | $spl0),
  inference(superposition,[],[f155,f928])).
fof(f1416,plain,(
  ( ! [X146,X144,X147,X145] : (s(X146,f(s(bool,happ(s(fun(sK0,bool),X147),s(sK0,sK2))))) != s(X146,g(s(bool,happ(s(fun(sK0,bool),X147),s(sK0,sK2))))) | s(fun(bool,X146),X144) = s(fun(bool,X146),X145) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X147,sK0)) | s(bool,'T_0') = s(bool,sK9(X144,X145,X146,bool))) ) | $spl0),
  inference(superposition,[],[f164,f928])).
fof(f1436,plain,(
  ( ! [X52,X50,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) | s(bool,X52) = s(bool,happ(s(fun(sK0,bool),X50),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X50,sK0)) | s(bool,'T_0') = s(bool,X51)) ) | $spl0),
  inference(superposition,[],[f154,f928])).
fof(f1448,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0)) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2)))) ) | $spl0),
  inference(equality_factoring,[],[f928])).
fof(f2280,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X3) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X3)),bool))) ) | $spl0),
  inference(superposition,[],[f1348,f163])).
fof(f2283,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),sK7('const.option.SOME_1'(s(fun(sK0,bool),X14)),fun(sK0,bool)),sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X15)),bool))) ) | $spl0),
  inference(superposition,[],[f1348,f946])).
fof(f2349,plain,(
  ( ! [X103,X101,X102] : (s(bool,X103) != s(bool,happ(s(fun(sK0,bool),X102),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X102,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool))) ) | $spl0),
  inference(superposition,[],[f155,f1348])).
fof(f2366,plain,(
  ( ! [X167,X169,X166,X168] : (s(X168,f(s(bool,happ(s(fun(sK0,bool),X169),s(sK0,sK2))))) != s(X168,g(s(bool,happ(s(fun(sK0,bool),X169),s(sK0,sK2))))) | s(fun(bool,X168),X166) = s(fun(bool,X168),X167) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X169,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK9(X166,X167,X168,bool))),bool))) ) | $spl0),
  inference(superposition,[],[f164,f1348])).
fof(f2393,plain,(
  ( ! [X70,X72,X71] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X71))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X72))) | s(bool,X72) = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X70,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X71)),bool))) ) | $spl0),
  inference(superposition,[],[f154,f1348])).
fof(f2408,plain,(
  ( ! [X10,X8,X9] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X9))) | s(sK0,X9) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X8,sK0)) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X10)),bool))) ) | $spl0),
  inference(superposition,[],[f154,f1348])).
fof(f2410,plain,(
  ( ! [X14,X15,X16] : (s(sK0,sK2) != s(sK0,X15) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X15))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X14,sK0)))) | s(bool,X16) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK2))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X16)),bool))) ) | $spl0),
  inference(superposition,[],[f155,f1348])).
fof(f2440,plain,(
  ( ! [X59,X60,X58] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X59))) | s(bool,X59) = s(bool,sK7('const.option.SOME_1'(s(bool,X58)),bool)) | s(bool,X58) = s(bool,happ(s(fun(sK0,bool),X60),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X60,sK0))) ) | $spl0),
  inference(superposition,[],[f154,f1348])).
fof(f2515,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,X10) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X11,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f2514,f946])).
fof(f2514,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,sK7('const.option.SOME_1'(s(bool,X9)),bool)) | s(bool,X9) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X11,sK0))) ) | $spl0),
  inference(subsumption_resolution,[],[f2423,f111])).
fof(f2423,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,sK7('const.option.SOME_1'(s(bool,X9)),bool)) | s(bool,X9) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X11,sK0))) ) | $spl0),
  inference(superposition,[],[f162,f1348])).
fof(f2516,plain,(
  ( ! [X47,X45,X46] : (s(bool,X46) = s(bool,sK7('const.option.SOME_1'(s(bool,X45)),bool)) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X47,sK0))) ) | $spl0),
  inference(subsumption_resolution,[],[f2435,f111])).
fof(f2435,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | s(bool,X46) = s(bool,sK7('const.option.SOME_1'(s(bool,X45)),bool)) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X47,sK0))) ) | $spl0),
  inference(superposition,[],[f1199,f1348])).
fof(f2517,plain,(
  ( ! [X54,X55] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) | s(bool,X54) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X55,sK0))) ) | $spl0),
  inference(subsumption_resolution,[],[f2438,f151])).
fof(f2438,plain,(
  ( ! [X54,X55] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) | s(bool,X54) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X55,sK0))) ) | $spl0),
  inference(superposition,[],[f146,f1348])).
fof(f2518,plain,(
  ( ! [X66,X64,X65] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65))) | s(bool,'T_0') != s(bool,X65) | s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X66,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f2442,f946])).
fof(f2442,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') != s(bool,X65) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X64)),bool)))) | s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X66,sK0))) ) | $spl0),
  inference(superposition,[],[f155,f1348])).
fof(f9381,plain,(
  ( ! [X19,X17,X20,X18] : (s(X18,happ(s(fun(bool,X18),X19),s(bool,'F_0'))) = s(X18,'const.option.option_CASE_3'(s('type.option.option'(bool),X17),s(X18,X20),s(fun(bool,X18),X19))) | s(bool,'T_0') = s(bool,sK7(X17,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X17)) )),
  inference(superposition,[],[f165,f610])).
fof(f9809,plain,(
  ( ! [X101,X102] : (s('type.option.option'(bool),X102) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK7(X101,bool)) = s(bool,sK7(X102,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X102) | s('type.option.option'(bool),X101) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X101)) )),
  inference(superposition,[],[f611,f610])).
fof(f3337,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') = s(bool,sK8('F_0',X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3321,f111])).
fof(f3321,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,sK8('F_0',X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f570,f190])).
fof(f492,plain,(
  ( ! [X6,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f154,f114])).
fof(f493,plain,(
  ( ! [X8,X7,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,X7) = s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f154,f190])).
fof(f495,plain,(
  ( ! [X14,X15,X13] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X14))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f154,f190])).
fof(f497,plain,(
  ( ! [X19,X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f154,f114])).
fof(f532,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f155,f114])).
fof(f533,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,X8) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f155,f190])).
fof(f535,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) != s(bool,X15) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f155,f190])).
fof(f537,plain,(
  ( ! [X19,X20] : (s(bool,'F_0') != s(bool,X20) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f155,f114])).
fof(f538,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X22) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f155,f190])).
fof(f961,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,X1) = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f603,f190])).
fof(f963,plain,(
  ( ! [X4,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X4)),bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f603,f190])).
fof(f1176,plain,(
  ( ! [X21,X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21))) | s(bool,X21) = s(bool,sK7('const.option.SOME_1'(s(bool,X20)),bool)) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f154,f1093])).
fof(f1201,plain,(
  ( ! [X24,X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X24))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25))) | s(bool,'T_0') != s(bool,X25) | s(bool,'F_0') = s(bool,X24)) )),
  inference(forward_demodulation,[],[f1178,f946])).
fof(f1178,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,X25) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X24)),bool)))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f155,f1093])).
fof(f1308,plain,(
  ( ! [X54,X52,X55,X53] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X55))) | s(bool,X55) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X52),'const.option.NONE_0'),s(bool,X53),s(fun(X52,bool),X54))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f154,f239])).
fof(f1980,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,'F_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(bool,X27),s(fun(X26,bool),X28)))) )),
  inference(superposition,[],[f492,f163])).
fof(f1983,plain,(
  ( ! [X41,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X41))) | s(bool,'F_0') = s(bool,X41) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X40)),bool))) )),
  inference(superposition,[],[f492,f946])).
fof(f1995,plain,(
  ( ! [X30,X28,X31,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X28),'const.option.NONE_0'),s(bool,X29),s(fun(X28,bool),X30))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f492,f163])).
fof(f1998,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'F_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X42)),bool)) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f492,f946])).
fof(f2161,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1200,f190])).
fof(f2167,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f1200,f1200])).
fof(f2175,plain,(
  ( ! [X19,X20,X18] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X18),'const.option.NONE_0'),s(bool,X19),s(fun(X18,bool),X20)))) )),
  inference(superposition,[],[f1200,f163])).
fof(f2178,plain,(
  ( ! [X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,'F_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X29)),bool))) )),
  inference(superposition,[],[f1200,f946])).
fof(f2183,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(bool,X7),X8),s(bool,'T_0'))) = s(X7,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))),s(X7,X9),s(fun(bool,X7),X8))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f165,f1200])).
fof(f2189,plain,(
  ( ! [X21,X20] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X21))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))))) | s('type.option.option'(bool),X21) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f154,f1200])).
fof(f2206,plain,(
  ( ! [X10,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10))) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f154,f1200])).
fof(f2214,plain,(
  ( ! [X24,X25] : (s('type.option.option'(bool),X25) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X25))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X24))))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f155,f1200])).
fof(f3099,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f3098])).
fof(f3098,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f3097])).
fof(f3097,plain,(
  ( ! [X6,X7,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f3073,f190])).
fof(f3073,plain,(
  ( ! [X6,X7,X5] : (s(bool,X6) != s(bool,X7) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f532,f190])).
fof(f3186,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(trivial_inequality_removal,[],[f3166])).
fof(f3166,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f537,f114])).
fof(f3242,plain,(
  ( ! [X14,X12,X15,X13] : (s(X13,happ(s(fun(bool,X13),X14),s(bool,'F_0'))) = s(X13,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))),s(X13,X15),s(fun(bool,X13),X14))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f165,f3099])).
fof(f9460,plain,(
  ( ! [X101,X102] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X102))) | s(bool,X102) = s(bool,sK7(X101,bool)) | s('type.option.option'(bool),X101) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X101)) )),
  inference(superposition,[],[f154,f610])).
fof(f10053,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,X33) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.NONE_0'),s(bool,X32),s(fun(X34,bool),X35))) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f491,f190])).
fof(f10054,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(bool,X39) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X40),'const.option.NONE_0'),s(bool,X37),s(fun(X40,bool),X41))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f491,f190])).
fof(f10067,plain,(
  ( ! [X103,X101,X99,X102,X100] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X99))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X100))) | s(bool,X100) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X101),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X101,bool),X102))) | s(bool,X99) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f491,f190])).
fof(f10069,plain,(
  ( ! [X111,X109,X112,X110] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X109))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X110))) | s(bool,X110) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X111),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X111,bool),X112))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f491,f114])).
fof(f10107,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,X32) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.NONE_0'),s(bool,X33),s(fun(X34,bool),X35))) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f491,f190])).
fof(f10108,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(bool,X37) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X40),'const.option.NONE_0'),s(bool,X39),s(fun(X40,bool),X41))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f491,f190])).
fof(f10121,plain,(
  ( ! [X103,X101,X99,X102,X100] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X99))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X100))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X101),'const.option.NONE_0'),s(bool,X100),s(fun(X101,bool),X102))) | s(bool,X99) = s(bool,X103) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f491,f190])).
fof(f10152,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.option.option'(bool),X22) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X24),'const.option.NONE_0'),s(bool,X23),s(fun(X24,bool),X25))) | s(bool,'T_0') = s(bool,sK7(X22,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X22)) )),
  inference(superposition,[],[f491,f610])).
fof(f2191,plain,(
  ( ! [X24,X25] : (s('type.option.option'(bool),X25) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X24))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X25))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f155,f1200])).
fof(f2205,plain,(
  ( ! [X6,X8,X7,X5] : (s(X6,happ(s(fun(bool,X6),X7),s(bool,X5))) = s(X6,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(X6,X8),s(fun(bool,X6),X7))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f165,f1200])).
fof(f2212,plain,(
  ( ! [X21,X20] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X21))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X21) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f154,f1200])).
fof(f9474,plain,(
  ( ! [X98] : (s('type.option.option'(bool),X98) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X98) | s('type.option.option'(bool),X98) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f9457])).
fof(f9457,plain,(
  ( ! [X98] : (s('type.option.option'(bool),X98) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X98) | s('type.option.option'(bool),X98) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X98)) )),
  inference(superposition,[],[f146,f610])).
fof(f239,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.NONE_0'),s(bool,X5),s(fun(X4,bool),X6))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f163,f114])).
fof(f610,plain,(
  ( ! [X8] : (s('type.option.option'(bool),X8) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X8,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X8)) )),
  inference(superposition,[],[f146,f114])).
fof(f933,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK9(X16,X17,X18,bool))) )),
  inference(superposition,[],[f164,f114])).
fof(f1093,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X10)),bool)) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f946,f114])).
fof(f1195,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(bool,X16),s(fun(X15,bool),X17)))) )),
  inference(forward_demodulation,[],[f1156,f946])).
fof(f1156,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X16)),bool)) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(bool,X16),s(fun(X15,bool),X17)))) )),
  inference(superposition,[],[f1093,f163])).
fof(f1199,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) )),
  inference(forward_demodulation,[],[f1198,f946])).
fof(f1198,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,sK7('const.option.SOME_1'(s(bool,X5)),bool)) | s(bool,'F_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1167,f111])).
fof(f1167,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | s(bool,X6) = s(bool,sK7('const.option.SOME_1'(s(bool,X5)),bool)) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f162,f1093])).
fof(f1253,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,sK7('const.option.SOME_1'(s(bool,X22)),bool)) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1240,f111])).
fof(f1240,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) = s(bool,sK7('const.option.SOME_1'(s(bool,X22)),bool)) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1199,f1093])).
fof(f1283,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.NONE_0'),s(bool,X39),s(fun(X41,bool),X42))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X38),'const.option.NONE_0'),s(bool,X39),s(fun(X38,bool),X40)))) )),
  inference(superposition,[],[f239,f163])).
fof(f1331,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X45) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X42),'const.option.NONE_0'),s(bool,X43),s(fun(X42,bool),X44))) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f1305,f111])).
fof(f1305,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | s(bool,X45) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X42),'const.option.NONE_0'),s(bool,X43),s(fun(X42,bool),X44))) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f1199,f239])).
fof(f1984,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK7(X0,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f492,f146])).
fof(f1999,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,sK7(X0,bool)) | s(bool,'T_0') = s(bool,X1) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f492,f146])).
fof(f2179,plain,(
  ( ! [X0] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X0,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f1200,f146])).
fof(f2252,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,sK7(X0,bool)) | s(bool,'T_0') = s(bool,X1) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f497,f146])).
fof(f8548,plain,(
  ( ! [X4,X2,X3] : (s(bool,X3) = s(bool,sK8(X3,X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8531,f111])).
fof(f8531,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,sK8(X3,X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f575,f190])).
fof(f9825,plain,(
  ( ! [X12,X13] : (s('type.option.option'(bool),X12) != s('type.option.option'(bool),X13) | s(bool,'F_0') = s(bool,sK7(X13,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X13) | s(bool,'T_0') = s(bool,sK7(X12,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X12)) )),
  inference(superposition,[],[f611,f610])).
fof(f9836,plain,(
  ( ! [X12,X13] : (s('type.option.option'(bool),X12) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) | s(bool,X13) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK7(X12,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X12)) )),
  inference(subsumption_resolution,[],[f9767,f151])).
fof(f9767,plain,(
  ( ! [X12,X13] : (s('type.option.option'(bool),X12) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) | s(bool,X13) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X12,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X12)) )),
  inference(superposition,[],[f611,f610])).
fof(f10098,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.option.option'(bool),X22) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,X23) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X24),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X24,bool),X25))) | s(bool,'T_0') = s(bool,sK7(X22,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X22)) )),
  inference(superposition,[],[f491,f610])).
fof(f588,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,sK8(X2,X1,X0)) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f580,f111])).
fof(f580,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X2) = s(X0,sK8(X2,X1,X0)) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f157,f190])).
fof(f10144,plain,(
  ( ! [X206,X204,X207,X205,X203] : (s('type.option.option'('type.option.option'(X203)),'const.option.SOME_1'(s('type.option.option'(X203),X204))) != s('type.option.option'('type.option.option'(X203)),'const.option.SOME_1'(s('type.option.option'(X203),X205))) | s('type.option.option'(X203),'const.option.SOME_1'(s(X203,sK7(X204,X203)))) = s('type.option.option'(X203),'const.option.option_CASE_3'(s('type.option.option'(X206),'const.option.NONE_0'),s('type.option.option'(X203),X205),s(fun(X206,'type.option.option'(X203)),X207))) | s('type.option.option'(X203),'const.option.NONE_0') = s('type.option.option'(X203),X204)) )),
  inference(superposition,[],[f491,f146])).
fof(f10105,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s('type.option.option'(X23),'const.option.SOME_1'(s(X23,X24))) != s('type.option.option'(X23),'const.option.SOME_1'(s(X23,X25))) | s(X23,sK7('const.option.SOME_1'(s(X23,X24)),X23)) = s(X23,'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(X23,X25),s(fun(X26,X23),X27)))) )),
  inference(superposition,[],[f491,f946])).
fof(f10104,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.option.option'(X15),'const.option.SOME_1'(s(X15,X20))) != s('type.option.option'(X15),'const.option.SOME_1'(s(X15,'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(X15,X18),s(fun(X16,X15),X19))))) | s(X15,'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.NONE_0'),s(X15,X20),s(fun(X21,X15),X22))) = s(X15,happ(s(fun(fun(X16,X15),X15),happ(s(fun(X15,fun(fun(X16,X15),X15)),happ(s(fun('type.option.option'(X16),fun(X15,fun(fun(X16,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(X15,X18))),s(fun(X16,X15),X19)))) )),
  inference(superposition,[],[f491,f166])).
fof(f10103,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X12))) != s('type.option.option'(X7),'const.option.SOME_1'(s(X7,happ(s(fun(X8,X7),X11),s(X8,X9))))) | s(X7,'const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.NONE_0'),s(X7,X12),s(fun(X13,X7),X14))) = s(X7,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(X7,X10),s(fun(X8,X7),X11)))) )),
  inference(superposition,[],[f491,f165])).
fof(f10102,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X4))) | s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.NONE_0'),s(X0,X4),s(fun(X5,X0),X6)))) )),
  inference(superposition,[],[f491,f163])).
fof(f10100,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s('type.option.option'(X30),X31) != s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X32))) | s(X30,X32) = s(X30,'const.option.option_CASE_3'(s('type.option.option'(X33),'const.option.NONE_0'),s(X30,sK7(X31,X30)),s(fun(X33,X30),X34))) | s('type.option.option'(X30),'const.option.NONE_0') = s('type.option.option'(X30),X31)) )),
  inference(superposition,[],[f491,f146])).
fof(f10090,plain,(
  ( ! [X206,X204,X207,X205,X203] : (s('type.option.option'('type.option.option'(X203)),'const.option.SOME_1'(s('type.option.option'(X203),X204))) != s('type.option.option'('type.option.option'(X203)),'const.option.SOME_1'(s('type.option.option'(X203),X205))) | s('type.option.option'(X203),X205) = s('type.option.option'(X203),'const.option.option_CASE_3'(s('type.option.option'(X206),'const.option.NONE_0'),s('type.option.option'(X203),'const.option.SOME_1'(s(X203,sK7(X204,X203)))),s(fun(X206,'type.option.option'(X203)),X207))) | s('type.option.option'(X203),'const.option.NONE_0') = s('type.option.option'(X203),X204)) )),
  inference(superposition,[],[f491,f146])).
fof(f9827,plain,(
  ( ! [X17,X18,X16] : (s('type.option.option'(X16),X17) != s('type.option.option'(X16),X18) | s(X16,sK7(X17,X16)) = s(X16,sK7(X18,X16)) | s('type.option.option'(X16),'const.option.NONE_0') = s('type.option.option'(X16),X18) | s('type.option.option'(X16),'const.option.NONE_0') = s('type.option.option'(X16),X17)) )),
  inference(superposition,[],[f611,f146])).
fof(f9817,plain,(
  ( ! [X121,X120,X119] : (s('type.option.option'('type.option.option'(X119)),X121) != s('type.option.option'('type.option.option'(X119)),'const.option.SOME_1'(s('type.option.option'(X119),X120))) | s('type.option.option'(X119),sK7(X121,'type.option.option'(X119))) = s('type.option.option'(X119),'const.option.SOME_1'(s(X119,sK7(X120,X119)))) | s('type.option.option'('type.option.option'(X119)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(X119)),X121) | s('type.option.option'(X119),'const.option.NONE_0') = s('type.option.option'(X119),X120)) )),
  inference(superposition,[],[f611,f146])).
fof(f9781,plain,(
  ( ! [X26,X27,X25] : (s('type.option.option'(bool),X27) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s(bool,X25) = s(bool,sK7(X27,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X27) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f611,f190])).
fof(f9780,plain,(
  ( ! [X24,X23,X22] : (s('type.option.option'(bool),X23) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X22) = s(bool,sK7(X23,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f611,f190])).
fof(f9778,plain,(
  ( ! [X19,X17,X18] : (s('type.option.option'(X17),X19) != s('type.option.option'(X17),'const.option.SOME_1'(s(X17,X18))) | s(X17,sK7(X19,X17)) = s(X17,sK7('const.option.SOME_1'(s(X17,X18)),X17)) | s('type.option.option'(X17),'const.option.NONE_0') = s('type.option.option'(X17),X19)) )),
  inference(superposition,[],[f611,f946])).
fof(f9777,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s('type.option.option'(X11),X16) != s('type.option.option'(X11),'const.option.SOME_1'(s(X11,'const.option.option_CASE_3'(s('type.option.option'(X12),X13),s(X11,X14),s(fun(X12,X11),X15))))) | s(X11,sK7(X16,X11)) = s(X11,happ(s(fun(fun(X12,X11),X11),happ(s(fun(X11,fun(fun(X12,X11),X11)),happ(s(fun('type.option.option'(X12),fun(X11,fun(fun(X12,X11),X11))),'const.option.option_CASE_0'),s('type.option.option'(X12),X13))),s(X11,X14))),s(fun(X12,X11),X15))) | s('type.option.option'(X11),'const.option.NONE_0') = s('type.option.option'(X11),X16)) )),
  inference(superposition,[],[f611,f166])).
fof(f9776,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X5),X10) != s('type.option.option'(X5),'const.option.SOME_1'(s(X5,happ(s(fun(X6,X5),X9),s(X6,X7))))) | s(X5,sK7(X10,X5)) = s(X5,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7))),s(X5,X8),s(fun(X6,X5),X9))) | s('type.option.option'(X5),'const.option.NONE_0') = s('type.option.option'(X5),X10)) )),
  inference(superposition,[],[f611,f165])).
fof(f9775,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),X4) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,sK7(X4,X0)) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X4)) )),
  inference(superposition,[],[f611,f163])).
fof(f9841,plain,(
  ( ! [X39,X38,X40] : (s('type.option.option'(X38),'const.option.NONE_0') = s('type.option.option'(X38),X39) | s('type.option.option'(X38),X39) != s('type.option.option'(X38),'const.option.SOME_1'(s(X38,X40))) | s(X38,X40) = s(X38,sK7(sK7('const.option.SOME_1'(s('type.option.option'(X38),X39)),'type.option.option'(X38)),X38))) )),
  inference(forward_demodulation,[],[f9774,f946])).
fof(f9774,plain,(
  ( ! [X39,X38,X40] : (s('type.option.option'(X38),X39) != s('type.option.option'(X38),'const.option.SOME_1'(s(X38,X40))) | s(X38,X40) = s(X38,sK7(sK7('const.option.SOME_1'(s('type.option.option'(X38),X39)),'type.option.option'(X38)),X38)) | s('type.option.option'(X38),'const.option.NONE_0') = s('type.option.option'(X38),sK7('const.option.SOME_1'(s('type.option.option'(X38),X39)),'type.option.option'(X38)))) )),
  inference(superposition,[],[f611,f946])).
fof(f9840,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s('type.option.option'(X32),'const.option.NONE_0') = s('type.option.option'(X32),'const.option.option_CASE_3'(s('type.option.option'(X33),X34),s('type.option.option'(X32),X35),s(fun(X33,'type.option.option'(X32)),X36))) | s('type.option.option'(X32),'const.option.SOME_1'(s(X32,X37))) != s('type.option.option'(X32),'const.option.option_CASE_3'(s('type.option.option'(X33),X34),s('type.option.option'(X32),X35),s(fun(X33,'type.option.option'(X32)),X36))) | s(X32,X37) = s(X32,sK7(happ(s(fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32)),happ(s(fun('type.option.option'(X32),fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32))),happ(s(fun('type.option.option'(X33),fun('type.option.option'(X32),fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32)))),'const.option.option_CASE_0'),s('type.option.option'(X33),X34))),s('type.option.option'(X32),X35))),s(fun(X33,'type.option.option'(X32)),X36)),X32))) )),
  inference(forward_demodulation,[],[f9773,f166])).
fof(f9773,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s('type.option.option'(X32),'const.option.SOME_1'(s(X32,X37))) != s('type.option.option'(X32),'const.option.option_CASE_3'(s('type.option.option'(X33),X34),s('type.option.option'(X32),X35),s(fun(X33,'type.option.option'(X32)),X36))) | s(X32,X37) = s(X32,sK7(happ(s(fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32)),happ(s(fun('type.option.option'(X32),fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32))),happ(s(fun('type.option.option'(X33),fun('type.option.option'(X32),fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32)))),'const.option.option_CASE_0'),s('type.option.option'(X33),X34))),s('type.option.option'(X32),X35))),s(fun(X33,'type.option.option'(X32)),X36)),X32)) | s('type.option.option'(X32),'const.option.NONE_0') = s('type.option.option'(X32),happ(s(fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32)),happ(s(fun('type.option.option'(X32),fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32))),happ(s(fun('type.option.option'(X33),fun('type.option.option'(X32),fun(fun(X33,'type.option.option'(X32)),'type.option.option'(X32)))),'const.option.option_CASE_0'),s('type.option.option'(X33),X34))),s('type.option.option'(X32),X35))),s(fun(X33,'type.option.option'(X32)),X36)))) )),
  inference(superposition,[],[f611,f166])).
fof(f9839,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s('type.option.option'(X26),'const.option.NONE_0') = s('type.option.option'(X26),happ(s(fun(X27,'type.option.option'(X26)),X30),s(X27,X28))) | s('type.option.option'(X26),'const.option.SOME_1'(s(X26,X31))) != s('type.option.option'(X26),happ(s(fun(X27,'type.option.option'(X26)),X30),s(X27,X28))) | s(X26,X31) = s(X26,sK7('const.option.option_CASE_3'(s('type.option.option'(X27),'const.option.SOME_1'(s(X27,X28))),s('type.option.option'(X26),X29),s(fun(X27,'type.option.option'(X26)),X30)),X26))) )),
  inference(forward_demodulation,[],[f9772,f165])).
fof(f9772,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s('type.option.option'(X26),'const.option.SOME_1'(s(X26,X31))) != s('type.option.option'(X26),happ(s(fun(X27,'type.option.option'(X26)),X30),s(X27,X28))) | s(X26,X31) = s(X26,sK7('const.option.option_CASE_3'(s('type.option.option'(X27),'const.option.SOME_1'(s(X27,X28))),s('type.option.option'(X26),X29),s(fun(X27,'type.option.option'(X26)),X30)),X26)) | s('type.option.option'(X26),'const.option.NONE_0') = s('type.option.option'(X26),'const.option.option_CASE_3'(s('type.option.option'(X27),'const.option.SOME_1'(s(X27,X28))),s('type.option.option'(X26),X29),s(fun(X27,'type.option.option'(X26)),X30)))) )),
  inference(superposition,[],[f611,f165])).
fof(f9838,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.option.option'(X21),'const.option.NONE_0') = s('type.option.option'(X21),X23) | s('type.option.option'(X21),X23) != s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X25))) | s(X21,X25) = s(X21,sK7('const.option.option_CASE_3'(s('type.option.option'(X22),'const.option.NONE_0'),s('type.option.option'(X21),X23),s(fun(X22,'type.option.option'(X21)),X24)),X21))) )),
  inference(forward_demodulation,[],[f9771,f163])).
fof(f9771,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.option.option'(X21),X23) != s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X25))) | s(X21,X25) = s(X21,sK7('const.option.option_CASE_3'(s('type.option.option'(X22),'const.option.NONE_0'),s('type.option.option'(X21),X23),s(fun(X22,'type.option.option'(X21)),X24)),X21)) | s('type.option.option'(X21),'const.option.NONE_0') = s('type.option.option'(X21),'const.option.option_CASE_3'(s('type.option.option'(X22),'const.option.NONE_0'),s('type.option.option'(X21),X23),s(fun(X22,'type.option.option'(X21)),X24)))) )),
  inference(superposition,[],[f611,f163])).
fof(f9837,plain,(
  ( ! [X17,X18,X16] : (s('type.option.option'(X16),X17) != s('type.option.option'(X16),'const.option.SOME_1'(s(X16,X18))) | s(X16,X18) = s(X16,sK7('const.option.SOME_1'(s(X16,sK7(X17,X16))),X16)) | s('type.option.option'(X16),'const.option.NONE_0') = s('type.option.option'(X16),X17)) )),
  inference(subsumption_resolution,[],[f9769,f151])).
fof(f9769,plain,(
  ( ! [X17,X18,X16] : (s('type.option.option'(X16),X17) != s('type.option.option'(X16),'const.option.SOME_1'(s(X16,X18))) | s(X16,X18) = s(X16,sK7('const.option.SOME_1'(s(X16,sK7(X17,X16))),X16)) | s('type.option.option'(X16),'const.option.NONE_0') = s('type.option.option'(X16),'const.option.SOME_1'(s(X16,sK7(X17,X16)))) | s('type.option.option'(X16),'const.option.NONE_0') = s('type.option.option'(X16),X17)) )),
  inference(superposition,[],[f611,f146])).
fof(f9374,plain,(
  ( ! [X4,X5,X3] : (s('type.option.option'(bool),X3) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,sK7(X3,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X3)) )),
  inference(superposition,[],[f497,f610])).
fof(f9311,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),X1) | s(bool,'T_0') = s(bool,sK7(X1,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X1) | s(bool,'T_0') = s(bool,sK7(X0,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f610,f610])).
fof(f1200,plain,(
  ( ! [X18] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1174,f151])).
fof(f1174,plain,(
  ( ! [X18] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f146,f1093])).
fof(f4196,plain,(
  ( ! [X37,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X37)) )),
  inference(forward_demodulation,[],[f3178,f946])).
fof(f3178,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,'F_0') | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X37)),bool)))) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X37)) )),
  inference(superposition,[],[f537,f1093])).
fof(f3920,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22) | s(bool,X20) = s(bool,X22) | s(bool,X20) != s(bool,X21)) )),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f588,f621,f900,f571,f898,f895,f893,f673,f671,f576,f573,f1070,f1048,f788,f760,f781,f1060,f1143,f1135,f1119,f912,f735,f733,f683,f681,f585,f568,f158,f157,f156,f1062,f1061,f1032,f1018,f1016,f934,f936,f1335,f1331,f1308,f1285,f1284,f1283,f1253,f1204,f1203,f1201,f1200,f1199,f1197,f1196,f1195,f1176,f1093,f1020,f938,f897,f892,f737,f675,f575,f570,f537,f492,f239,f1980,f1981,f1982,f1983,f1984,f1995,f1996,f1997,f1998,f1999,f2161,f2167,f2175,f2176,f2177,f2178,f2179,f2183,f2189,f2191,f2193,f2194,f2205,f2206,f2212,f2214,f2216,f2217,f2249,f2250,f3068,f3069,f3186,f3230,f3231,f3252,f3253,f3278,f3279,f3297,f3298,f3299,f3300,f3337,f2263,f3276,f3274,f3267,f3250,f3248,f3242,f3229,f3099,f3070,f3067,f2252,f2251,f2248,f965,f933,f610,f532,f497,f113,f2260,f3078])).
fof(f3078,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f532,f190])).
fof(f113,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb',aHL_FALSITY)).
fof(f3058,plain,(
  ( ! [X19,X17,X18] : (s(bool,X19) != s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,'T_0') = s(bool,X17) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X17)),bool))) ) | $spl0),
  inference(superposition,[],[f532,f1348])).
fof(f3223,plain,(
  ( ! [X12,X11] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X11) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X12,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X11)),bool))) ) | $spl0),
  inference(superposition,[],[f3099,f1348])).
fof(f1433,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X44)),bool)) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X43),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X43,sK0)) | s(bool,'T_0') = s(bool,X44)) ) | $spl0),
  inference(superposition,[],[f1093,f928])).
fof(f1432,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.NONE_0'),s(bool,X40),s(fun(X41,bool),X42))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X39,sK0)) | s(bool,'T_0') = s(bool,X40)) ) | $spl0),
  inference(superposition,[],[f239,f928])).
fof(f1246,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1232,f111])).
fof(f1232,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1199,f190])).
fof(f1208,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'('type.option.option'(X0))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.NONE_0'))))) != s('type.option.option'('type.option.option'('type.option.option'(X0))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))))) )),
  inference(unit_resulting_resolution,[],[f488,f154])).
fof(f1091,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X6)),bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f946,f190])).
fof(f1089,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,sK7('const.option.SOME_1'(s(bool,X2)),bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f946,f190])).
fof(f1022,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X5),'const.option.SOME_1'(s(X5,X10))) != s('type.option.option'(X5),'const.option.SOME_1'(s(X5,'const.option.option_CASE_3'(s('type.option.option'(X6),X7),s(X5,X8),s(fun(X6,X5),X9))))) | s(X5,X10) = s(X5,happ(s(fun(fun(X6,X5),X5),happ(s(fun(X5,fun(fun(X6,X5),X5)),happ(s(fun('type.option.option'(X6),fun(X5,fun(fun(X6,X5),X5))),'const.option.option_CASE_0'),s('type.option.option'(X6),X7))),s(X5,X8))),s(fun(X6,X5),X9)))) )),
  inference(superposition,[],[f154,f166])).
fof(f967,plain,(
  ( ! [X8,X7,X9] : (s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X8))) != s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X9))) | s(X7,X9) = s(X7,sK7('const.option.SOME_1'(s(X7,X8)),X7))) )),
  inference(superposition,[],[f154,f603])).
fof(f946,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK7('const.option.SOME_1'(s(X0,X1)),X0))) )),
  inference(unit_resulting_resolution,[],[f603,f154])).
fof(f931,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK9(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f164,f190])).
fof(f929,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK9(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f164,f190])).
fof(f813,plain,(
  ( ! [X0] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X0))) | s(sK0,X0) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)))) ) | $spl0),
  inference(superposition,[],[f154,f170])).
fof(f739,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X5),'const.option.SOME_1'(s(X5,X10))) != s('type.option.option'(X5),'const.option.SOME_1'(s(X5,happ(s(fun(X6,X5),X9),s(X6,X7))))) | s(X5,X10) = s(X5,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7))),s(X5,X8),s(fun(X6,X5),X9)))) )),
  inference(superposition,[],[f154,f165])).
fof(f716,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(X5,happ(s(fun(X3,X5),X6),s(X3,sK7(X4,X3)))) = s(X5,'const.option.option_CASE_3'(s('type.option.option'(X3),X4),s(X5,X7),s(fun(X3,X5),X6))) | s('type.option.option'(X3),'const.option.NONE_0') = s('type.option.option'(X3),X4)) )),
  inference(superposition,[],[f165,f146])).
fof(f619,plain,(
  ( ! [X19,X20,X18] : (s('type.option.option'(X18),X19) != s('type.option.option'(X18),X20) | s('type.option.option'('type.option.option'(X18)),'const.option.SOME_1'(s('type.option.option'(X18),X20))) = s('type.option.option'('type.option.option'(X18)),'const.option.SOME_1'(s('type.option.option'(X18),'const.option.SOME_1'(s(X18,sK7(X19,X18)))))) | s('type.option.option'(X18),'const.option.NONE_0') = s('type.option.option'(X18),X19)) )),
  inference(superposition,[],[f155,f146])).
fof(f617,plain,(
  ( ! [X14,X12,X13] : (s('type.option.option'('type.option.option'(X12)),'const.option.SOME_1'(s('type.option.option'(X12),X13))) != s('type.option.option'('type.option.option'(X12)),'const.option.SOME_1'(s('type.option.option'(X12),X14))) | s('type.option.option'(X12),X14) = s('type.option.option'(X12),'const.option.SOME_1'(s(X12,sK7(X13,X12)))) | s('type.option.option'(X12),'const.option.NONE_0') = s('type.option.option'(X12),X13)) )),
  inference(superposition,[],[f154,f146])).
fof(f611,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X2) = s(X0,sK7(X1,X0)) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f154,f146])).
fof(f608,plain,(
  ( ! [X4,X5] : (s('type.option.option'(bool),X4) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X4) | s(bool,'T_0') = s(bool,sK7(X4,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f146,f190])).
fof(f606,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0) | s(bool,X1) = s(bool,sK7(X0,bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f146,f190])).
fof(f603,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK7('const.option.SOME_1'(s(X0,X1)),X0))))) )),
  inference(unit_resulting_resolution,[],[f151,f146])).
fof(f491,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X4))) | s(X0,X4) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)))) )),
  inference(superposition,[],[f154,f163])).
fof(f488,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.NONE_0'))) != s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))) )),
  inference(unit_resulting_resolution,[],[f151,f154])).
fof(f482,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f346,f111])).
fof(f346,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f162,f190])).
fof(f290,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X17) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X14),'const.option.NONE_0'),s(bool,X15),s(fun(X14,bool),X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f190,f163])).
fof(f280,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f190,f190])).
fof(f270,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(bool,X21),s(fun(X20,bool),X22)))) )),
  inference(superposition,[],[f190,f163])).
fof(f221,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f215,f113])).
fof(f215,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f159,f114])).
fof(f192,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f113,f114])).
fof(f190,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f114,f114])).
fof(f166,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7] : ! [X20,X21,X22] : s(X7,'const.option.option_CASE_3'(s('type.option.option'(X6),X20),s(X7,X21),s(fun(X6,X7),X22))) = s(X7,happ(s(fun(fun(X6,X7),X7),happ(s(fun(X7,fun(fun(X6,X7),X7)),happ(s(fun('type.option.option'(X6),fun(X7,fun(fun(X6,X7),X7))),'const.option.option_CASE_0'),s('type.option.option'(X6),X20))),s(X7,X21))),s(fun(X6,X7),X22)))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','ARITY_0_3_const.option.option_CASE')).
fof(f165,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(X0,X3),s(fun(X1,X0),X4)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(X0,X3),s(fun(X1,X0),X4)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(X0,X3),s(fun(X1,X0),X4)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X7,X6] : ! [X12,X8,X19] : s(X7,happ(s(fun(X6,X7),X19),s(X6,X12))) = s(X7,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))),s(X7,X8),s(fun(X6,X7),X19)))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','thm.option.option_case_def|split|1')).
fof(f164,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb',aHL_EXT)).
fof(f159,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f162,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f151,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : ~s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X12] : ~s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12)))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','thm.option.NOT_NONE_SOME')).
fof(f114,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb',aHL_BOOL_CASES)).
fof(f111,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb',aHL_TRUTH)).
fof(f28914,plain,(
  $false | ($spl0 | $spl2 | $spl188)),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f603,f606,f608,f611,f617,f619,f716,f739,f172,f170,f813,f929,f931,f946,f967,f1022,f1089,f1091,f1208,f1246,f1472,f1467,f1460,f1479,f1481,f1432,f1485,f1433,f1486,f2131,f2466,f3570,f3325,f3575,f3327,f3223,f3290,f3102,f3581,f3058,f3091,f2263,f113,f2260,f3078,f3920,f4196,f1200,f9311,f9374,f9837,f9838,f9839,f9840,f9841,f9775,f9776,f9777,f9778,f9780,f9781,f9817,f9827,f10090,f10100,f10102,f10103,f10104,f10105,f10144,f11648,f11649,f11650,f11651,f11652,f588,f10098,f9836,f9825,f8548,f2252,f2179,f1999,f1984,f1331,f1283,f1253,f1199,f1195,f1093,f933,f610,f239,f9474,f2212,f2205,f2191,f10152,f10121,f10108,f10107,f10069,f10067,f10054,f10053,f9460,f3242,f3186,f3099,f2214,f2206,f2189,f2183,f2178,f2175,f2167,f2161,f1998,f1995,f1983,f1980,f1308,f1201,f1176,f963,f961,f538,f537,f535,f533,f532,f497,f495,f493,f492,f3337,f9809,f9381,f2518,f2517,f2516,f2515,f2440,f2410,f2408,f2393,f2366,f2349,f2283,f2280,f1448,f1436,f1416,f1399,f1368,f1367,f1365,f1363,f1349,f1348,f1345,f1271,f1116,f928,f921,f869,f1045,f2282,f1351,f922,f2281,f1350,f900,f621,f10603,f2194,f2193,f2217,f2216,f3298,f10943,f8516,f3299,f10944,f8517,f1402,f10565,f8518,f8515,f575,f10911,f10894,f10566,f2521,f1335,f1204,f571,f10963,f10945,f10942,f10927,f10924,f10910,f10895,f10582,f10580,f10567,f3300,f3297,f2520,f1441,f1440,f1415,f1401,f1203,f898,f897,f895,f893,f892,f675,f673,f671,f576,f573,f570,f1048,f1070,f10559,f760,f788,f10558,f781,f1060,f10925,f10658,f10560,f10557,f1346,f1284,f1196,f1143,f1135,f1119,f912,f737,f735,f733,f683,f681,f585,f568,f158,f157,f156,f10926,f10620,f2479,f1347,f1285,f1197,f1062,f1061,f1032,f1020,f1018,f1016,f2478,f10124,f9797,f2477,f2359,f1409,f938,f936,f934,f2176,f1996,f1981,f3188,f2395,f2363,f2347,f1438,f1413,f1397,f2177,f1997,f1982,f10923,f10604,f10601,f10595,f10145,f10142,f10136,f10091,f10088,f10082,f10052,f9818,f9815,f9514,f9494,f9493,f9492,f9491,f9465,f9464,f9462,f9392,f9391,f9389,f9387,f9371,f9370,f9368,f9364,f9357,f9356,f9354,f9352,f3279,f3278,f3276,f3274,f3267,f3253,f3252,f3250,f3248,f3231,f3230,f3229,f3070,f3069,f3068,f3067,f2251,f2250,f2249,f2248,f965,f26859,f26860,f26861,f26874,f26876,f26904,f26905,f26906,f26907,f26909,f26910,f26927,f26997,f27046,f27202,f27058,f27060,f27886,f27887,f27888,f27814,f27815,f27828,f27924,f27925,f27860,f27861,f27862,f27863,f27937,f27880,f28336,f28688,f28338,f28339,f28340,f28353,f28700,f28701,f28359,f28702,f28361,f28703,f28447,f28448,f28451,f28453,f28455,f28456,f28462,f28463,f28464,f28465,f28466,f28467,f28469,f28470,f28471,f28472,f28474,f28476,f28481,f28483,f28484,f28485,f28486,f28487,f28488,f28489,f28490,f28491,f28492,f28544,f28545,f28546,f28547,f28548,f28549,f28550,f28552,f28553,f28554,f28559,f28563,f28567,f28569,f28571,f28572,f28578,f28579,f28580,f28581,f28582,f28583,f28589,f28591,f28596,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f28605,f28606,f28607,f28659,f28660,f28661,f28662,f28663,f28664,f28665,f28667,f28668,f28669,f28674,f28677,f28679,f28714,f28750,f28756,f28781,f28784,f28790,f28788,f28793,f28795,f28796,f28797,f28798,f28799,f28800,f28801,f28802,f28803,f28804,f28805,f28806,f28807,f28808,f28809,f28810,f28811,f28812,f28813,f28814,f28815,f28816,f28817,f28818,f28819,f28820,f28821,f28822,f28823,f28824,f28825,f28826,f28835,f28836,f28837,f28838,f28839,f28840,f28841,f28842,f28843,f28844,f28845,f28846,f28847,f28848,f28849,f28850,f28851,f28852,f28853,f28854,f28855,f28856,f28857,f28858,f28863,f28865,f28902,f28904,f28905,f28906,f28907,f28908,f28909,f28910,f28911,f28874,f28872,f28898])).
fof(f28913,plain,(
  $false | ($spl0 | $spl2 | $spl188)),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f603,f606,f608,f611,f617,f619,f716,f739,f172,f170,f813,f929,f931,f946,f967,f1022,f1089,f1091,f1208,f1246,f1472,f1467,f1460,f1479,f1481,f1432,f1485,f1433,f1486,f2131,f2466,f3570,f3325,f3575,f3327,f3223,f3290,f3102,f3581,f3058,f3091,f2263,f113,f2260,f3078,f3920,f4196,f1200,f9311,f9374,f9837,f9838,f9839,f9840,f9841,f9775,f9776,f9777,f9778,f9780,f9781,f9817,f9827,f10090,f10100,f10102,f10103,f10104,f10105,f10144,f11648,f11649,f11650,f11651,f11652,f588,f10098,f9836,f9825,f8548,f2252,f2179,f1999,f1984,f1331,f1283,f1253,f1199,f1195,f1093,f933,f610,f239,f9474,f2212,f2205,f2191,f10152,f10121,f10108,f10107,f10069,f10067,f10054,f10053,f9460,f3242,f3186,f3099,f2214,f2206,f2189,f2183,f2178,f2175,f2167,f2161,f1998,f1995,f1983,f1980,f1308,f1201,f1176,f963,f961,f538,f537,f535,f533,f532,f497,f495,f493,f492,f3337,f9809,f9381,f2518,f2517,f2516,f2515,f2440,f2410,f2408,f2393,f2366,f2349,f2283,f2280,f1448,f1436,f1416,f1399,f1368,f1367,f1365,f1363,f1349,f1348,f1345,f1271,f1116,f928,f921,f869,f1045,f2282,f1351,f922,f2281,f1350,f900,f621,f10603,f2194,f2193,f2217,f2216,f3298,f10943,f8516,f3299,f10944,f8517,f1402,f10565,f8518,f8515,f575,f10911,f10894,f10566,f2521,f1335,f1204,f571,f10963,f10945,f10942,f10927,f10924,f10910,f10895,f10582,f10580,f10567,f3300,f3297,f2520,f1441,f1440,f1415,f1401,f1203,f898,f897,f895,f893,f892,f675,f673,f671,f576,f573,f570,f1048,f1070,f10559,f760,f788,f10558,f781,f1060,f10925,f10658,f10560,f10557,f1346,f1284,f1196,f1143,f1135,f1119,f912,f737,f735,f733,f683,f681,f585,f568,f158,f157,f156,f10926,f10620,f2479,f1347,f1285,f1197,f1062,f1061,f1032,f1020,f1018,f1016,f2478,f10124,f9797,f2477,f2359,f1409,f938,f936,f934,f2176,f1996,f1981,f3188,f2395,f2363,f2347,f1438,f1413,f1397,f2177,f1997,f1982,f10923,f10604,f10601,f10595,f10145,f10142,f10136,f10091,f10088,f10082,f10052,f9818,f9815,f9514,f9494,f9493,f9492,f9491,f9465,f9464,f9462,f9392,f9391,f9389,f9387,f9371,f9370,f9368,f9364,f9357,f9356,f9354,f9352,f3279,f3278,f3276,f3274,f3267,f3253,f3252,f3250,f3248,f3231,f3230,f3229,f3070,f3069,f3068,f3067,f2251,f2250,f2249,f2248,f965,f26859,f26860,f26861,f26874,f26876,f26904,f26905,f26906,f26907,f26909,f26910,f26927,f26997,f27046,f27202,f27058,f27060,f27886,f27887,f27888,f27814,f27815,f27828,f27924,f27925,f27860,f27861,f27862,f27863,f27937,f27880,f28336,f28688,f28338,f28339,f28340,f28353,f28700,f28701,f28359,f28702,f28361,f28703,f28447,f28448,f28451,f28453,f28455,f28456,f28462,f28463,f28464,f28465,f28466,f28467,f28469,f28470,f28471,f28472,f28474,f28476,f28481,f28483,f28484,f28485,f28486,f28487,f28488,f28489,f28490,f28491,f28492,f28544,f28545,f28546,f28547,f28548,f28549,f28550,f28552,f28553,f28554,f28559,f28563,f28567,f28569,f28571,f28572,f28578,f28579,f28580,f28581,f28582,f28583,f28589,f28591,f28596,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f28605,f28606,f28607,f28659,f28660,f28661,f28662,f28663,f28664,f28665,f28667,f28668,f28669,f28674,f28677,f28679,f28714,f28750,f28756,f28781,f28784,f28790,f28788,f28793,f28795,f28796,f28797,f28798,f28799,f28800,f28801,f28802,f28803,f28804,f28805,f28806,f28807,f28808,f28809,f28810,f28811,f28812,f28813,f28814,f28815,f28816,f28817,f28818,f28819,f28820,f28821,f28822,f28823,f28824,f28825,f28826,f28835,f28836,f28837,f28838,f28839,f28840,f28841,f28842,f28843,f28844,f28845,f28846,f28847,f28848,f28849,f28850,f28851,f28852,f28853,f28854,f28855,f28856,f28857,f28858,f28863,f28865,f28902,f28904,f28905,f28906,f28907,f28908,f28909,f28910,f28911,f28874,f28872])).
fof(f28912,plain,(
  $false | ($spl0 | $spl2 | $spl188)),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f603,f606,f608,f611,f617,f619,f716,f739,f172,f170,f813,f929,f931,f946,f967,f1022,f1089,f1091,f1208,f1246,f1472,f1467,f1460,f1479,f1481,f1432,f1485,f1433,f1486,f2131,f2466,f3570,f3325,f3575,f3327,f3223,f3290,f3102,f3581,f3058,f3091,f2263,f113,f2260,f3078,f3920,f4196,f1200,f9311,f9374,f9837,f9838,f9839,f9840,f9841,f9775,f9776,f9777,f9778,f9780,f9781,f9817,f9827,f10090,f10100,f10102,f10103,f10104,f10105,f10144,f11648,f11649,f11650,f11651,f11652,f588,f10098,f9836,f9825,f8548,f2252,f2179,f1999,f1984,f1331,f1283,f1253,f1199,f1195,f1093,f933,f610,f239,f9474,f2212,f2205,f2191,f10152,f10121,f10108,f10107,f10069,f10067,f10054,f10053,f9460,f3242,f3186,f3099,f2214,f2206,f2189,f2183,f2178,f2175,f2167,f2161,f1998,f1995,f1983,f1980,f1308,f1201,f1176,f963,f961,f538,f537,f535,f533,f532,f497,f495,f493,f492,f3337,f9809,f9381,f2518,f2517,f2516,f2515,f2440,f2410,f2408,f2393,f2366,f2349,f2283,f2280,f1448,f1436,f1416,f1399,f1368,f1367,f1365,f1363,f1349,f1348,f1345,f1271,f1116,f928,f921,f869,f1045,f2282,f1351,f922,f2281,f1350,f900,f621,f10603,f2194,f2193,f2217,f2216,f3298,f10943,f8516,f3299,f10944,f8517,f1402,f10565,f8518,f8515,f575,f10911,f10894,f10566,f2521,f1335,f1204,f571,f10963,f10945,f10942,f10927,f10924,f10910,f10895,f10582,f10580,f10567,f3300,f3297,f2520,f1441,f1440,f1415,f1401,f1203,f898,f897,f895,f893,f892,f675,f673,f671,f576,f573,f570,f1048,f1070,f10559,f760,f788,f10558,f781,f1060,f10925,f10658,f10560,f10557,f1346,f1284,f1196,f1143,f1135,f1119,f912,f737,f735,f733,f683,f681,f585,f568,f158,f157,f156,f10926,f10620,f2479,f1347,f1285,f1197,f1062,f1061,f1032,f1020,f1018,f1016,f2478,f10124,f9797,f2477,f2359,f1409,f938,f936,f934,f2176,f1996,f1981,f3188,f2395,f2363,f2347,f1438,f1413,f1397,f2177,f1997,f1982,f10923,f10604,f10601,f10595,f10145,f10142,f10136,f10091,f10088,f10082,f10052,f9818,f9815,f9514,f9494,f9493,f9492,f9491,f9465,f9464,f9462,f9392,f9391,f9389,f9387,f9371,f9370,f9368,f9364,f9357,f9356,f9354,f9352,f3279,f3278,f3276,f3274,f3267,f3253,f3252,f3250,f3248,f3231,f3230,f3229,f3070,f3069,f3068,f3067,f2251,f2250,f2249,f2248,f965,f26859,f26860,f26861,f26874,f26876,f26904,f26905,f26906,f26907,f26909,f26910,f26927,f26997,f27046,f27202,f27058,f27060,f27886,f27887,f27888,f27814,f27815,f27828,f27924,f27925,f27860,f27861,f27862,f27863,f27937,f27880,f28336,f28688,f28338,f28339,f28340,f28353,f28700,f28701,f28359,f28702,f28361,f28703,f28447,f28448,f28451,f28453,f28455,f28456,f28462,f28463,f28464,f28465,f28466,f28467,f28469,f28470,f28471,f28472,f28474,f28476,f28481,f28483,f28484,f28485,f28486,f28487,f28488,f28489,f28490,f28491,f28492,f28544,f28545,f28546,f28547,f28548,f28549,f28550,f28552,f28553,f28554,f28559,f28563,f28567,f28569,f28571,f28572,f28578,f28579,f28580,f28581,f28582,f28583,f28589,f28591,f28596,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f28605,f28606,f28607,f28659,f28660,f28661,f28662,f28663,f28664,f28665,f28667,f28668,f28669,f28674,f28677,f28679,f28714,f28750,f28756,f28781,f28784,f28790,f28788,f28793,f28795,f28796,f28797,f28798,f28799,f28800,f28801,f28802,f28803,f28804,f28805,f28806,f28807,f28808,f28809,f28810,f28811,f28812,f28813,f28814,f28815,f28816,f28817,f28818,f28819,f28820,f28821,f28822,f28823,f28824,f28825,f28826,f28835,f28836,f28837,f28838,f28839,f28840,f28841,f28842,f28843,f28844,f28845,f28846,f28847,f28848,f28849,f28850,f28851,f28852,f28853,f28854,f28855,f28856,f28857,f28858,f28863,f28865,f28902,f28904,f28905,f28906,f28907,f28908,f28909,f28910,f28911,f28874])).
fof(f28901,plain,(
  $false | ($spl0 | $spl2 | $spl7 | $spl188)),
  inference(subsumption_resolution,[],[f28874,f176])).
fof(f28900,plain,(
  $false | ($spl0 | $spl2 | $spl7 | $spl188)),
  inference(subsumption_resolution,[],[f28872,f176])).
fof(f28899,plain,(
  $false | ($spl0 | $spl2 | $spl7 | $spl188)),
  inference(subsumption_resolution,[],[f28898,f792])).
fof(f792,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5))) | $spl7),
  inference(unit_resulting_resolution,[],[f176,f154])).
fof(f28897,plain,(
  $false | ($spl0 | $spl2 | $spl7 | $spl188)),
  inference(subsumption_resolution,[],[f28896,f792])).
fof(f173,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK4) | $spl4),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK4) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f28714,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1)))) | $spl188),
  inference(cnf_transformation,[],[f28714_D])).
fof(f28714_D,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1)))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f28677,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1)))) | s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK4) | $spl0),
  inference(superposition,[],[f170,f716])).
fof(f28448,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK7(sK4,sK1)))) | s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK4) | $spl0),
  inference(superposition,[],[f716,f170])).
fof(f219,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f25968,plain,(
  ( ! [X43] : (s(bool,'F_0') != s(bool,X43) | s(bool,'F_0') = s(bool,g(s(bool,X43)))) ) | $spl186),
  inference(cnf_transformation,[],[f25968_D])).
fof(f25968_D,plain,(
  ( ! [X43] : (s(bool,'F_0') != s(bool,X43) | s(bool,'F_0') = s(bool,g(s(bool,X43)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f25981,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') != s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,'F_0') = s(bool,g(s(bool,X82)))) ) | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f25619,f23113])).
fof(f23113,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(duplicate_literal_removal,[],[f23112])).
fof(f23112,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(condensation,[],[f23111])).
fof(f23111,plain,(
  ( ! [X127,X125,X126] : (s(bool,X126) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X126) | s(bool,X125) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f22931,f10463])).
fof(f10463,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl30 | $spl33)),
  inference(condensation,[],[f10462])).
fof(f10462,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl30 | $spl33)),
  inference(duplicate_literal_removal,[],[f10348])).
fof(f10348,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f7397,f114])).
fof(f7397,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,sK7('const.option.SOME_1'(s(bool,X3)),bool)) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f7298,f190])).
fof(f7298,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl30 | $spl33)),
  inference(subsumption_resolution,[],[f6823,f4267])).
fof(f4267,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(superposition,[],[f2948,f114])).
fof(f2948,plain,(
  s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | $spl33),
  inference(cnf_transformation,[],[f2948_D])).
fof(f2948_D,plain,(
  s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f6823,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool))) ) | $spl30),
  inference(equality_factoring,[],[f2476])).
fof(f22931,plain,(
  ( ! [X127,X125,X126] : (s(bool,X126) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X126) | s(bool,X125) = s(bool,X127) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X125)),bool))) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(superposition,[],[f7397,f22612])).
fof(f22612,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(trivial_inequality_removal,[],[f22594])).
fof(f22594,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(superposition,[],[f21538,f21950])).
fof(f21950,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) ) | $spl178),
  inference(superposition,[],[f21839,f114])).
fof(f21538,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(forward_demodulation,[],[f21335,f14863])).
fof(f21335,plain,(
  ( ! [X0] : (s(bool,g(s(bool,'T_0'))) != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(equality_factoring,[],[f17848])).
fof(f17848,plain,(
  ( ! [X58] : (s(bool,X58) = s(bool,g(s(bool,'T_0'))) | s(bool,X58) = s(bool,f(s(bool,X58)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(subsumption_resolution,[],[f17163,f111])).
fof(f17163,plain,(
  ( ! [X58] : (~p(s(bool,'T_0')) | s(bool,X58) = s(bool,g(s(bool,'T_0'))) | s(bool,X58) = s(bool,f(s(bool,X58)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(superposition,[],[f17014,f14863])).
fof(f17014,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | s(bool,X20) = s(bool,X21) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl109)),
  inference(forward_demodulation,[],[f17013,f946])).
fof(f17013,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | s(bool,X21) = s(bool,sK7('const.option.SOME_1'(s(bool,X20)),bool)) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl109)),
  inference(subsumption_resolution,[],[f16901,f111])).
fof(f16901,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | s(bool,X21) = s(bool,sK7('const.option.SOME_1'(s(bool,X20)),bool)) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl109)),
  inference(superposition,[],[f162,f13917])).
fof(f13917,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X5)),bool)) | s(bool,X5) = s(bool,f(s(bool,X5)))) ) | ($spl30 | $spl33 | $spl109)),
  inference(subsumption_resolution,[],[f13849,f12624])).
fof(f12624,plain,(
  s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | ($spl33 | $spl109)),
  inference(forward_demodulation,[],[f11838,f11817])).
fof(f11817,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl109),
  inference(unit_resulting_resolution,[],[f11666,f114])).
fof(f11838,plain,(
  s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,f(s(bool,'F_0')))),bool)) | ($spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f11666,f4267])).
fof(f13849,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X5)),bool))) ) | ($spl30 | $spl109)),
  inference(superposition,[],[f11817,f2476])).
fof(f25619,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') != s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,X83) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X82)))) ) | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(superposition,[],[f9636,f23043])).
fof(f23043,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X29)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f22818,f22596])).
fof(f22596,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,X2) = s(bool,g(s(bool,'T_0')))) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(superposition,[],[f21538,f21839])).
fof(f22818,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X29))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(superposition,[],[f22612,f19249])).
fof(f19249,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl15 | $spl119)),
  inference(subsumption_resolution,[],[f19079,f4197])).
fof(f4197,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl15),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f200,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f588,f1331,f1308,f1283,f1253,f1201,f1200,f1199,f1195,f1176,f1093,f879,f537,f492,f490,f239,f1980,f1983,f1984,f1995,f1999,f2161,f2167,f2175,f2179,f2183,f2189,f2191,f2205,f2206,f2212,f2214,f3100,f3186,f2263,f3276,f3274,f3267,f3250,f3248,f3242,f3229,f3099,f3070,f3067,f2252,f2251,f2248,f965,f933,f610,f532,f497,f113,f2260,f3078,f3920,f938,f936,f934,f1016,f1018,f1020,f1032,f1061,f1062,f1197,f1285,f1982,f1997,f2177,f2250,f3069,f3231,f156,f157,f158,f568,f585,f681,f683,f733,f735,f737,f912,f1119,f1135,f1143,f1196,f1284,f1981,f1996,f2176,f2249,f3068,f3230,f1060,f781,f1048,f1070,f760,f788,f570,f573,f576,f671,f673,f675,f892,f893,f895,f897,f898,f1203,f3297,f3300,f571,f1204,f1335,f575,f3298,f3299,f900,f621,f3253,f3252,f2194,f2193,f2217,f2216,f3279,f3278,f3337,f1998,f2178,f4196])).
fof(f3100,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl15),
  inference(subsumption_resolution,[],[f3078,f490])).
fof(f490,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl15),
  inference(unit_resulting_resolution,[],[f200,f154])).
fof(f879,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | $spl15),
  inference(unit_resulting_resolution,[],[f490,f154])).
fof(f19079,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X18)) ) | $spl119),
  inference(superposition,[],[f19019,f114])).
fof(f19019,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl119),
  inference(trivial_inequality_removal,[],[f19012])).
fof(f19012,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl119),
  inference(superposition,[],[f11672,f190])).
fof(f11672,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,'F_0'))))) | $spl119),
  inference(cnf_transformation,[],[f11672_D])).
fof(f11672_D,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,'F_0'))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f9636,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl33 | $spl36)),
  inference(condensation,[],[f9635])).
fof(f9635,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f9571,f4267])).
fof(f9571,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(superposition,[],[f4347,f114])).
fof(f4347,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl36),
  inference(superposition,[],[f2954,f190])).
fof(f3569,plain,(
  ( ! [X109] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X109)))) ) | $spl52),
  inference(cnf_transformation,[],[f3569_D])).
fof(f3569_D,plain,(
  ( ! [X109] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X109)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f25967,plain,(
  ( ! [X43,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,'F_0') != s(bool,X43) | s(bool,'F_0') = s(bool,g(s(bool,X43)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f25966,f155])).
fof(f25966,plain,(
  ( ! [X43,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,'F_0') != s(bool,X43) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,g(s(bool,X43)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(forward_demodulation,[],[f25603,f14863])).
fof(f25603,plain,(
  ( ! [X43,X44] : (s(bool,'F_0') != s(bool,X43) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,g(s(bool,X43)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(superposition,[],[f537,f23043])).
fof(f25965,plain,(
  ( ! [X37] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) | s(bool,'F_0') = s(bool,g(s(bool,X37)))) ) | $spl184),
  inference(cnf_transformation,[],[f25965_D])).
fof(f25965_D,plain,(
  ( ! [X37] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) | s(bool,'F_0') = s(bool,g(s(bool,X37)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f25964,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,X38) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) | s(bool,'F_0') = s(bool,g(s(bool,X37)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(forward_demodulation,[],[f25963,f14863])).
fof(f25963,plain,(
  ( ! [X37,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) | s(bool,X38) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X37)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f25600,f22596])).
fof(f25600,plain,(
  ( ! [X37,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) | s(bool,X38) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,g(s(bool,X37)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(superposition,[],[f497,f23043])).
fof(f25954,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,g(s(bool,X12)))) ) | $spl182),
  inference(cnf_transformation,[],[f25954_D])).
fof(f25954_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,g(s(bool,X12)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f25956,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X14)) | s(bool,'F_0') = s(bool,g(s(bool,X14)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f25955,f192])).
fof(f25955,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') = s(bool,g(s(bool,X14)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(forward_demodulation,[],[f25588,f14863])).
fof(f25588,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X14)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(superposition,[],[f159,f23043])).
fof(f25953,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,g(s(bool,X12)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f25952,f192])).
fof(f25952,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,g(s(bool,X12)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(forward_demodulation,[],[f25587,f14863])).
fof(f25587,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X13) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X12)))) ) | ($spl15 | $spl30 | $spl33 | $spl109 | $spl119 | $spl124 | $spl178)),
  inference(superposition,[],[f159,f23043])).
fof(f24999,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41) | s(bool,'F_0') = s(bool,f(s(bool,X41)))) ) | $spl180),
  inference(cnf_transformation,[],[f24999_D])).
fof(f24999_D,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41) | s(bool,'F_0') = s(bool,f(s(bool,X41)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f25024,plain,(
  ( ! [X80,X81] : (s(bool,'F_0') != s(bool,X80) | s(bool,'T_0') = s(bool,X81) | s(bool,'F_0') = s(bool,f(s(bool,X80)))) ) | ($spl30 | $spl33 | $spl36 | $spl109 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f24622,f23113])).
fof(f24622,plain,(
  ( ! [X80,X81] : (s(bool,'F_0') != s(bool,X80) | s(bool,'T_0') = s(bool,X81) | s(bool,X81) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X80)))) ) | ($spl30 | $spl33 | $spl36 | $spl109 | $spl124 | $spl178)),
  inference(superposition,[],[f9636,f22815])).
fof(f22815,plain,(
  ( ! [X26] : (s(bool,X26) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,f(s(bool,X26)))) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(superposition,[],[f22612,f21839])).
fof(f24998,plain,(
  ( ! [X41,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) | s(bool,'F_0') != s(bool,X41) | s(bool,'F_0') = s(bool,f(s(bool,X41)))) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(subsumption_resolution,[],[f24997,f155])).
fof(f24997,plain,(
  ( ! [X41,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) | s(bool,'F_0') != s(bool,X41) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,f(s(bool,X41)))) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(forward_demodulation,[],[f24606,f14863])).
fof(f24606,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') != s(bool,X41) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,f(s(bool,X41)))) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl178)),
  inference(superposition,[],[f537,f22815])).
fof(f21839,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,f(s(bool,X25)))) ) | $spl178),
  inference(cnf_transformation,[],[f21839_D])).
fof(f21839_D,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,f(s(bool,X25)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f16206,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X37)) ) | $spl132),
  inference(cnf_transformation,[],[f16206_D])).
fof(f16206_D,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X37)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f21838,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,f(s(bool,X25)))) ) | ($spl30 | $spl33 | $spl109 | $spl124 | $spl154)),
  inference(subsumption_resolution,[],[f21787,f17782])).
fof(f21787,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,g(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,f(s(bool,X25)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(superposition,[],[f532,f17848])).
fof(f17782,plain,(
  ( ! [X55] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X55))) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | $spl154),
  inference(cnf_transformation,[],[f17782_D])).
fof(f17782_D,plain,(
  ( ! [X55] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X55))) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f21601,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') = s(bool,X41) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40))) | s(bool,X40) = s(bool,f(s(bool,X40)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(duplicate_literal_removal,[],[f21600])).
fof(f21600,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') = s(bool,X41) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,f(s(bool,X40)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(forward_demodulation,[],[f21568,f14863])).
fof(f21568,plain,(
  ( ! [X41,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40))) | s(bool,X41) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,f(s(bool,X40)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(superposition,[],[f497,f17848])).
fof(f21480,plain,(
  ( ! [X65] : (s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,g(s(bool,'T_0')))) ) | $spl176),
  inference(cnf_transformation,[],[f21480_D])).
fof(f21480_D,plain,(
  ( ! [X65] : (s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,g(s(bool,'T_0')))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f17787,plain,(
  ( ! [X64] : (s(bool,'F_0') != s(bool,X64) | s(bool,X64) = s(bool,f(s(bool,X64)))) ) | $spl156),
  inference(cnf_transformation,[],[f17787_D])).
fof(f17787_D,plain,(
  ( ! [X64] : (s(bool,'F_0') != s(bool,X64) | s(bool,X64) = s(bool,f(s(bool,X64)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f21259,plain,(
  ( ! [X64,X65] : (s(bool,'F_0') != s(bool,X64) | s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,g(s(bool,'T_0'))) | s(bool,X64) = s(bool,f(s(bool,X64)))) ) | ($spl30 | $spl33 | $spl36 | $spl109 | $spl124)),
  inference(superposition,[],[f9636,f17848])).
fof(f17773,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,X17) = s(bool,f(s(bool,X17)))) ) | $spl150),
  inference(cnf_transformation,[],[f17773_D])).
fof(f17773_D,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,X17) = s(bool,f(s(bool,X17)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f21457,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X6)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(subsumption_resolution,[],[f21456,f192])).
fof(f21456,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) = s(bool,f(s(bool,X6)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(forward_demodulation,[],[f21233,f14863])).
fof(f21233,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,X6) = s(bool,f(s(bool,X6)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(superposition,[],[f159,f17848])).
fof(f21455,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | p(s(bool,X4)) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(subsumption_resolution,[],[f21454,f192])).
fof(f21454,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(forward_demodulation,[],[f21232,f14863])).
fof(f21232,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X5) = s(bool,g(s(bool,'T_0'))) | s(bool,X4) = s(bool,f(s(bool,X4)))) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(superposition,[],[f159,f17848])).
fof(f21367,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl174),
  inference(cnf_transformation,[],[f21367_D])).
fof(f21367_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f17678,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl136),
  inference(cnf_transformation,[],[f17678_D])).
fof(f17678_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f21102,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(superposition,[],[f17848,f190])).
fof(f21099,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl30 | $spl33 | $spl109 | $spl124)),
  inference(superposition,[],[f17848,f190])).
fof(f18752,plain,(
  $false | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18743,f18706])).
fof(f18706,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(condensation,[],[f18705])).
fof(f18705,plain,(
  ( ! [X15,X11] : (s(bool,X11) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(global_subsumption,[],[f18635,f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f3186,f2263,f113,f2260,f3078,f3920,f2948,f4196,f1200,f4267,f4437,f4652,f4903,f4983,f4989,f4990,f4991,f5158,f5194,f5195,f5500,f5528,f5529,f5554,f5559,f5560,f5581,f5636,f5638,f5652,f5655,f5707,f5709,f5731,f5737,f5739,f5753,f5760,f5762,f5776,f5799,f6026,f5855,f5857,f5871,f5874,f6126,f6302,f6313,f6314,f7662,f8224,f8206,f8406,f8008,f8085,f8491,f8148,f8508,f9311,f9374,f9832,f9833,f9837,f9838,f9839,f9840,f9841,f9775,f9776,f9777,f9778,f9780,f9781,f9844,f9790,f9791,f9792,f9793,f9811,f9813,f9817,f9827,f10053,f10054,f10063,f10064,f10065,f10066,f10067,f10178,f10084,f10086,f10090,f10100,f10102,f10103,f10104,f10105,f10107,f10108,f10193,f10117,f10118,f10119,f10120,f10121,f10138,f10140,f10144,f11648,f11649,f11650,f11651,f11652,f11655,f11656,f11657,f588,f6031,f8419,f5726,f8415,f5723,f5764,f5765,f10597,f10599,f5742,f5741,f900,f621,f10603,f10576,f10566,f5860,f571,f10580,f10579,f10578,f10577,f10567,f8473,f5873,f5859,f5802,f5801,f5725,f5712,f5711,f5654,f5641,f5640,f898,f895,f893,f673,f671,f576,f573,f760,f788,f10558,f1048,f1070,f10559,f781,f1060,f10666,f10658,f10560,f10557,f8282,f6030,f5582,f5555,f5501,f5168,f4438,f1143,f1135,f1119,f912,f735,f733,f683,f681,f585,f568,f158,f157,f156,f10620,f8283,f5583,f5556,f5502,f4439,f1062,f1061,f1032,f1018,f1016,f5867,f5719,f5648,f936,f934,f10971,f10963,f10959,f10958,f10955,f10945,f10944,f10943,f10942,f10927,f10926,f10925,f10924,f10923,f10912,f10911,f10910,f10906,f10905,f10904,f10895,f10894,f10656,f10655,f10636,f10634,f10604,f10602,f10601,f10596,f10595,f10582,f10565,f10204,f10202,f10201,f10196,f10195,f10194,f10192,f10188,f10185,f10184,f10171,f10170,f10168,f10152,f10145,f10143,f10142,f10137,f10136,f10124,f10098,f10091,f10089,f10088,f10083,f10082,f10069,f10052,f9852,f9851,f9847,f9846,f9845,f9843,f9836,f9831,f9825,f9822,f9818,f9816,f9815,f9810,f9809,f9797,f9514,f9494,f9493,f9492,f9491,f9476,f9474,f9465,f9464,f9462,f9460,f9392,f9391,f9389,f9387,f9381,f9371,f9370,f9368,f9364,f9357,f9356,f9354,f9352,f8734,f8675,f8556,f8548,f8518,f8517,f8516,f8515,f7694,f7681,f6207,f6206,f6204,f6202,f6195,f6180,f6179,f6177,f6175,f6169,f5777,f5420,f5200,f5199,f5197,f5196,f5188,f5187,f5184,f5177,f5176,f5175,f5173,f5109,f5068,f5023,f4955,f4949,f4882,f4881,f4863,f4858,f4851,f4845,f4783,f4753,f4716,f4578,f4520,f4486,f4452,f4280,f3337,f3300,f3299,f3298,f3297,f3279,f3278,f3276,f3274,f3267,f3253,f3252,f3250,f3248,f3242,f3231,f3230,f3229,f3099,f3070,f3069,f3068,f3067,f2252,f2251,f2250,f2249,f2248,f2217,f2216,f2214,f2212,f2206,f2205,f2194,f2193,f2191,f2189,f2183,f2179,f2177,f2176,f2175,f2167,f2161,f1999,f1997,f1996,f1995,f1984,f1983,f1982,f1981,f1980,f1335,f1331,f1308,f1285,f1284,f1283,f1253,f1204,f1203,f1201,f1199,f1197,f1196,f1195,f1176,f1093,f1020,f965,f938,f933,f897,f892,f737,f675,f610,f575,f570,f537,f532,f497,f492,f239,f10198,f10133,f10132,f10131,f9806,f9805,f9804,f9803,f8308,f8307,f8306,f8232,f8018,f2178,f1998,f11659,f11673,f17893,f17891,f17895,f17958,f18062,f18063,f18064,f18065,f17971,f17972,f17973,f17974,f17975,f17976,f17977,f17978,f17979,f18066,f18067,f18068,f17999,f18000,f18001,f18002,f18003,f18004,f18005,f18006,f18007,f18008,f18009,f18010,f18011,f18012,f18013,f18069,f18070,f18071,f18072,f18073,f18019,f18022,f18024,f18074,f18075,f18030,f18032,f18076,f18036,f18038,f18042,f18044,f18046,f18047,f18077,f18083,f18088,f18411,f18455])).
fof(f18411,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK9(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl101 | $spl118)),
  inference(condensation,[],[f18410])).
fof(f18410,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK9(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18109,f114])).
fof(f18109,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK9(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f18088,f936])).
fof(f18088,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18087,f17958])).
fof(f18087,plain,(
  ( ! [X2] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17915,f11659])).
fof(f17915,plain,(
  ( ! [X2] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl118),
  inference(superposition,[],[f492,f11673])).
fof(f18047,plain,(
  ( ! [X81,X82] : (s('type.option.option'(bool),X81) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ~p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7(X81,bool),X82,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X81)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10923])).
fof(f18046,plain,(
  ( ! [X210,X208,X211,X209] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X209),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | s('type.option.option'(bool),X208) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X210),'const.option.NONE_0'),s('type.option.option'(bool),X208),s(fun(X210,'type.option.option'(bool)),X211)),X209,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,sK7(X208,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X208)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10604])).
fof(f18044,plain,(
  ( ! [X198,X196,X197,X195] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X197),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(fun(X197,'type.option.option'(bool)),X198)),X196,'type.option.option'(bool))) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X196),s('type.option.option'(bool),X195)))) | s(bool,'T_0') = s(bool,sK7(X195,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X195)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10601])).
fof(f18042,plain,(
  ( ! [X171,X169,X172,X170] : (s('type.option.option'(bool),X169) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ~p(s(bool,happ(s(fun(bool,bool),X170),s(bool,'T_0')))) | s(bool,sK7(X169,bool)) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X171),'const.option.NONE_0'),s(bool,sK7(X169,bool)),s(fun(X171,bool),X172)),X170,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X169)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10595])).
fof(f18038,plain,(
  ( ! [X210,X208,X211,X209] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X209))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X208) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X210),'const.option.NONE_0'),s('type.option.option'(bool),X209),s(fun(X210,'type.option.option'(bool)),X211))) | s(bool,'T_0') = s(bool,sK7(X208,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X208)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10145])).
fof(f18036,plain,(
  ( ! [X198,X196,X197,X195] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X197),'const.option.NONE_0'),s('type.option.option'(bool),X196),s(fun(X197,'type.option.option'(bool)),X198))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X195))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X196))) | s(bool,'T_0') = s(bool,sK7(X195,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X195)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10142])).
fof(f18076,plain,(
  ( ! [X171,X169,X172,X170] : (s('type.option.option'(bool),X169) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK7(X169,bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X171),'const.option.NONE_0'),s(bool,X170),s(fun(X171,bool),X172))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X169)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18034,f17958])).
fof(f18034,plain,(
  ( ! [X171,X169,X172,X170] : (s('type.option.option'(bool),X169) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X170))) | s(bool,sK7(X169,bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X171),'const.option.NONE_0'),s(bool,X170),s(fun(X171,bool),X172))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X169)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10136])).
fof(f18032,plain,(
  ( ! [X210,X208,X211,X209] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X209))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X209) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X210),'const.option.NONE_0'),s('type.option.option'(bool),X208),s(fun(X210,'type.option.option'(bool)),X211))) | s(bool,'T_0') = s(bool,sK7(X208,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X208)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10091])).
fof(f18030,plain,(
  ( ! [X198,X196,X197,X195] : (s('type.option.option'(bool),X196) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X197),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(fun(X197,'type.option.option'(bool)),X198))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X195))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X196))) | s(bool,'T_0') = s(bool,sK7(X195,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X195)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10088])).
fof(f18075,plain,(
  ( ! [X171,X169,X172,X170] : (s('type.option.option'(bool),X169) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X170) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X171),'const.option.NONE_0'),s(bool,sK7(X169,bool)),s(fun(X171,bool),X172))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X169)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18028,f17958])).
fof(f18028,plain,(
  ( ! [X171,X169,X172,X170] : (s('type.option.option'(bool),X169) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X170))) | s(bool,X170) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X171),'const.option.NONE_0'),s(bool,sK7(X169,bool)),s(fun(X171,bool),X172))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X169)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10082])).
fof(f18074,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.NONE_0'),s(bool,X28),s(fun(X30,bool),X31))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18027,f17958])).
fof(f18027,plain,(
  ( ! [X30,X28,X31,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,X29) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.NONE_0'),s(bool,X28),s(fun(X30,bool),X31))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f10052])).
fof(f18024,plain,(
  ( ! [X123,X122] : (s('type.option.option'('type.option.option'(bool)),X123) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X122) = s('type.option.option'(bool),sK7(X123,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X123) | s(bool,'T_0') = s(bool,sK7(X122,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X122)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9818])).
fof(f18022,plain,(
  ( ! [X116,X115] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK7(X116,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),X116) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X115))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X116) | s(bool,'T_0') = s(bool,sK7(X115,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X115)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9815])).
fof(f18019,plain,(
  ( ! [X26,X25] : (s('type.option.option'(bool),X25) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,sK7(X25,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X25)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9514])).
fof(f18073,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X21) | s(bool,'T_0') = s(bool,sK7(sK7('const.option.SOME_1'(s('type.option.option'(bool),X21)),'type.option.option'(bool)),bool))) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18018,f9841])).
fof(f18018,plain,(
  ( ! [X21] : (s('type.option.option'(bool),X21) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X21) | s(bool,'T_0') = s(bool,sK7(sK7('const.option.SOME_1'(s('type.option.option'(bool),X21)),'type.option.option'(bool)),bool))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9494])).
fof(f18072,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X17),X18),s('type.option.option'(bool),X19),s(fun(X17,'type.option.option'(bool)),X20))) | s(bool,'T_0') = s(bool,sK7(happ(s(fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)),happ(s(fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool))),happ(s(fun('type.option.option'(X17),fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)))),'const.option.option_CASE_0'),s('type.option.option'(X17),X18))),s('type.option.option'(bool),X19))),s(fun(X17,'type.option.option'(bool)),X20)),bool))) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18017,f9840])).
fof(f18017,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X17),X18),s('type.option.option'(bool),X19),s(fun(X17,'type.option.option'(bool)),X20))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X17),X18),s('type.option.option'(bool),X19),s(fun(X17,'type.option.option'(bool)),X20))) | s(bool,'T_0') = s(bool,sK7(happ(s(fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)),happ(s(fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool))),happ(s(fun('type.option.option'(X17),fun('type.option.option'(bool),fun(fun(X17,'type.option.option'(bool)),'type.option.option'(bool)))),'const.option.option_CASE_0'),s('type.option.option'(X17),X18))),s('type.option.option'(bool),X19))),s(fun(X17,'type.option.option'(bool)),X20)),bool))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9493])).
fof(f18071,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),happ(s(fun(X13,'type.option.option'(bool)),X16),s(X13,X14))) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X14))),s('type.option.option'(bool),X15),s(fun(X13,'type.option.option'(bool)),X16)),bool))) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18016,f9839])).
fof(f18016,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),happ(s(fun(X13,'type.option.option'(bool)),X16),s(X13,X14))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),happ(s(fun(X13,'type.option.option'(bool)),X16),s(X13,X14))) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X14))),s('type.option.option'(bool),X15),s(fun(X13,'type.option.option'(bool)),X16)),bool))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9492])).
fof(f18070,plain,(
  ( ! [X12,X10,X11] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X11) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X10),'const.option.NONE_0'),s('type.option.option'(bool),X11),s(fun(X10,'type.option.option'(bool)),X12)),bool))) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18015,f9838])).
fof(f18015,plain,(
  ( ! [X12,X10,X11] : (s('type.option.option'(bool),X11) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X11) | s(bool,'T_0') = s(bool,sK7('const.option.option_CASE_3'(s('type.option.option'(X10),'const.option.NONE_0'),s('type.option.option'(bool),X11),s(fun(X10,'type.option.option'(bool)),X12)),bool))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9491])).
fof(f18069,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X66),'const.option.NONE_0'),s(bool,sK7(X65,bool)),s(fun(X66,bool),X67))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X65)) ) | ($spl33 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18014,f10100])).
fof(f18014,plain,(
  ( ! [X66,X67,X65] : (s('type.option.option'(bool),X65) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X66),'const.option.NONE_0'),s(bool,sK7(X65,bool)),s(fun(X66,bool),X67))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X65)) ) | ($spl33 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9476])).
fof(f18013,plain,(
  ( ! [X112,X113] : (s('type.option.option'(bool),X112) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ~p(s(bool,happ(s(fun(bool,bool),X113),s(bool,'T_0')))) | s(bool,sK7(X112,bool)) = s(bool,sK8(sK7(X112,bool),X113,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X112)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9465])).
fof(f18012,plain,(
  ( ! [X111,X109,X110] : (s('type.option.option'(bool),X109) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ~p(s(bool,happ(s(fun(bool,bool),X110),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X110),s(bool,X111)))) | s(bool,X111) != s(bool,sK7(X109,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X109)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9464])).
fof(f18011,plain,(
  ( ! [X105,X106] : (s('type.option.option'(bool),X105) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X106) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X106))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7(X105,bool)))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X105)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9462])).
fof(f18010,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,'F_0')),X43,'type.option.option'(bool))) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X43),s('type.option.option'(bool),X42)))) | s(bool,'T_0') = s(bool,sK7(X42,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X42)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9392])).
fof(f18009,plain,(
  ( ! [X39,X41,X40] : (s('type.option.option'(bool),X41) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X40),s('type.option.option'(bool),X39)))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X40),s('type.option.option'(bool),X41)))) | s(bool,'T_0') = s(bool,sK7(X39,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X39)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9391])).
fof(f18008,plain,(
  ( ! [X35,X36] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X36))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X35) != s('type.option.option'(bool),X36) | s(bool,'T_0') = s(bool,sK7(X35,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X35)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9389])).
fof(f18007,plain,(
  ( ! [X31,X32] : (s('type.option.option'(bool),X32) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X31))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X32))) | s(bool,'T_0') = s(bool,sK7(X31,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X31)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9387])).
fof(f18006,plain,(
  ( ! [X80,X81,X82] : (s(X82,f(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s(X82,g(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s(fun('type.option.option'(bool),X82),X80) = s(fun('type.option.option'(bool),X82),X81) | s(bool,'T_0') = s(bool,sK7(sK9(X80,X81,X82,'type.option.option'(bool)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),sK9(X80,X81,X82,'type.option.option'(bool)))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9371])).
fof(f18005,plain,(
  ( ! [X78,X79] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X79),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X79),s('type.option.option'(bool),X78)))) | s(bool,'T_0') = s(bool,sK7(sK8(X78,X79,'type.option.option'(bool)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),sK8(X78,X79,'type.option.option'(bool)))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9370])).
fof(f18004,plain,(
  ( ! [X76] : (s('type.option.option'('type.option.option'(bool)),X76) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X76) | s(bool,'T_0') = s(bool,sK7(sK7(X76,'type.option.option'(bool)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),sK7(X76,'type.option.option'(bool)))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9368])).
fof(f18003,plain,(
  ( ! [X64,X62,X63] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),f(s(X62,sK9(X63,X64,'type.option.option'(bool),X62)))) | s(fun(X62,'type.option.option'(bool)),X63) = s(fun(X62,'type.option.option'(bool)),X64) | s(bool,'T_0') = s(bool,sK7(g(s(X62,sK9(X63,X64,'type.option.option'(bool),X62))),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),g(s(X62,sK9(X63,X64,'type.option.option'(bool),X62))))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9364])).
fof(f18002,plain,(
  ( ! [X41,X40] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X41),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | s('type.option.option'(bool),X40) = s('type.option.option'(bool),sK8(X40,X41,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,sK7(X40,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X40)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9357])).
fof(f18001,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X38),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X38),s('type.option.option'(bool),X39)))) | s('type.option.option'(bool),X37) != s('type.option.option'(bool),X39) | s(bool,'T_0') = s(bool,sK7(X37,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X37)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9356])).
fof(f18000,plain,(
  ( ! [X33,X34] : (s('type.option.option'(bool),X34) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X33))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X34))) | s(bool,'T_0') = s(bool,sK7(X33,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X33)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9354])).
fof(f17999,plain,(
  ( ! [X30,X29] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X30))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X29) = s('type.option.option'(bool),X30) | s(bool,'T_0') = s(bool,sK7(X29,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X29)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f9352])).
fof(f18068,plain,(
  ( ! [X99,X98] : (s(bool,X99) = s(bool,sK7('const.option.SOME_1'(s(bool,X98)),bool)) | s(bool,'T_0') = s(bool,X98)) ) | ($spl33 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17998,f4983])).
fof(f17998,plain,(
  ( ! [X99,X98] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X99))) | s(bool,X99) = s(bool,sK7('const.option.SOME_1'(s(bool,X98)),bool)) | s(bool,'T_0') = s(bool,X98)) ) | ($spl33 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f8675])).
fof(f18067,plain,(
  ( ! [X82] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X82) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X82,bool))),bool))) ) | ($spl33 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17983,f9837])).
fof(f17983,plain,(
  ( ! [X82] : (s('type.option.option'(bool),X82) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X82) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X82,bool))),bool))) ) | ($spl33 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f4858])).
fof(f18066,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X35)),bool))) ) | ($spl33 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17982,f4990])).
fof(f17982,plain,(
  ( ! [X35,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X35)),bool))) ) | ($spl33 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f4845])).
fof(f17979,plain,(
  ( ! [X37,X36] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X37),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,X36)),X37,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3279])).
fof(f17978,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X34),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X34),s('type.option.option'(bool),X35)))) | s('type.option.option'(bool),X35) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3278])).
fof(f17977,plain,(
  ( ! [X30,X29] : (s('type.option.option'(bool),X30) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X30))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3276])).
fof(f17976,plain,(
  ( ! [X26,X25] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X26))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X26) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3274])).
fof(f17975,plain,(
  ( ! [X12,X10,X13,X11] : (s(X11,happ(s(fun(bool,X11),X12),s(bool,X10))) = s(X11,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(X11,X13),s(fun(bool,X11),X12))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3267])).
fof(f17974,plain,(
  ( ! [X37,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,'F_0')),X38,'type.option.option'(bool))) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X38),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37)))))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3253])).
fof(f17973,plain,(
  ( ! [X35,X36,X34] : (s('type.option.option'(bool),X36) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X35),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34)))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X35),s('type.option.option'(bool),X36)))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3252])).
fof(f17972,plain,(
  ( ! [X30,X31] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X31))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X31) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3250])).
fof(f17971,plain,(
  ( ! [X26,X27] : (s('type.option.option'(bool),X27) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X27))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))))) | s(bool,'T_0') = s(bool,X26)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f3248])).
fof(f18065,plain,(
  ( ! [X45,X44] : (s(bool,X45) = s(bool,sK7('const.option.SOME_1'(s(bool,X44)),bool)) | s(bool,'T_0') = s(bool,X45)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17963,f961])).
fof(f17963,plain,(
  ( ! [X45,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,X45) = s(bool,sK7('const.option.SOME_1'(s(bool,X44)),bool)) | s(bool,'T_0') = s(bool,X45)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f2251])).
fof(f18064,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X43) = s(bool,happ(s(fun(fun(X39,bool),bool),happ(s(fun(bool,fun(fun(X39,bool),bool)),happ(s(fun('type.option.option'(X39),fun(bool,fun(fun(X39,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X39),X40))),s(bool,X41))),s(fun(X39,bool),X42))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17962,f17958])).
fof(f17962,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X39),X40),s(bool,X41),s(fun(X39,bool),X42))))) | s(bool,X43) = s(bool,happ(s(fun(fun(X39,bool),bool),happ(s(fun(bool,fun(fun(X39,bool),bool)),happ(s(fun('type.option.option'(X39),fun(bool,fun(fun(X39,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X39),X40))),s(bool,X41))),s(fun(X39,bool),X42))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f2250])).
fof(f18063,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X38) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.SOME_1'(s(X34,X35))),s(bool,X36),s(fun(X34,bool),X37))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17961,f17958])).
fof(f17961,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X34,bool),X37),s(X34,X35))))) | s(bool,X38) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.SOME_1'(s(X34,X35))),s(bool,X36),s(fun(X34,bool),X37))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f2249])).
fof(f18062,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X33) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.NONE_0'),s(bool,X31),s(fun(X30,bool),X32))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17960,f17958])).
fof(f17960,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,X33) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.NONE_0'),s(bool,X31),s(fun(X30,bool),X32))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f2248])).
fof(f17958,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17957,f11659])).
fof(f17957,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl118),
  inference(subsumption_resolution,[],[f17907,f155])).
fof(f17907,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl118),
  inference(superposition,[],[f11673,f190])).
fof(f17895,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl118),
  inference(superposition,[],[f11673,f114])).
fof(f17891,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(bool,g(s(bool,'F_0'))),s(fun(X0,bool),X1)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f11673,f491])).
fof(f17893,plain,(
  s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(bool,'F_0')))),bool)) | $spl118),
  inference(unit_resulting_resolution,[],[f151,f11673,f611])).
fof(f11659,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | $spl101),
  inference(cnf_transformation,[],[f11659_D])).
fof(f11659_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f8018,plain,(
  ( ! [X241,X239,X242,X240] : (s(bool,sK7('const.option.SOME_1'(s(bool,X242)),bool)) != s(bool,f(s(X239,sK9(X240,X241,bool,X239)))) | s(fun(X239,bool),X240) = s(fun(X239,bool),X241) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(X239,sK9(X240,X241,bool,X239))))),bool)) | s(bool,'T_0') = s(bool,X242)) ) | $spl33),
  inference(superposition,[],[f164,f4652])).
fof(f8232,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(condensation,[],[f8231])).
fof(f8231,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(subsumption_resolution,[],[f7815,f4267])).
fof(f7815,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(superposition,[],[f4652,f114])).
fof(f8306,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'T_0') = s(bool,X74) | s(bool,sK7('const.option.SOME_1'(s(bool,X76)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,X74),s(fun(X73,bool),X75))) | s(bool,'T_0') = s(bool,X76)) ) | $spl33),
  inference(forward_demodulation,[],[f7923,f946])).
fof(f7923,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X74)),bool)) | s(bool,sK7('const.option.SOME_1'(s(bool,X76)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,X74),s(fun(X73,bool),X75))) | s(bool,'T_0') = s(bool,X76)) ) | $spl33),
  inference(superposition,[],[f4652,f163])).
fof(f8307,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') = s(bool,happ(s(fun(X77,bool),X80),s(X77,X78))) | s(bool,sK7('const.option.SOME_1'(s(bool,X81)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.SOME_1'(s(X77,X78))),s(bool,X79),s(fun(X77,bool),X80))) | s(bool,'T_0') = s(bool,X81)) ) | $spl33),
  inference(forward_demodulation,[],[f7924,f946])).
fof(f7924,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,happ(s(fun(X77,bool),X80),s(X77,X78)))),bool)) | s(bool,sK7('const.option.SOME_1'(s(bool,X81)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.SOME_1'(s(X77,X78))),s(bool,X79),s(fun(X77,bool),X80))) | s(bool,'T_0') = s(bool,X81)) ) | $spl33),
  inference(superposition,[],[f4652,f165])).
fof(f8308,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X82),X83),s(bool,X84),s(fun(X82,bool),X85))) | s(bool,sK7('const.option.SOME_1'(s(bool,X86)),bool)) = s(bool,happ(s(fun(fun(X82,bool),bool),happ(s(fun(bool,fun(fun(X82,bool),bool)),happ(s(fun('type.option.option'(X82),fun(bool,fun(fun(X82,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X82),X83))),s(bool,X84))),s(fun(X82,bool),X85))) | s(bool,'T_0') = s(bool,X86)) ) | $spl33),
  inference(forward_demodulation,[],[f7925,f946])).
fof(f7925,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X82),X83),s(bool,X84),s(fun(X82,bool),X85)))),bool)) | s(bool,sK7('const.option.SOME_1'(s(bool,X86)),bool)) = s(bool,happ(s(fun(fun(X82,bool),bool),happ(s(fun(bool,fun(fun(X82,bool),bool)),happ(s(fun('type.option.option'(X82),fun(bool,fun(fun(X82,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X82),X83))),s(bool,X84))),s(fun(X82,bool),X85))) | s(bool,'T_0') = s(bool,X86)) ) | $spl33),
  inference(superposition,[],[f4652,f166])).
fof(f9803,plain,(
  ( ! [X85,X86] : (s('type.option.option'(bool),X86) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK7(X86,bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X85)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X86) | s(bool,X85) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f611,f4452])).
fof(f9804,plain,(
  ( ! [X88,X87,X89] : (s('type.option.option'(bool),X89) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X88))) | s(bool,sK7(X89,bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X87)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X89) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X88)),bool)) | s(bool,'T_0') = s(bool,X87)) ) | $spl33),
  inference(superposition,[],[f611,f4652])).
fof(f9805,plain,(
  ( ! [X92,X90,X91] : (s('type.option.option'(bool),X91) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK7(X91,bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X90)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X91) | s(bool,X90) = s(bool,sK7('const.option.SOME_1'(s(bool,X92)),bool)) | s(bool,'T_0') = s(bool,X92)) ) | $spl33),
  inference(superposition,[],[f611,f4652])).
fof(f9806,plain,(
  ( ! [X94,X93] : (s('type.option.option'(bool),X94) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,sK7(X94,bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X93)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X94) | s(bool,'T_0') = s(bool,X93)) ) | $spl33),
  inference(superposition,[],[f611,f8232])).
fof(f10131,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X146))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X147))) | s(bool,sK7('const.option.SOME_1'(s(bool,X145)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X148),'const.option.NONE_0'),s(bool,X147),s(fun(X148,bool),X149))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X146)),bool)) | s(bool,'T_0') = s(bool,X145)) ) | $spl33),
  inference(superposition,[],[f491,f4652])).
fof(f10132,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X151))) | s(bool,sK7('const.option.SOME_1'(s(bool,X150)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X152),'const.option.NONE_0'),s(bool,X151),s(fun(X152,bool),X153))) | s(bool,X150) = s(bool,sK7('const.option.SOME_1'(s(bool,X154)),bool)) | s(bool,'T_0') = s(bool,X154)) ) | $spl33),
  inference(superposition,[],[f491,f4652])).
fof(f10133,plain,(
  ( ! [X158,X156,X157,X155] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X156))) | s(bool,sK7('const.option.SOME_1'(s(bool,X155)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X157),'const.option.NONE_0'),s(bool,X156),s(fun(X157,bool),X158))) | s(bool,'T_0') = s(bool,X155)) ) | $spl33),
  inference(superposition,[],[f491,f8232])).
fof(f10198,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'F_0') = s(bool,X141) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X142))) | s(bool,sK7('const.option.SOME_1'(s(bool,X141)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X143),'const.option.NONE_0'),s(bool,X142),s(fun(X143,bool),X144)))) ) | $spl33),
  inference(forward_demodulation,[],[f10130,f946])).
fof(f10130,plain,(
  ( ! [X144,X142,X143,X141] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X142))) | s(bool,sK7('const.option.SOME_1'(s(bool,X141)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X143),'const.option.NONE_0'),s(bool,X142),s(fun(X143,bool),X144))) | s(bool,X141) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4452])).
fof(f4280,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(trivial_inequality_removal,[],[f4274])).
fof(f4274,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(superposition,[],[f2948,f190])).
fof(f4452,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X25)),bool))) ) | $spl33),
  inference(superposition,[],[f4280,f946])).
fof(f4486,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X32)) ) | $spl33),
  inference(superposition,[],[f155,f4280])).
fof(f4520,plain,(
  ( ! [X26,X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s(bool,X26) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X25)) ) | $spl33),
  inference(superposition,[],[f154,f4280])).
fof(f4578,plain,(
  ( ! [X33,X34] : (s(bool,'F_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,'F_0')),bool),X34,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X33)))) | s(bool,'T_0') = s(bool,X33)) ) | $spl33),
  inference(forward_demodulation,[],[f4524,f946])).
fof(f4524,plain,(
  ( ! [X33,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X33)))) | s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,'F_0')),bool),X34,bool)) | s(bool,'T_0') = s(bool,X33)) ) | $spl33),
  inference(superposition,[],[f157,f4280])).
fof(f4716,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X39)),bool))) ) | $spl33),
  inference(superposition,[],[f155,f4452])).
fof(f4753,plain,(
  ( ! [X30,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) | s(bool,X30) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X29)),bool))) ) | $spl33),
  inference(superposition,[],[f154,f4452])).
fof(f4783,plain,(
  ( ! [X26,X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s(bool,X26) = s(bool,sK7('const.option.SOME_1'(s(bool,X25)),bool)) | s(bool,X25) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f154,f4452])).
fof(f4845,plain,(
  ( ! [X35,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X35)),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f4714,f946])).
fof(f4714,plain,(
  ( ! [X35,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X35)),bool))) ) | $spl33),
  inference(superposition,[],[f154,f4452])).
fof(f4851,plain,(
  ( ! [X70,X68,X69] : (s(bool,'F_0') != s(bool,f(s(X68,sK9(X69,X70,bool,X68)))) | s(fun(X68,bool),X69) = s(fun(X68,bool),X70) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(X68,sK9(X69,X70,bool,X68))))),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f4726,f946])).
fof(f4726,plain,(
  ( ! [X70,X68,X69] : (s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) != s(bool,f(s(X68,sK9(X69,X70,bool,X68)))) | s(fun(X68,bool),X69) = s(fun(X68,bool),X70) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(X68,sK9(X69,X70,bool,X68))))),bool))) ) | $spl33),
  inference(superposition,[],[f164,f4452])).
fof(f4858,plain,(
  ( ! [X82] : (s('type.option.option'(bool),X82) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X82) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X82,bool))),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f4730,f946])).
fof(f4730,plain,(
  ( ! [X82] : (s('type.option.option'(bool),X82) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X82) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X82,bool))),bool))) ) | $spl33),
  inference(superposition,[],[f146,f4452])).
fof(f4863,plain,(
  ( ! [X88,X87,X86] : (s(X88,f(s(bool,'F_0'))) != s(X88,g(s(bool,'F_0'))) | s(fun(bool,X88),X86) = s(fun(bool,X88),X87) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK9(X86,X87,X88,bool))),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f4733,f946])).
fof(f4733,plain,(
  ( ! [X88,X87,X86] : (s(X88,f(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) != s(X88,g(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(fun(bool,X88),X86) = s(fun(bool,X88),X87) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK9(X86,X87,X88,bool))),bool))) ) | $spl33),
  inference(superposition,[],[f164,f4452])).
fof(f4881,plain,(
  ( ! [X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,X23) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(subsumption_resolution,[],[f4781,f151])).
fof(f4781,plain,(
  ( ! [X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,X23) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f146,f4452])).
fof(f4882,plain,(
  ( ! [X30,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) | s(bool,'T_0') != s(bool,X30) | s(bool,X29) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f4785,f946])).
fof(f4785,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,X30) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X29)),bool)))) | s(bool,X29) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f155,f4452])).
fof(f4949,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl33),
  inference(superposition,[],[f4881,f190])).
fof(f4955,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,X0) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f4881,f4881])).
fof(f5023,plain,(
  ( ! [X6,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,X5) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f154,f4881])).
fof(f5068,plain,(
  ( ! [X47,X48] : (s(bool,X48) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47)))) ) | $spl33),
  inference(superposition,[],[f155,f4881])).
fof(f5109,plain,(
  ( ! [X37,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) | s(bool,X38) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37)))) ) | $spl33),
  inference(superposition,[],[f154,f4881])).
fof(f5173,plain,(
  ( ! [X43,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,X43) = s(bool,X44) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43)))) ) | $spl33),
  inference(forward_demodulation,[],[f5066,f946])).
fof(f5066,plain,(
  ( ! [X43,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(bool,X43) = s(bool,X44) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43)))) ) | $spl33),
  inference(superposition,[],[f154,f4881])).
fof(f5175,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X52),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X52),s(bool,X53)))) | s(bool,X51) != s(bool,X53) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51)))) ) | $spl33),
  inference(forward_demodulation,[],[f5070,f946])).
fof(f5070,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X52),s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))))) | p(s(bool,happ(s(fun(bool,bool),X52),s(bool,X53)))) | s(bool,X51) != s(bool,X53) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51)))) ) | $spl33),
  inference(superposition,[],[f156,f4881])).
fof(f5176,plain,(
  ( ! [X54,X55] : (~p(s(bool,happ(s(fun(bool,bool),X55),s(bool,'F_0')))) | s(bool,X54) = s(bool,sK8(X54,X55,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54)))) ) | $spl33),
  inference(forward_demodulation,[],[f5071,f946])).
fof(f5071,plain,(
  ( ! [X54,X55] : (~p(s(bool,happ(s(fun(bool,bool),X55),s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))))) | s(bool,X54) = s(bool,sK8(X54,X55,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54)))) ) | $spl33),
  inference(superposition,[],[f157,f4881])).
fof(f5177,plain,(
  ( ! [X78,X76,X77] : (s(bool,'F_0') != s(bool,f(s(X76,sK9(X77,X78,bool,X76)))) | s(fun(X76,bool),X77) = s(fun(X76,bool),X78) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(X76,sK9(X77,X78,bool,X76))))))) ) | $spl33),
  inference(forward_demodulation,[],[f5078,f946])).
fof(f5078,plain,(
  ( ! [X78,X76,X77] : (s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) != s(bool,f(s(X76,sK9(X77,X78,bool,X76)))) | s(fun(X76,bool),X77) = s(fun(X76,bool),X78) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(X76,sK9(X77,X78,bool,X76))))))) ) | $spl33),
  inference(superposition,[],[f164,f4881])).
fof(f5184,plain,(
  ( ! [X90] : (s('type.option.option'(bool),X90) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X90) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7(X90,bool))))) ) | $spl33),
  inference(forward_demodulation,[],[f5082,f946])).
fof(f5082,plain,(
  ( ! [X90] : (s('type.option.option'(bool),X90) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X90) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7(X90,bool))))) ) | $spl33),
  inference(superposition,[],[f146,f4881])).
fof(f5187,plain,(
  ( ! [X92,X93] : (p(s(bool,happ(s(fun(bool,bool),X93),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X93),s(bool,X92)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK8(X92,X93,bool))))) ) | $spl33),
  inference(forward_demodulation,[],[f5084,f946])).
fof(f5084,plain,(
  ( ! [X92,X93] : (p(s(bool,happ(s(fun(bool,bool),X93),s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))))) | ~p(s(bool,happ(s(fun(bool,bool),X93),s(bool,X92)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK8(X92,X93,bool))))) ) | $spl33),
  inference(superposition,[],[f158,f4881])).
fof(f5188,plain,(
  ( ! [X94,X95,X96] : (s(X96,f(s(bool,'F_0'))) != s(X96,g(s(bool,'F_0'))) | s(fun(bool,X96),X94) = s(fun(bool,X96),X95) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK9(X94,X95,X96,bool))))) ) | $spl33),
  inference(forward_demodulation,[],[f5085,f946])).
fof(f5085,plain,(
  ( ! [X94,X95,X96] : (s(X96,f(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) != s(X96,g(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(fun(bool,X96),X94) = s(fun(bool,X96),X95) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK9(X94,X95,X96,bool))))) ) | $spl33),
  inference(superposition,[],[f164,f4881])).
fof(f5196,plain,(
  ( ! [X35] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) | $spl33),
  inference(subsumption_resolution,[],[f5107,f151])).
fof(f5107,plain,(
  ( ! [X35] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) | $spl33),
  inference(superposition,[],[f146,f4881])).
fof(f5197,plain,(
  ( ! [X41,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) | s(bool,X41) != s(bool,X42) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X41)))) ) | $spl33),
  inference(forward_demodulation,[],[f5111,f946])).
fof(f5111,plain,(
  ( ! [X41,X42] : (s(bool,X41) != s(bool,X42) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X41)))) ) | $spl33),
  inference(superposition,[],[f155,f4881])).
fof(f5199,plain,(
  ( ! [X47,X45,X46] : (s(bool,'F_0') != s(bool,X47) | ~p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X45)))) | p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X47)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X45)))) ) | $spl33),
  inference(forward_demodulation,[],[f5113,f946])).
fof(f5113,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X45)))) | p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X47)))) | s(bool,X47) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X45)))) ) | $spl33),
  inference(superposition,[],[f156,f4881])).
fof(f5200,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,'F_0')),bool),X49,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,X48)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48)))) ) | $spl33),
  inference(forward_demodulation,[],[f5114,f946])).
fof(f5114,plain,(
  ( ! [X48,X49] : (~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,X48)))) | s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,'F_0')),bool),X49,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48)))) ) | $spl33),
  inference(superposition,[],[f157,f4881])).
fof(f5420,plain,(
  ( ! [X30,X31] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X31) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30)))) ) | $spl33),
  inference(subsumption_resolution,[],[f5369,f2948])).
fof(f5369,plain,(
  ( ! [X30,X31] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30)))) ) | $spl33),
  inference(superposition,[],[f492,f4881])).
fof(f5777,plain,(
  ( ! [X24,X23,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,X23) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,X22) = s(bool,X24) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X24)))) ) | $spl33),
  inference(superposition,[],[f4881,f4991])).
fof(f6169,plain,(
  ( ! [X14,X17,X15,X16] : (s(X15,happ(s(fun(bool,X15),X16),s(bool,'F_0'))) = s(X15,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X14))),s(X15,X17),s(fun(bool,X15),X16))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X14)))) ) | $spl33),
  inference(superposition,[],[f165,f5196])).
fof(f6175,plain,(
  ( ! [X28,X29] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X29))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28))))) | s('type.option.option'(bool),X29) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28)))) ) | $spl33),
  inference(superposition,[],[f154,f5196])).
fof(f6177,plain,(
  ( ! [X33,X32] : (s('type.option.option'(bool),X33) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X33))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32)))) ) | $spl33),
  inference(superposition,[],[f155,f5196])).
fof(f6179,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X37),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36)))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X37),s('type.option.option'(bool),X38)))) | s('type.option.option'(bool),X38) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36)))) ) | $spl33),
  inference(superposition,[],[f156,f5196])).
fof(f6180,plain,(
  ( ! [X39,X40] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X40),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,'F_0')),X40,'type.option.option'(bool))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39)))) ) | $spl33),
  inference(superposition,[],[f157,f5196])).
fof(f6195,plain,(
  ( ! [X14,X12,X15,X13] : (s(X13,happ(s(fun(bool,X13),X14),s(bool,X12))) = s(X13,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(X13,X15),s(fun(bool,X13),X14))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) | $spl33),
  inference(superposition,[],[f165,f5196])).
fof(f6202,plain,(
  ( ! [X28,X27] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X28))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X28) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) | $spl33),
  inference(superposition,[],[f154,f5196])).
fof(f6204,plain,(
  ( ! [X31,X32] : (s('type.option.option'(bool),X32) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X32))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31)))) ) | $spl33),
  inference(superposition,[],[f155,f5196])).
fof(f6206,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X36),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X36),s('type.option.option'(bool),X37)))) | s('type.option.option'(bool),X37) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) | $spl33),
  inference(superposition,[],[f156,f5196])).
fof(f6207,plain,(
  ( ! [X39,X38] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X39),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,X38)),X39,'type.option.option'(bool))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38)))) ) | $spl33),
  inference(superposition,[],[f157,f5196])).
fof(f7681,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(bool,bool),X48),s(bool,X47)))) | s(bool,'F_0') = s(bool,sK8('F_0',X48,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47)))) ) | $spl33),
  inference(subsumption_resolution,[],[f7639,f2948])).
fof(f7639,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(bool,bool),X48),s(bool,X47)))) | s(bool,'F_0') = s(bool,sK8('F_0',X48,bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47)))) ) | $spl33),
  inference(superposition,[],[f570,f4881])).
fof(f7694,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') = s(bool,sK8('F_0',X26,bool)) | s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),X26),s(bool,X27))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28)))) ) | $spl33),
  inference(subsumption_resolution,[],[f7659,f111])).
fof(f7659,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,sK8('F_0',X26,bool)) | s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),X26),s(bool,X27))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28)))) ) | $spl33),
  inference(superposition,[],[f570,f4991])).
fof(f8556,plain,(
  ( ! [X30,X31,X29] : (s(bool,X30) = s(bool,sK8(X30,X29,bool)) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31)))) ) | $spl33),
  inference(subsumption_resolution,[],[f8541,f111])).
fof(f8541,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,sK8(X30,X29,bool)) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31)))) ) | $spl33),
  inference(superposition,[],[f575,f4991])).
fof(f8675,plain,(
  ( ! [X99,X98] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X99))) | s(bool,X99) = s(bool,sK7('const.option.SOME_1'(s(bool,X98)),bool)) | s(bool,'T_0') = s(bool,X98)) ) | $spl33),
  inference(superposition,[],[f154,f8232])).
fof(f8734,plain,(
  ( ! [X109,X110] : (s(bool,X109) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X109)),bool),X110,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X110),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X109)) ) | $spl33),
  inference(forward_demodulation,[],[f8680,f946])).
fof(f8680,plain,(
  ( ! [X109,X110] : (~p(s(bool,happ(s(fun(bool,bool),X110),s(bool,'F_0')))) | s(bool,sK7('const.option.SOME_1'(s(bool,X109)),bool)) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X109)),bool),X110,bool)) | s(bool,'T_0') = s(bool,X109)) ) | $spl33),
  inference(superposition,[],[f157,f8232])).
fof(f9476,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X66),'const.option.NONE_0'),s(bool,sK7(X65,bool)),s(fun(X66,bool),X67))) | s('type.option.option'(bool),X65) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X65)) ) | $spl33),
  inference(trivial_inequality_removal,[],[f9437])).
fof(f9437,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X66),'const.option.NONE_0'),s(bool,sK7(X65,bool)),s(fun(X66,bool),X67))) | s('type.option.option'(bool),X65) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X65)) ) | $spl33),
  inference(superposition,[],[f4437,f610])).
fof(f9810,plain,(
  ( ! [X103,X104] : (s('type.option.option'('type.option.option'(bool)),X104) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) = s('type.option.option'(bool),sK7(X104,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X104) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103)))) ) | $spl33),
  inference(superposition,[],[f611,f5196])).
fof(f9816,plain,(
  ( ! [X118,X117] : (s('type.option.option'('type.option.option'(bool)),X118) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X117))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK7(X118,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X118) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X117)))) ) | $spl33),
  inference(superposition,[],[f611,f5196])).
fof(f9822,plain,(
  ( ! [X6,X5] : (s('type.option.option'(bool),X6) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X5) = s(bool,sK7(X6,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X6) | s(bool,X5) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f611,f4881])).
fof(f9831,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | $spl33),
  inference(subsumption_resolution,[],[f9762,f151])).
fof(f9762,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | $spl33),
  inference(superposition,[],[f611,f5196])).
fof(f9843,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),X43) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X42) = s(bool,sK7(X43,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X43) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X42)),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f9787,f946])).
fof(f9787,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),X43) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(bool,X42) = s(bool,sK7(X43,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X43) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X42)),bool))) ) | $spl33),
  inference(superposition,[],[f611,f4452])).
fof(f9845,plain,(
  ( ! [X47,X48] : (s('type.option.option'(bool),X48) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X47) = s(bool,sK7(X48,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X48) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47)))) ) | $spl33),
  inference(forward_demodulation,[],[f9789,f946])).
fof(f9789,plain,(
  ( ! [X47,X48] : (s('type.option.option'(bool),X48) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(bool,X47) = s(bool,sK7(X48,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X48) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47)))) ) | $spl33),
  inference(superposition,[],[f611,f4881])).
fof(f9846,plain,(
  ( ! [X74,X73] : (s(bool,'F_0') = s(bool,sK7(X74,bool)) | s('type.option.option'(bool),X74) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X73))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X74) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X73)))) ) | $spl33),
  inference(forward_demodulation,[],[f9798,f946])).
fof(f9798,plain,(
  ( ! [X74,X73] : (s('type.option.option'(bool),X74) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X73))) | s(bool,sK7(X74,bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X74) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X73)))) ) | $spl33),
  inference(superposition,[],[f611,f4881])).
fof(f9847,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') = s(bool,sK7(X76,bool)) | s('type.option.option'(bool),X76) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X76) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X75)),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f9799,f946])).
fof(f9799,plain,(
  ( ! [X76,X75] : (s('type.option.option'(bool),X76) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75))) | s(bool,sK7(X76,bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X76) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X75)),bool))) ) | $spl33),
  inference(superposition,[],[f611,f4452])).
fof(f9851,plain,(
  ( ! [X109,X108] : (s(bool,'F_0') = s(bool,X108) | s('type.option.option'('type.option.option'(bool)),X109) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X108))) = s('type.option.option'(bool),sK7(X109,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X109)) ) | $spl33),
  inference(forward_demodulation,[],[f9812,f946])).
fof(f9812,plain,(
  ( ! [X109,X108] : (s('type.option.option'('type.option.option'(bool)),X109) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X108))) = s('type.option.option'(bool),sK7(X109,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X109) | s(bool,X108) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f611,f4881])).
fof(f9852,plain,(
  ( ! [X114,X113] : (s(bool,'F_0') = s(bool,X113) | s('type.option.option'('type.option.option'(bool)),X114) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X113))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK7(X114,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X114)) ) | $spl33),
  inference(forward_demodulation,[],[f9814,f946])).
fof(f9814,plain,(
  ( ! [X114,X113] : (s('type.option.option'('type.option.option'(bool)),X114) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X113))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK7(X114,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X114) | s(bool,X113) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f611,f4881])).
fof(f10083,plain,(
  ( ! [X175,X173,X176,X174] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X174))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X174) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X175),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X173))),s(fun(X175,'type.option.option'(bool)),X176))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X173)))) ) | $spl33),
  inference(superposition,[],[f491,f5196])).
fof(f10089,plain,(
  ( ! [X202,X200,X199,X201] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X200))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199))))) | s('type.option.option'(bool),X200) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X201),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(fun(X201,'type.option.option'(bool)),X202))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199)))) ) | $spl33),
  inference(superposition,[],[f491,f5196])).
fof(f10137,plain,(
  ( ! [X175,X173,X176,X174] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X174))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X173))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X175),'const.option.NONE_0'),s('type.option.option'(bool),X174),s(fun(X175,'type.option.option'(bool)),X176))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X173)))) ) | $spl33),
  inference(superposition,[],[f491,f5196])).
fof(f10143,plain,(
  ( ! [X202,X200,X199,X201] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X200))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X201),'const.option.NONE_0'),s('type.option.option'(bool),X200),s(fun(X201,'type.option.option'(bool)),X202))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199)))) ) | $spl33),
  inference(superposition,[],[f491,f5196])).
fof(f10168,plain,(
  ( ! [X78,X76,X77,X75] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X76))) | s(bool,X76) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.NONE_0'),s(bool,X75),s(fun(X77,bool),X78))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75)))) ) | $spl33),
  inference(forward_demodulation,[],[f10062,f946])).
fof(f10062,plain,(
  ( ! [X78,X76,X77,X75] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X76))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(bool,X76) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.NONE_0'),s(bool,X75),s(fun(X77,bool),X78))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75)))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10170,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,X120) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X121),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X121,bool),X122))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X120))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119)))) ) | $spl33),
  inference(forward_demodulation,[],[f10071,f946])).
fof(f10071,plain,(
  ( ! [X121,X122,X120,X119] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X120))) | s(bool,X120) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X121),'const.option.NONE_0'),s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)),s(fun(X121,bool),X122))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119)))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10171,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,X124) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X125),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X125,bool),X126))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X124))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X123)),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f10072,f946])).
fof(f10072,plain,(
  ( ! [X125,X123,X126,X124] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X124))) | s(bool,X124) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X125),'const.option.NONE_0'),s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)),s(fun(X125,bool),X126))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X123)),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4452])).
fof(f10184,plain,(
  ( ! [X185,X182,X184,X183] : (s(bool,'F_0') = s(bool,X182) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X183))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X183) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X184),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X182))),s(fun(X184,'type.option.option'(bool)),X185)))) ) | $spl33),
  inference(forward_demodulation,[],[f10085,f946])).
fof(f10085,plain,(
  ( ! [X185,X182,X184,X183] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X183))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X183) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X184),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X182))),s(fun(X184,'type.option.option'(bool)),X185))) | s(bool,X182) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10185,plain,(
  ( ! [X191,X194,X192,X193] : (s(bool,'F_0') = s(bool,X191) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X192))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X191))))) | s('type.option.option'(bool),X192) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X193),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(fun(X193,'type.option.option'(bool)),X194)))) ) | $spl33),
  inference(forward_demodulation,[],[f10087,f946])).
fof(f10087,plain,(
  ( ! [X191,X194,X192,X193] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X192))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X191))))) | s('type.option.option'(bool),X192) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X193),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(fun(X193,'type.option.option'(bool)),X194))) | s(bool,X191) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10188,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') = s(bool,X18) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,X19) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X20,bool),X21)))) ) | $spl33),
  inference(forward_demodulation,[],[f10097,f946])).
fof(f10097,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,X19) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X20,bool),X21))) | s(bool,X18) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10192,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X67))) | s(bool,X66) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X68),'const.option.NONE_0'),s(bool,X67),s(fun(X68,bool),X69))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X66)),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f10114,f946])).
fof(f10114,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X67))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(bool,X66) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X68),'const.option.NONE_0'),s(bool,X67),s(fun(X68,bool),X69))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X66)),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4452])).
fof(f10194,plain,(
  ( ! [X78,X76,X77,X75] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X76))) | s(bool,X75) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.NONE_0'),s(bool,X76),s(fun(X77,bool),X78))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75)))) ) | $spl33),
  inference(forward_demodulation,[],[f10116,f946])).
fof(f10116,plain,(
  ( ! [X78,X76,X77,X75] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X76))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | s(bool,X75) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.NONE_0'),s(bool,X76),s(fun(X77,bool),X78))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75)))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10195,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X121),'const.option.NONE_0'),s(bool,X120),s(fun(X121,bool),X122))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X120))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119)))) ) | $spl33),
  inference(forward_demodulation,[],[f10125,f946])).
fof(f10125,plain,(
  ( ! [X121,X122,X120,X119] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X120))) | s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X121),'const.option.NONE_0'),s(bool,X120),s(fun(X121,bool),X122))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119)))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10196,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X125),'const.option.NONE_0'),s(bool,X124),s(fun(X125,bool),X126))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X124))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X123)),bool))) ) | $spl33),
  inference(forward_demodulation,[],[f10126,f946])).
fof(f10126,plain,(
  ( ! [X125,X123,X126,X124] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X124))) | s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X125),'const.option.NONE_0'),s(bool,X124),s(fun(X125,bool),X126))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X123)),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4452])).
fof(f10201,plain,(
  ( ! [X185,X182,X184,X183] : (s(bool,'F_0') = s(bool,X182) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X183))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X182))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X184),'const.option.NONE_0'),s('type.option.option'(bool),X183),s(fun(X184,'type.option.option'(bool)),X185)))) ) | $spl33),
  inference(forward_demodulation,[],[f10139,f946])).
fof(f10139,plain,(
  ( ! [X185,X182,X184,X183] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X183))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X182))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X184),'const.option.NONE_0'),s('type.option.option'(bool),X183),s(fun(X184,'type.option.option'(bool)),X185))) | s(bool,X182) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10202,plain,(
  ( ! [X191,X194,X192,X193] : (s(bool,'F_0') = s(bool,X191) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X192))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X191))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X193),'const.option.NONE_0'),s('type.option.option'(bool),X192),s(fun(X193,'type.option.option'(bool)),X194)))) ) | $spl33),
  inference(forward_demodulation,[],[f10141,f946])).
fof(f10141,plain,(
  ( ! [X191,X194,X192,X193] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X192))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X191))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X193),'const.option.NONE_0'),s('type.option.option'(bool),X192),s(fun(X193,'type.option.option'(bool)),X194))) | s(bool,X191) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10204,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') = s(bool,X9) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10))) | s(bool,X9) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X11),'const.option.NONE_0'),s(bool,X10),s(fun(X11,bool),X12)))) ) | $spl33),
  inference(forward_demodulation,[],[f10149,f946])).
fof(f10149,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10))) | s(bool,X9) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X11),'const.option.NONE_0'),s(bool,X10),s(fun(X11,bool),X12))) | s(bool,X9) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f491,f4881])).
fof(f10596,plain,(
  ( ! [X175,X173,X176,X174] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X174),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X173))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X175),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X173))),s(fun(X175,'type.option.option'(bool)),X176)),X174,'type.option.option'(bool))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X173)))) ) | $spl33),
  inference(superposition,[],[f585,f5196])).
fof(f10602,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X200),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X201),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(fun(X201,'type.option.option'(bool)),X202)),X200,'type.option.option'(bool))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199)))) ) | $spl33),
  inference(superposition,[],[f585,f5196])).
fof(f10634,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,happ(s(fun(bool,bool),X76),s(bool,'F_0')))) | s(bool,X75) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.NONE_0'),s(bool,X75),s(fun(X77,bool),X78)),X76,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75)))) ) | $spl33),
  inference(forward_demodulation,[],[f10575,f946])).
fof(f10575,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,happ(s(fun(bool,bool),X76),s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))))) | s(bool,X75) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.NONE_0'),s(bool,X75),s(fun(X77,bool),X78)),X76,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75)))) ) | $spl33),
  inference(superposition,[],[f585,f4881])).
fof(f10636,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X121),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X121,bool),X122)),X120,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X120),s(bool,X119)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119)))) ) | $spl33),
  inference(forward_demodulation,[],[f10584,f946])).
fof(f10584,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,happ(s(fun(bool,bool),X120),s(bool,X119)))) | s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X121),'const.option.NONE_0'),s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)),s(fun(X121,bool),X122)),X120,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119)))) ) | $spl33),
  inference(superposition,[],[f585,f4881])).
fof(f10655,plain,(
  ( ! [X185,X182,X184,X183] : (s(bool,'F_0') = s(bool,X182) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X183),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X182))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X184),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X182))),s(fun(X184,'type.option.option'(bool)),X185)),X183,'type.option.option'(bool)))) ) | $spl33),
  inference(forward_demodulation,[],[f10598,f946])).
fof(f10598,plain,(
  ( ! [X185,X182,X184,X183] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X183),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X182))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X184),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X182))),s(fun(X184,'type.option.option'(bool)),X185)),X183,'type.option.option'(bool))) | s(bool,X182) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f585,f4881])).
fof(f10656,plain,(
  ( ! [X191,X194,X192,X193] : (s(bool,'F_0') = s(bool,X191) | ~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X192),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X191)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X193),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(fun(X193,'type.option.option'(bool)),X194)),X192,'type.option.option'(bool)))) ) | $spl33),
  inference(forward_demodulation,[],[f10600,f946])).
fof(f10600,plain,(
  ( ! [X191,X194,X192,X193] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X192),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X191)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X193),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(fun(X193,'type.option.option'(bool)),X194)),X192,'type.option.option'(bool))) | s(bool,X191) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f585,f4881])).
fof(f10904,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X29,X30,bool)) | s(bool,X29) = s(bool,X31) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31)))) ) | $spl33),
  inference(superposition,[],[f675,f4991])).
fof(f10905,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X33)))) | p(s(bool,happ(s(fun(bool,bool),X34),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X32,X34,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X33)) ) | $spl33),
  inference(superposition,[],[f675,f4991])).
fof(f10906,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X36)))) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X35,X37,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,'T_0') = s(bool,X35)) ) | $spl33),
  inference(superposition,[],[f675,f4991])).
fof(f10912,plain,(
  ( ! [X54,X53] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,X53)))) | p(s(bool,happ(s(fun(bool,bool),X54),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,'F_0')),bool),X54,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X53)))) ) | $spl33),
  inference(superposition,[],[f675,f4881])).
fof(f10955,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') = s(bool,X55) | ~p(s(bool,happ(s(fun(bool,bool),X56),s(bool,X55)))) | p(s(bool,happ(s(fun(bool,bool),X56),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,'F_0')),bool),X56,bool))) ) | $spl33),
  inference(forward_demodulation,[],[f10913,f946])).
fof(f10913,plain,(
  ( ! [X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X56),s(bool,X55)))) | p(s(bool,happ(s(fun(bool,bool),X56),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,'F_0')),bool),X56,bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X55)),bool))) ) | $spl33),
  inference(superposition,[],[f675,f4452])).
fof(f10958,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') = s(bool,X65) | ~p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X65)),bool),X66,bool))) ) | $spl33),
  inference(forward_demodulation,[],[f10917,f946])).
fof(f10917,plain,(
  ( ! [X66,X65] : (~p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X65)),bool),X66,bool)) | s(bool,X65) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl33),
  inference(superposition,[],[f675,f4452])).
fof(f10959,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') = s(bool,X68) | ~p(s(bool,happ(s(fun(bool,bool),X69),s(bool,X68)))) | p(s(bool,happ(s(fun(bool,bool),X69),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X67)),bool),X69,bool)) | s(bool,'T_0') = s(bool,X67)) ) | $spl33),
  inference(forward_demodulation,[],[f10918,f946])).
fof(f10918,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,happ(s(fun(bool,bool),X69),s(bool,X68)))) | p(s(bool,happ(s(fun(bool,bool),X69),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X67)),bool),X69,bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X68)),bool)) | s(bool,'T_0') = s(bool,X67)) ) | $spl33),
  inference(superposition,[],[f675,f4652])).
fof(f10971,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X30,X29,bool)) | s(bool,X31) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31)))) ) | $spl33),
  inference(subsumption_resolution,[],[f10939,f111])).
fof(f10939,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(X30,X29,bool)) | s(bool,X31) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31)))) ) | $spl33),
  inference(superposition,[],[f675,f4991])).
fof(f5648,plain,(
  ( ! [X177,X180,X178,X179] : (s(bool,X180) != s(bool,f(s(X177,sK9(X178,X179,bool,X177)))) | s(fun(X177,bool),X178) = s(fun(X177,bool),X179) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X180))) | s(bool,'T_0') = s(bool,g(s(X177,sK9(X178,X179,bool,X177))))) ) | $spl33),
  inference(superposition,[],[f164,f4991])).
fof(f5719,plain,(
  ( ! [X177,X180,X178,X179] : (s(bool,X180) != s(bool,f(s(X177,sK9(X178,X179,bool,X177)))) | s(fun(X177,bool),X178) = s(fun(X177,bool),X179) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(X177,sK9(X178,X179,bool,X177)))))) | s(bool,'T_0') = s(bool,X180)) ) | $spl33),
  inference(superposition,[],[f164,f4991])).
fof(f5867,plain,(
  ( ! [X167,X165,X166,X168] : (s(bool,'T_0') != s(bool,f(s(X165,sK9(X166,X167,bool,X165)))) | s(fun(X165,bool),X166) = s(fun(X165,bool),X167) | s(bool,X168) = s(bool,g(s(X165,sK9(X166,X167,bool,X165)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X168)))) ) | $spl33),
  inference(superposition,[],[f164,f4991])).
fof(f4439,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X19),X20),s(bool,X21),s(fun(X19,bool),X22))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X19,bool),bool),happ(s(fun(bool,fun(fun(X19,bool),bool)),happ(s(fun('type.option.option'(X19),fun(bool,fun(fun(X19,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X19),X20))),s(bool,X21))),s(fun(X19,bool),X22)))) ) | $spl33),
  inference(forward_demodulation,[],[f4421,f946])).
fof(f4421,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X19),X20),s(bool,X21),s(fun(X19,bool),X22)))),bool)) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X19,bool),bool),happ(s(fun(bool,fun(fun(X19,bool),bool)),happ(s(fun('type.option.option'(X19),fun(bool,fun(fun(X19,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X19),X20))),s(bool,X21))),s(fun(X19,bool),X22)))) ) | $spl33),
  inference(superposition,[],[f4267,f166])).
fof(f5502,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X62) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X58),X59),s(bool,X60),s(fun(X58,bool),X61))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X58,bool),bool),happ(s(fun(bool,fun(fun(X58,bool),bool)),happ(s(fun('type.option.option'(X58),fun(bool,fun(fun(X58,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X58),X59))),s(bool,X60))),s(fun(X58,bool),X61)))) ) | $spl33),
  inference(superposition,[],[f4991,f166])).
fof(f5556,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),X48),s(bool,X49),s(fun(X47,bool),X50))))) | s(bool,X51) = s(bool,happ(s(fun(fun(X47,bool),bool),happ(s(fun(bool,fun(fun(X47,bool),bool)),happ(s(fun('type.option.option'(X47),fun(bool,fun(fun(X47,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X47),X48))),s(bool,X49))),s(fun(X47,bool),X50))) | s(bool,'T_0') = s(bool,X51)) ) | $spl33),
  inference(superposition,[],[f4991,f166])).
fof(f5583,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),X48),s(bool,X49),s(fun(X47,bool),X50))) | s(bool,X51) = s(bool,happ(s(fun(fun(X47,bool),bool),happ(s(fun(bool,fun(fun(X47,bool),bool)),happ(s(fun('type.option.option'(X47),fun(bool,fun(fun(X47,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X47),X48))),s(bool,X49))),s(fun(X47,bool),X50))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51)))) ) | $spl33),
  inference(superposition,[],[f4991,f166])).
fof(f8283,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X83) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X79),X80),s(bool,X81),s(fun(X79,bool),X82))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X83)),bool)) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X79,bool),bool),happ(s(fun(bool,fun(fun(X79,bool),bool)),happ(s(fun('type.option.option'(X79),fun(bool,fun(fun(X79,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X79),X80))),s(bool,X81))),s(fun(X79,bool),X82)))) ) | $spl33),
  inference(forward_demodulation,[],[f7871,f946])).
fof(f7871,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X83) = s(bool,sK7('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X79),X80),s(bool,X81),s(fun(X79,bool),X82)))),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X83)),bool)) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X79,bool),bool),happ(s(fun(bool,fun(fun(X79,bool),bool)),happ(s(fun('type.option.option'(X79),fun(bool,fun(fun(X79,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X79),X80))),s(bool,X81))),s(fun(X79,bool),X82)))) ) | $spl33),
  inference(superposition,[],[f4652,f166])).
fof(f4438,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X18),s(X15,X16))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.SOME_1'(s(X15,X16))),s(bool,X17),s(fun(X15,bool),X18)))) ) | $spl33),
  inference(forward_demodulation,[],[f4420,f946])).
fof(f4420,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,happ(s(fun(X15,bool),X18),s(X15,X16)))),bool)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.SOME_1'(s(X15,X16))),s(bool,X17),s(fun(X15,bool),X18)))) ) | $spl33),
  inference(superposition,[],[f4267,f165])).
fof(f5168,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,happ(s(fun(X34,bool),X35),s(X34,X36)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X34,bool),X35),s(X34,sK8(X36,X35,X34))))))) ) | $spl33),
  inference(subsumption_resolution,[],[f5167,f113])).
fof(f5167,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X34,bool),X35),s(X34,X36)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X34,bool),X35),s(X34,sK8(X36,X35,X34))))))) ) | $spl33),
  inference(forward_demodulation,[],[f5058,f946])).
fof(f5058,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) | ~p(s(bool,happ(s(fun(X34,bool),X35),s(X34,X36)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X34,bool),X35),s(X34,sK8(X36,X35,X34))))))) ) | $spl33),
  inference(superposition,[],[f158,f4881])).
fof(f5501,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X57) = s(bool,happ(s(fun(X53,bool),X56),s(X53,X54))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X53),'const.option.SOME_1'(s(X53,X54))),s(bool,X55),s(fun(X53,bool),X56)))) ) | $spl33),
  inference(superposition,[],[f4991,f165])).
fof(f5555,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X42,bool),X45),s(X42,X43))))) | s(bool,X46) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X42),'const.option.SOME_1'(s(X42,X43))),s(bool,X44),s(fun(X42,bool),X45))) | s(bool,'T_0') = s(bool,X46)) ) | $spl33),
  inference(superposition,[],[f4991,f165])).
fof(f5582,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),X45),s(X42,X43))) | s(bool,X46) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X42),'const.option.SOME_1'(s(X42,X43))),s(bool,X44),s(fun(X42,bool),X45))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46)))) ) | $spl33),
  inference(superposition,[],[f4991,f165])).
fof(f6030,plain,(
  ( ! [X94,X92,X95,X93,X91] : (p(s(bool,happ(s(fun(X91,bool),X92),s(X91,X94)))) | s(X91,X93) != s(X91,X94) | s(bool,X95) = s(bool,happ(s(fun(X91,bool),X92),s(X91,X93))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X95)))) ) | $spl33),
  inference(subsumption_resolution,[],[f5844,f111])).
fof(f5844,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X91,bool),X92),s(X91,X94)))) | s(X91,X93) != s(X91,X94) | s(bool,X95) = s(bool,happ(s(fun(X91,bool),X92),s(X91,X93))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X95)))) ) | $spl33),
  inference(superposition,[],[f156,f4991])).
fof(f8282,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,X78) = s(bool,happ(s(fun(X74,bool),X77),s(X74,X75))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X78)),bool)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X74),'const.option.SOME_1'(s(X74,X75))),s(bool,X76),s(fun(X74,bool),X77)))) ) | $spl33),
  inference(forward_demodulation,[],[f7870,f946])).
fof(f7870,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,X78) = s(bool,sK7('const.option.SOME_1'(s(bool,happ(s(fun(X74,bool),X77),s(X74,X75)))),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X78)),bool)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X74),'const.option.SOME_1'(s(X74,X75))),s(bool,X76),s(fun(X74,bool),X77)))) ) | $spl33),
  inference(superposition,[],[f4652,f165])).
fof(f10666,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (s(X62,X64) = s(X62,sK8('const.option.option_CASE_3'(s('type.option.option'(X65),'const.option.NONE_0'),s(X62,X64),s(fun(X65,X62),X66)),X63,X62)) | s(bool,X67) = s(bool,happ(s(fun(X62,bool),X63),s(X62,X64))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X67)))) ) | $spl33),
  inference(subsumption_resolution,[],[f10617,f111])).
fof(f10617,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (~p(s(bool,'T_0')) | s(X62,X64) = s(X62,sK8('const.option.option_CASE_3'(s('type.option.option'(X65),'const.option.NONE_0'),s(X62,X64),s(fun(X65,X62),X66)),X63,X62)) | s(bool,X67) = s(bool,happ(s(fun(X62,bool),X63),s(X62,X64))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X67)))) ) | $spl33),
  inference(superposition,[],[f585,f4991])).
fof(f5640,plain,(
  ( ! [X146,X144,X145,X143] : (~p(s(bool,happ(s(fun(bool,bool),X145),s(bool,X144)))) | p(s(bool,happ(s(fun(bool,bool),X145),s(bool,X146)))) | s(bool,X143) != s(bool,X146) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X144))) | s(bool,'T_0') = s(bool,X143)) ) | $spl33),
  inference(superposition,[],[f156,f4991])).
fof(f5641,plain,(
  ( ! [X149,X147,X148] : (~p(s(bool,happ(s(fun(bool,bool),X149),s(bool,X148)))) | s(bool,X147) = s(bool,sK8(X147,X149,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X148))) | s(bool,'T_0') = s(bool,X147)) ) | $spl33),
  inference(superposition,[],[f157,f4991])).
fof(f5654,plain,(
  ( ! [X200,X199,X201] : (p(s(bool,happ(s(fun(bool,bool),X200),s(bool,X201)))) | ~p(s(bool,happ(s(fun(bool,bool),X200),s(bool,X199)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X201))) | s(bool,'T_0') = s(bool,sK8(X199,X200,bool))) ) | $spl33),
  inference(superposition,[],[f158,f4991])).
fof(f5711,plain,(
  ( ! [X146,X144,X145,X143] : (~p(s(bool,happ(s(fun(bool,bool),X145),s(bool,X144)))) | p(s(bool,happ(s(fun(bool,bool),X145),s(bool,X146)))) | s(bool,X143) != s(bool,X146) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X143))) | s(bool,'T_0') = s(bool,X144)) ) | $spl33),
  inference(superposition,[],[f156,f4991])).
fof(f5712,plain,(
  ( ! [X149,X147,X148] : (~p(s(bool,happ(s(fun(bool,bool),X149),s(bool,X148)))) | s(bool,X147) = s(bool,sK8(X147,X149,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X147))) | s(bool,'T_0') = s(bool,X148)) ) | $spl33),
  inference(superposition,[],[f157,f4991])).
fof(f5725,plain,(
  ( ! [X200,X199,X201] : (p(s(bool,happ(s(fun(bool,bool),X200),s(bool,X201)))) | ~p(s(bool,happ(s(fun(bool,bool),X200),s(bool,X199)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK8(X199,X200,bool)))) | s(bool,'T_0') = s(bool,X201)) ) | $spl33),
  inference(superposition,[],[f158,f4991])).
fof(f5801,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,X90)))) | p(s(bool,happ(s(fun(bool,bool),X91),s(bool,X92)))) | s(bool,'T_0') != s(bool,X92) | s(bool,X90) = s(bool,X93) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X93)))) ) | $spl33),
  inference(superposition,[],[f156,f4991])).
fof(f5802,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,happ(s(fun(bool,bool),X95),s(bool,X94)))) | s(bool,'T_0') = s(bool,sK8('T_0',X95,bool)) | s(bool,X94) = s(bool,X96) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X96)))) ) | $spl33),
  inference(superposition,[],[f157,f4991])).
fof(f5859,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,happ(s(fun(bool,bool),X133),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X133),s(bool,X134)))) | s(bool,X132) != s(bool,X134) | s(bool,X132) = s(bool,X135) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X135)))) ) | $spl33),
  inference(superposition,[],[f156,f4991])).
fof(f5873,plain,(
  ( ! [X189,X187,X188] : (p(s(bool,happ(s(fun(bool,bool),X188),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X188),s(bool,X187)))) | s(bool,X189) = s(bool,sK8(X187,X188,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X189)))) ) | $spl33),
  inference(superposition,[],[f158,f4991])).
fof(f8473,plain,(
  ( ! [X165,X163,X164] : (s(bool,'T_0') = s(bool,X164) | s(bool,X163) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X163)),bool),X165,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X165),s(bool,X164)))) | s(bool,'T_0') = s(bool,X163)) ) | $spl33),
  inference(forward_demodulation,[],[f8472,f946])).
fof(f8472,plain,(
  ( ! [X165,X163,X164] : (s(bool,X163) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X163)),bool),X165,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X165),s(bool,X164)))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X164)),bool)) | s(bool,'T_0') = s(bool,X163)) ) | $spl33),
  inference(forward_demodulation,[],[f8090,f946])).
fof(f8090,plain,(
  ( ! [X165,X163,X164] : (~p(s(bool,happ(s(fun(bool,bool),X165),s(bool,X164)))) | s(bool,sK7('const.option.SOME_1'(s(bool,X163)),bool)) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X163)),bool),X165,bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X164)),bool)) | s(bool,'T_0') = s(bool,X163)) ) | $spl33),
  inference(superposition,[],[f157,f4652])).
fof(f10577,plain,(
  ( ! [X88,X87,X85,X86,X84] : (~p(s(bool,happ(s(fun(bool,bool),X86),s(bool,X85)))) | s(bool,X84) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X87),'const.option.NONE_0'),s(bool,X84),s(fun(X87,bool),X88)),X86,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X84))) | s(bool,'T_0') = s(bool,X85)) ) | $spl33),
  inference(superposition,[],[f585,f4991])).
fof(f10578,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,X90)))) | s(bool,X89) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X92),'const.option.NONE_0'),s(bool,X89),s(fun(X92,bool),X93)),X91,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X90))) | s(bool,'T_0') = s(bool,X89)) ) | $spl33),
  inference(superposition,[],[f585,f4991])).
fof(f10579,plain,(
  ( ! [X94,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(bool,bool),X95),s(bool,X94)))) | s(bool,'T_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X96),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X96,bool),X97)),X95,bool)) | s(bool,X94) = s(bool,X98) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X98)))) ) | $spl33),
  inference(superposition,[],[f585,f4991])).
fof(f5860,plain,(
  ( ! [X138,X136,X137] : (~p(s(bool,happ(s(fun(bool,bool),X137),s(bool,'T_0')))) | s(bool,X136) = s(bool,sK8(X136,X137,bool)) | s(bool,X136) = s(bool,X138) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X138)))) ) | $spl33),
  inference(superposition,[],[f157,f4991])).
fof(f10576,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,happ(s(fun(bool,bool),X80),s(bool,'T_0')))) | s(bool,X79) = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X81),'const.option.NONE_0'),s(bool,X79),s(fun(X81,bool),X82)),X80,bool)) | s(bool,X79) = s(bool,X83) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X83)))) ) | $spl33),
  inference(superposition,[],[f585,f4991])).
fof(f5741,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X44),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43)))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X44),s('type.option.option'(bool),X45)))) | s('type.option.option'(bool),X45) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl33),
  inference(superposition,[],[f156,f4991])).
fof(f5742,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X48),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,'T_0')),X48,'type.option.option'(bool))) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) ) | $spl33),
  inference(superposition,[],[f157,f4991])).
fof(f10599,plain,(
  ( ! [X189,X187,X190,X188,X186] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X187),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X186)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X188),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(fun(X188,'type.option.option'(bool)),X189)),X187,'type.option.option'(bool))) | s(bool,X186) = s(bool,X190) | s(bool,'T_0') = s(bool,X190)) ) | $spl33),
  inference(superposition,[],[f585,f4991])).
fof(f10597,plain,(
  ( ! [X177,X180,X178,X181,X179] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X178),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X177))) = s('type.option.option'(bool),sK8('const.option.option_CASE_3'(s('type.option.option'(X179),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X177))),s(fun(X179,'type.option.option'(bool)),X180)),X178,'type.option.option'(bool))) | s(bool,X177) = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) ) | $spl33),
  inference(superposition,[],[f585,f4991])).
fof(f5765,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X50),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49))) = s('type.option.option'(bool),sK8('const.option.SOME_1'(s(bool,X49)),X50,'type.option.option'(bool))) | s(bool,X49) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) ) | $spl33),
  inference(superposition,[],[f157,f4991])).
fof(f5764,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X46),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X46),s('type.option.option'(bool),X47)))) | s('type.option.option'(bool),X47) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X45))) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) ) | $spl33),
  inference(superposition,[],[f156,f4991])).
fof(f5723,plain,(
  ( ! [X196,X195] : (s('type.option.option'(bool),X195) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X196))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X195) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7(X195,bool)))) | s(bool,'T_0') = s(bool,X196)) ) | $spl33),
  inference(superposition,[],[f146,f4991])).
fof(f8415,plain,(
  ( ! [X257,X258] : (s('type.option.option'(bool),X257) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X258))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X257) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X257,bool))),bool)) | s(bool,'T_0') = s(bool,X258)) ) | $spl33),
  inference(forward_demodulation,[],[f8022,f946])).
fof(f8022,plain,(
  ( ! [X257,X258] : (s('type.option.option'(bool),X257) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X258)),bool)))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X257) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X257,bool))),bool)) | s(bool,'T_0') = s(bool,X258)) ) | $spl33),
  inference(superposition,[],[f146,f4652])).
fof(f5726,plain,(
  ( ! [X204,X202,X205,X203] : (s(X204,f(s(bool,X205))) != s(X204,g(s(bool,X205))) | s(fun(bool,X204),X202) = s(fun(bool,X204),X203) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK9(X202,X203,X204,bool)))) | s(bool,'T_0') = s(bool,X205)) ) | $spl33),
  inference(superposition,[],[f164,f4991])).
fof(f8419,plain,(
  ( ! [X265,X267,X264,X266] : (s(X266,f(s(bool,X267))) != s(X266,g(s(bool,X267))) | s(fun(bool,X266),X264) = s(fun(bool,X266),X265) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK9(X264,X265,X266,bool))),bool)) | s(bool,'T_0') = s(bool,X267)) ) | $spl33),
  inference(forward_demodulation,[],[f8025,f946])).
fof(f8025,plain,(
  ( ! [X265,X267,X264,X266] : (s(X266,f(s(bool,sK7('const.option.SOME_1'(s(bool,X267)),bool)))) != s(X266,g(s(bool,sK7('const.option.SOME_1'(s(bool,X267)),bool)))) | s(fun(bool,X266),X264) = s(fun(bool,X266),X265) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK9(X264,X265,X266,bool))),bool)) | s(bool,'T_0') = s(bool,X267)) ) | $spl33),
  inference(superposition,[],[f164,f4652])).
fof(f6031,plain,(
  ( ! [X99,X97,X98,X96] : (s(X96,X98) = s(X96,sK8(X98,X97,X96)) | s(bool,X99) = s(bool,happ(s(fun(X96,bool),X97),s(X96,X98))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X99)))) ) | $spl33),
  inference(subsumption_resolution,[],[f5845,f111])).
fof(f5845,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(X96,X98) = s(X96,sK8(X98,X97,X96)) | s(bool,X99) = s(bool,happ(s(fun(X96,bool),X97),s(X96,X98))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X99)))) ) | $spl33),
  inference(superposition,[],[f157,f4991])).
fof(f10140,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X187))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X186))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X188),'const.option.NONE_0'),s('type.option.option'(bool),X187),s(fun(X188,'type.option.option'(bool)),X189))) | s(bool,X186) = s(bool,X190) | s(bool,'T_0') = s(bool,X190)) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10138,plain,(
  ( ! [X177,X180,X178,X181,X179] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X178))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X177))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X179),'const.option.NONE_0'),s('type.option.option'(bool),X178),s(fun(X179,'type.option.option'(bool)),X180))) | s(bool,X177) = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10120,plain,(
  ( ! [X94,X97,X95,X98,X96] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X94))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X95))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X96),'const.option.NONE_0'),s(bool,X95),s(fun(X96,bool),X97))) | s(bool,X94) = s(bool,X98) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X98)))) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10119,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X90))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X91))) | s(bool,X89) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X92),'const.option.NONE_0'),s(bool,X91),s(fun(X92,bool),X93))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X90))) | s(bool,'T_0') = s(bool,X89)) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10118,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X85))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X86))) | s(bool,X84) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X87),'const.option.NONE_0'),s(bool,X86),s(fun(X87,bool),X88))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X84))) | s(bool,'T_0') = s(bool,X85)) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10117,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))) | s(bool,X79) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X81),'const.option.NONE_0'),s(bool,X80),s(fun(X81,bool),X82))) | s(bool,X79) = s(bool,X83) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X83)))) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10193,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X71))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X72))) | s(bool,X70) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,X72),s(fun(X73,bool),X74))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X70)),bool)) | s(bool,'T_0') = s(bool,X71)) ) | $spl33),
  inference(forward_demodulation,[],[f10115,f946])).
fof(f10115,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X72))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X71)),bool)))) | s(bool,X70) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,X72),s(fun(X73,bool),X74))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X70)),bool)) | s(bool,'T_0') = s(bool,X71)) ) | $spl33),
  inference(superposition,[],[f491,f4652])).
fof(f10086,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X187))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X186))))) | s('type.option.option'(bool),X187) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X188),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(fun(X188,'type.option.option'(bool)),X189))) | s(bool,X186) = s(bool,X190) | s(bool,'T_0') = s(bool,X190)) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10084,plain,(
  ( ! [X177,X180,X178,X181,X179] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X178))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X178) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X179),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X177))),s(fun(X179,'type.option.option'(bool)),X180))) | s(bool,X177) = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10178,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(bool,X151) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X152),'const.option.NONE_0'),s(bool,X150),s(fun(X152,bool),X153))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X151))) | s(bool,X150) = s(bool,sK7('const.option.SOME_1'(s(bool,X154)),bool)) | s(bool,'T_0') = s(bool,X154)) ) | $spl33),
  inference(forward_demodulation,[],[f10078,f946])).
fof(f10078,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X151))) | s(bool,X151) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X152),'const.option.NONE_0'),s(bool,sK7('const.option.SOME_1'(s(bool,X150)),bool)),s(fun(X152,bool),X153))) | s(bool,X150) = s(bool,sK7('const.option.SOME_1'(s(bool,X154)),bool)) | s(bool,'T_0') = s(bool,X154)) ) | $spl33),
  inference(superposition,[],[f491,f4652])).
fof(f10066,plain,(
  ( ! [X94,X97,X95,X98,X96] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X94))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X95))) | s(bool,X95) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X96),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X96,bool),X97))) | s(bool,X94) = s(bool,X98) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X98)))) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10065,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X90))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X91))) | s(bool,X91) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X92),'const.option.NONE_0'),s(bool,X89),s(fun(X92,bool),X93))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X90))) | s(bool,'T_0') = s(bool,X89)) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10064,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X85))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X86))) | s(bool,X86) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X87),'const.option.NONE_0'),s(bool,X84),s(fun(X87,bool),X88))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X84))) | s(bool,'T_0') = s(bool,X85)) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f10063,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X80))) | s(bool,X80) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X81),'const.option.NONE_0'),s(bool,X79),s(fun(X81,bool),X82))) | s(bool,X79) = s(bool,X83) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X83)))) ) | $spl33),
  inference(superposition,[],[f491,f4991])).
fof(f9813,plain,(
  ( ! [X111,X112,X110] : (s('type.option.option'('type.option.option'(bool)),X111) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X110))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),sK7(X111,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X111) | s(bool,X110) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) ) | $spl33),
  inference(superposition,[],[f611,f4991])).
fof(f9811,plain,(
  ( ! [X107,X105,X106] : (s('type.option.option'('type.option.option'(bool)),X106) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X105))) = s('type.option.option'(bool),sK7(X106,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X106) | s(bool,X105) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) ) | $spl33),
  inference(superposition,[],[f611,f4991])).
fof(f9793,plain,(
  ( ! [X59,X60,X58] : (s('type.option.option'(bool),X59) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s(bool,'T_0') = s(bool,sK7(X59,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X59) | s(bool,X58) = s(bool,X60) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X60)))) ) | $spl33),
  inference(superposition,[],[f611,f4991])).
fof(f9792,plain,(
  ( ! [X57,X56,X55] : (s('type.option.option'(bool),X57) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56))) | s(bool,X55) = s(bool,sK7(X57,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X57) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56))) | s(bool,'T_0') = s(bool,X55)) ) | $spl33),
  inference(superposition,[],[f611,f4991])).
fof(f9791,plain,(
  ( ! [X54,X52,X53] : (s('type.option.option'(bool),X54) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X53))) | s(bool,X52) = s(bool,sK7(X54,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X54) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) | s(bool,'T_0') = s(bool,X53)) ) | $spl33),
  inference(superposition,[],[f611,f4991])).
fof(f9790,plain,(
  ( ! [X50,X51,X49] : (s('type.option.option'(bool),X50) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X49) = s(bool,sK7(X50,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X50) | s(bool,X49) = s(bool,X51) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51)))) ) | $spl33),
  inference(superposition,[],[f611,f4991])).
fof(f9844,plain,(
  ( ! [X45,X46,X44] : (s('type.option.option'(bool),X46) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X45))) | s(bool,X44) = s(bool,sK7(X46,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X46) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X44)),bool)) | s(bool,'T_0') = s(bool,X45)) ) | $spl33),
  inference(forward_demodulation,[],[f9788,f946])).
fof(f9788,plain,(
  ( ! [X45,X46,X44] : (s('type.option.option'(bool),X46) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X45)),bool)))) | s(bool,X44) = s(bool,sK7(X46,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X46) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X44)),bool)) | s(bool,'T_0') = s(bool,X45)) ) | $spl33),
  inference(superposition,[],[f611,f4652])).
fof(f9833,plain,(
  ( ! [X8,X7,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,X8) = s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl33),
  inference(subsumption_resolution,[],[f9765,f151])).
fof(f9765,plain,(
  ( ! [X8,X7,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,X8) = s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl33),
  inference(superposition,[],[f611,f4991])).
fof(f9832,plain,(
  ( ! [X4,X2,X3] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3))) | s(bool,X3) = s(bool,sK7('const.option.SOME_1'(s(bool,X2)),bool)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl33),
  inference(subsumption_resolution,[],[f9763,f151])).
fof(f9763,plain,(
  ( ! [X4,X2,X3] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3))) | s(bool,X3) = s(bool,sK7('const.option.SOME_1'(s(bool,X2)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl33),
  inference(superposition,[],[f611,f4991])).
fof(f8508,plain,(
  ( ! [X140,X141,X139] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X139))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X140))) | s(bool,'T_0') != s(bool,X140) | s(bool,X139) = s(bool,sK7('const.option.SOME_1'(s(bool,X141)),bool)) | s(bool,'T_0') = s(bool,X141)) ) | $spl33),
  inference(forward_demodulation,[],[f8150,f946])).
fof(f8150,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X140) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X140))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X139)),bool)))) | s(bool,X139) = s(bool,sK7('const.option.SOME_1'(s(bool,X141)),bool)) | s(bool,'T_0') = s(bool,X141)) ) | $spl33),
  inference(superposition,[],[f155,f4652])).
fof(f8148,plain,(
  ( ! [X134,X135,X133] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X134))) | s(bool,X134) = s(bool,sK7('const.option.SOME_1'(s(bool,X133)),bool)) | s(bool,X133) = s(bool,sK7('const.option.SOME_1'(s(bool,X135)),bool)) | s(bool,'T_0') = s(bool,X135)) ) | $spl33),
  inference(superposition,[],[f154,f4652])).
fof(f8491,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X92),'const.option.NONE_0'),s(bool,X91),s(fun(X92,bool),X93))) | s(bool,X91) = s(bool,sK7('const.option.SOME_1'(s(bool,X94)),bool)) | s(bool,'T_0') = s(bool,X94)) ) | $spl33),
  inference(forward_demodulation,[],[f8166,f946])).
fof(f8166,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X92),'const.option.NONE_0'),s(bool,sK7('const.option.SOME_1'(s(bool,X91)),bool)),s(fun(X92,bool),X93))) | s(bool,X91) = s(bool,sK7('const.option.SOME_1'(s(bool,X94)),bool)) | s(bool,'T_0') = s(bool,X94)) ) | $spl33),
  inference(trivial_inequality_removal,[],[f8131])).
fof(f8131,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X92),'const.option.NONE_0'),s(bool,sK7('const.option.SOME_1'(s(bool,X91)),bool)),s(fun(X92,bool),X93))) | s(bool,X91) = s(bool,sK7('const.option.SOME_1'(s(bool,X94)),bool)) | s(bool,'T_0') = s(bool,X94)) ) | $spl33),
  inference(superposition,[],[f4437,f4652])).
fof(f8085,plain,(
  ( ! [X149,X147,X148] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X148))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X149))) | s(bool,X149) = s(bool,sK7('const.option.SOME_1'(s(bool,X147)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X148)),bool)) | s(bool,'T_0') = s(bool,X147)) ) | $spl33),
  inference(superposition,[],[f154,f4652])).
fof(f8008,plain,(
  ( ! [X200,X199,X201] : (s(bool,X201) != s(bool,sK7('const.option.SOME_1'(s(bool,X200)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X199))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X201))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X199)),bool)) | s(bool,'T_0') = s(bool,X200)) ) | $spl33),
  inference(superposition,[],[f155,f4652])).
fof(f8406,plain,(
  ( ! [X194,X195,X193] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X194))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X195))) | s(bool,X193) = s(bool,X195) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X193)),bool)) | s(bool,'T_0') = s(bool,X194)) ) | $spl33),
  inference(forward_demodulation,[],[f8006,f946])).
fof(f8006,plain,(
  ( ! [X194,X195,X193] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X195))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X194)),bool)))) | s(bool,X193) = s(bool,X195) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X193)),bool)) | s(bool,'T_0') = s(bool,X194)) ) | $spl33),
  inference(superposition,[],[f154,f4652])).
fof(f6126,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(superposition,[],[f5196,f114])).
fof(f5874,plain,(
  ( ! [X191,X192,X190,X193] : (s(X192,f(s(bool,'T_0'))) != s(X192,g(s(bool,'T_0'))) | s(fun(bool,X192),X190) = s(fun(bool,X192),X191) | s(bool,X193) = s(bool,sK9(X190,X191,X192,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X193)))) ) | $spl33),
  inference(superposition,[],[f164,f4991])).
fof(f5871,plain,(
  ( ! [X184,X183] : (s('type.option.option'(bool),X183) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X183) | s(bool,X184) = s(bool,sK7(X183,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X184)))) ) | $spl33),
  inference(superposition,[],[f146,f4991])).
fof(f5857,plain,(
  ( ! [X127,X128,X126] : (s(bool,'T_0') != s(bool,X127) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X126))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X127))) | s(bool,X126) = s(bool,X128) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X128)))) ) | $spl33),
  inference(superposition,[],[f155,f4991])).
fof(f5855,plain,(
  ( ! [X121,X122,X120] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X121))) | s(bool,X120) = s(bool,X121) | s(bool,X120) = s(bool,X122) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X122)))) ) | $spl33),
  inference(superposition,[],[f154,f4991])).
fof(f6026,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10)))) ) | $spl33),
  inference(subsumption_resolution,[],[f5812,f111])).
fof(f5812,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10)))) ) | $spl33),
  inference(superposition,[],[f162,f4991])).
fof(f5799,plain,(
  ( ! [X85,X86,X84] : (s(bool,X84) != s(bool,X85) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X85))) | s(bool,X84) = s(bool,X86) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X86)))) ) | $spl33),
  inference(superposition,[],[f155,f4991])).
fof(f5776,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X19),'const.option.NONE_0'),s(bool,X18),s(fun(X19,bool),X20))) | s(bool,X17) = s(bool,X21) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21)))) ) | $spl33),
  inference(superposition,[],[f4437,f4991])).
fof(f5762,plain,(
  ( ! [X39,X41,X40] : (s('type.option.option'(bool),X40) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X40))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))))) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl33),
  inference(superposition,[],[f155,f4991])).
fof(f5760,plain,(
  ( ! [X35,X33,X34] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X34))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s('type.option.option'(bool),X34) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl33),
  inference(superposition,[],[f154,f4991])).
fof(f5753,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(X12,happ(s(fun(bool,X12),X13),s(bool,X11))) = s(X12,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))),s(X12,X14),s(fun(bool,X12),X13))) | s(bool,X11) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl33),
  inference(superposition,[],[f165,f4991])).
fof(f5739,plain,(
  ( ! [X39,X37,X38] : (s('type.option.option'(bool),X38) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X38))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl33),
  inference(superposition,[],[f155,f4991])).
fof(f5737,plain,(
  ( ! [X33,X31,X32] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X32))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X31))))) | s('type.option.option'(bool),X32) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl33),
  inference(superposition,[],[f154,f4991])).
fof(f5731,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(X12,happ(s(fun(bool,X12),X13),s(bool,'T_0'))) = s(X12,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X11))),s(X12,X14),s(fun(bool,X12),X13))) | s(bool,X11) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl33),
  inference(superposition,[],[f165,f4991])).
fof(f5709,plain,(
  ( ! [X138,X139,X137] : (s(bool,X138) != s(bool,X139) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X137))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X139))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X137))) | s(bool,'T_0') = s(bool,X138)) ) | $spl33),
  inference(superposition,[],[f155,f4991])).
fof(f5707,plain,(
  ( ! [X132,X133,X131] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X132))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X133))) | s(bool,X131) = s(bool,X133) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X131))) | s(bool,'T_0') = s(bool,X132)) ) | $spl33),
  inference(superposition,[],[f154,f4991])).
fof(f5655,plain,(
  ( ! [X204,X202,X205,X203] : (s(X204,f(s(bool,X205))) != s(X204,g(s(bool,X205))) | s(fun(bool,X204),X202) = s(fun(bool,X204),X203) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X205))) | s(bool,'T_0') = s(bool,sK9(X202,X203,X204,bool))) ) | $spl33),
  inference(superposition,[],[f164,f4991])).
fof(f5652,plain,(
  ( ! [X196,X195] : (s('type.option.option'(bool),X195) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X196))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X195) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X196))) | s(bool,'T_0') = s(bool,sK7(X195,bool))) ) | $spl33),
  inference(superposition,[],[f146,f4991])).
fof(f5638,plain,(
  ( ! [X138,X139,X137] : (s(bool,X138) != s(bool,X139) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X137))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X139))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X138))) | s(bool,'T_0') = s(bool,X137)) ) | $spl33),
  inference(superposition,[],[f155,f4991])).
fof(f5636,plain,(
  ( ! [X132,X133,X131] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X132))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X133))) | s(bool,X131) = s(bool,X133) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X132))) | s(bool,'T_0') = s(bool,X131)) ) | $spl33),
  inference(superposition,[],[f154,f4991])).
fof(f5581,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X38),'const.option.NONE_0'),s(bool,X39),s(fun(X38,bool),X40))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X41)))) ) | $spl33),
  inference(superposition,[],[f4991,f163])).
fof(f5560,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl33),
  inference(superposition,[],[f4991,f190])).
fof(f5559,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,X0) = s(bool,X3) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3)))) ) | $spl33),
  inference(superposition,[],[f4991,f4991])).
fof(f5554,plain,(
  ( ! [X39,X41,X38,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(bool,X41) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X38),'const.option.NONE_0'),s(bool,X39),s(fun(X38,bool),X40))) | s(bool,'T_0') = s(bool,X41)) ) | $spl33),
  inference(superposition,[],[f4991,f163])).
fof(f5529,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl33),
  inference(superposition,[],[f4991,f190])).
fof(f5528,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3)))) ) | $spl33),
  inference(superposition,[],[f4991,f4991])).
fof(f5500,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X50) = s(bool,X52) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X49),'const.option.NONE_0'),s(bool,X50),s(fun(X49,bool),X51)))) ) | $spl33),
  inference(superposition,[],[f4991,f163])).
fof(f5195,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X27)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) | $spl33),
  inference(subsumption_resolution,[],[f5103,f2948])).
fof(f5103,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X27)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) | $spl33),
  inference(superposition,[],[f4267,f4881])).
fof(f5194,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22)))) ) | $spl33),
  inference(subsumption_resolution,[],[f5099,f2948])).
fof(f5099,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22)))) ) | $spl33),
  inference(superposition,[],[f192,f4881])).
fof(f5158,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1)))) ) | $spl33),
  inference(subsumption_resolution,[],[f5157,f113])).
fof(f5157,plain,(
  ( ! [X2,X1] : (p(s(bool,'F_0')) | p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1)))) ) | $spl33),
  inference(forward_demodulation,[],[f5042,f946])).
fof(f5042,plain,(
  ( ! [X2,X1] : (p(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) | p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1)))) ) | $spl33),
  inference(superposition,[],[f159,f4881])).
fof(f4991,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,X5) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) ) | $spl33),
  inference(superposition,[],[f4881,f4280])).
fof(f4990,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X2)),bool))) ) | $spl33),
  inference(superposition,[],[f4881,f4452])).
fof(f4989,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | $spl33),
  inference(superposition,[],[f4881,f4881])).
fof(f4983,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(superposition,[],[f4881,f114])).
fof(f4903,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X2),'const.option.NONE_0'),s(bool,X1),s(fun(X2,bool),X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl33),
  inference(superposition,[],[f4437,f190])).
fof(f4652,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(superposition,[],[f4452,f114])).
fof(f4437,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X13) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X12),'const.option.NONE_0'),s(bool,X13),s(fun(X12,bool),X14)))) ) | $spl33),
  inference(forward_demodulation,[],[f4419,f946])).
fof(f4419,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X13)),bool)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X12),'const.option.NONE_0'),s(bool,X13),s(fun(X12,bool),X14)))) ) | $spl33),
  inference(superposition,[],[f4267,f163])).
fof(f18635,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(X12,happ(s(fun(bool,X12),X13),s(bool,X11))) = s(X12,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.NONE_0'),s(X12,X14),s(fun(bool,X12),X13))) | s(bool,X11) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(backward_demodulation,[],[f18458,f5753])).
fof(f18458,plain,(
  ( ! [X82] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X82)) ) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18368,f12624])).
fof(f18368,plain,(
  ( ! [X82] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X82)) ) | ($spl33 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f18088,f18067])).
fof(f18743,plain,(
  s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(backward_demodulation,[],[f18706,f12624])).
fof(f18751,plain,(
  $false | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18742,f111])).
fof(f18742,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(backward_demodulation,[],[f18706,f12618])).
fof(f12618,plain,(
  ~p(s(bool,'F_0')) | $spl109),
  inference(forward_demodulation,[],[f11828,f11817])).
fof(f11828,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl109),
  inference(unit_resulting_resolution,[],[f11666,f192])).
fof(f18750,plain,(
  $false | ($spl33 | $spl67 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18740,f18706])).
fof(f18740,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl33 | $spl67 | $spl101 | $spl109 | $spl118)),
  inference(backward_demodulation,[],[f18706,f11835])).
fof(f11835,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | ($spl67 | $spl109)),
  inference(unit_resulting_resolution,[],[f3917,f11666,f532])).
fof(f18749,plain,(
  $false | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18738,f18706])).
fof(f18738,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(backward_demodulation,[],[f18706,f11659])).
fof(f18748,plain,(
  $false | ($spl33 | $spl96 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18736,f18706])).
fof(f18736,plain,(
  ( ! [X95,X96] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X95,bool),X96)))) ) | ($spl33 | $spl96 | $spl101 | $spl109 | $spl118)),
  inference(backward_demodulation,[],[f18706,f11308])).
fof(f18696,plain,(
  ( ! [X2,X3,X1] : (s(X1,happ(s(fun(bool,X1),X2),s(bool,'T_0'))) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.NONE_0'),s(X1,X3),s(fun(bool,X1),X2)))) ) | $spl172),
  inference(cnf_transformation,[],[f18696_D])).
fof(f18696_D,plain,(
  ( ! [X2,X3,X1] : (s(X1,happ(s(fun(bool,X1),X2),s(bool,'T_0'))) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.NONE_0'),s(X1,X3),s(fun(bool,X1),X2)))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f18670,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,happ(s(fun(bool,X1),X2),s(bool,'T_0'))) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.NONE_0'),s(X1,X3),s(fun(bool,X1),X2)))) ) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(condensation,[],[f18530])).
fof(f18530,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(X12,happ(s(fun(bool,X12),X13),s(bool,'T_0'))) = s(X12,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.NONE_0'),s(X12,X14),s(fun(bool,X12),X13))) | s(bool,X11) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(backward_demodulation,[],[f18458,f5731])).
fof(f18695,plain,(
  ( ! [X14,X15,X13] : (s(X13,happ(s(fun(bool,X13),X14),s(bool,'F_0'))) = s(X13,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.NONE_0'),s(X13,X15),s(fun(bool,X13),X14)))) ) | $spl170),
  inference(cnf_transformation,[],[f18695_D])).
fof(f18695_D,plain,(
  ( ! [X14,X15,X13] : (s(X13,happ(s(fun(bool,X13),X14),s(bool,'F_0'))) = s(X13,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.NONE_0'),s(X13,X15),s(fun(bool,X13),X14)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f18493,plain,(
  ( ! [X14,X12,X15,X13] : (s(X13,happ(s(fun(bool,X13),X14),s(bool,'F_0'))) = s(X13,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.NONE_0'),s(X13,X15),s(fun(bool,X13),X14))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(backward_demodulation,[],[f18458,f3242])).
fof(f18456,plain,(
  ( ! [X105] : (s('type.option.option'(bool),X105) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X105)) ) | $spl166),
  inference(cnf_transformation,[],[f18456_D])).
fof(f18456_D,plain,(
  ( ! [X105] : (s('type.option.option'(bool),X105) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X105)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f18457,plain,(
  ( ! [X106] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X106)))) ) | $spl168),
  inference(cnf_transformation,[],[f18457_D])).
fof(f18457_D,plain,(
  ( ! [X106] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X106)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f18455,plain,(
  ( ! [X105,X106] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X106))) | s('type.option.option'(bool),X105) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X105)) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18362,f3099])).
fof(f18362,plain,(
  ( ! [X105,X106] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X106))) | s('type.option.option'(bool),X105) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X106) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X105)) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f18088,f18011])).
fof(f2475,plain,(
  ( ! [X0] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X0,sK0))) ) | $spl28),
  inference(cnf_transformation,[],[f2475_D])).
fof(f2475_D,plain,(
  ( ! [X0] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X0,sK0))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f18452,plain,(
  ( ! [X167,X166,X168] : (s(X168,f(s(bool,'F_0'))) != s(X168,g(s(bool,'F_0'))) | s(fun(bool,X168),X166) = s(fun(bool,X168),X167)) ) | $spl164),
  inference(cnf_transformation,[],[f18452_D])).
fof(f18452_D,plain,(
  ( ! [X167,X166,X168] : (s(X168,f(s(bool,'F_0'))) != s(X168,g(s(bool,'F_0'))) | s(fun(bool,X168),X166) = s(fun(bool,X168),X167)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f18451,plain,(
  ( ! [X167,X169,X166,X168] : (s(X168,f(s(bool,'F_0'))) != s(X168,g(s(bool,'F_0'))) | s(fun(bool,X168),X166) = s(fun(bool,X168),X167) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X169,sK0))) ) | ($spl0 | $spl33 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18450,f12624])).
fof(f18450,plain,(
  ( ! [X167,X169,X166,X168] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(X168,f(s(bool,'F_0'))) != s(X168,g(s(bool,'F_0'))) | s(fun(bool,X168),X166) = s(fun(bool,X168),X167) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X169,sK0))) ) | ($spl0 | $spl101 | $spl118)),
  inference(forward_demodulation,[],[f18341,f18088])).
fof(f18341,plain,(
  ( ! [X167,X169,X166,X168] : (s(X168,f(s(bool,'F_0'))) != s(X168,g(s(bool,'F_0'))) | s(fun(bool,X168),X166) = s(fun(bool,X168),X167) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X169,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK9(X166,X167,X168,bool))),bool))) ) | ($spl0 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f18088,f2366])).
fof(f3578,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X11)),bool))) ) | $spl60),
  inference(cnf_transformation,[],[f3578_D])).
fof(f3578_D,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X11)),bool))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f18449,plain,(
  ( ! [X101,X102] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X102,sK0))) ) | ($spl0 | $spl33 | $spl36 | $spl101 | $spl118)),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f170,f813,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f1472,f1467,f1460,f1479,f1481,f1432,f1485,f1433,f1486,f2131,f2466,f3186,f3570,f3325,f3575,f3327,f3223,f3290,f3102,f3581,f3058,f3091,f2263,f113,f2260,f3078,f3920,f2948,f4196,f1200,f2954,f4267,f4347,f4437,f4652,f4903,f4983,f4989,f4990,f4991,f5158,f5194,f5195,f5500,f5528,f5529,f5554,f5559,f5560,f5581,f5636,f5638,f5652,f5655,f5707,f5709,f5731,f5737,f5739,f5753,f5760,f5762,f5775,f5776,f5799,f6026,f5855,f5857,f5871,f5874,f6126,f6302,f6313,f6314,f7662,f8224,f8206,f8406,f8008,f8085,f8491,f8148,f8508,f9311,f9374,f9567,f9570,f9636,f9607,f9832,f9833,f9837,f9838,f9839,f9840,f9841,f9775,f9776,f9777,f9778,f9780,f9781,f9844,f9790,f9791,f9792,f9793,f9811,f9813,f9817,f9827,f9974,f10053,f10054,f10063,f10064,f10065,f10066,f10067,f10178,f10084,f10086,f10090,f10100,f10102,f10103,f10104,f10105,f10107,f10108,f10193,f10117,f10118,f10119,f10120,f10121,f10138,f10140,f10144,f10245,f11259,f11253,f11245,f11307,f11313,f11239,f11350,f11353,f11226,f11398,f11224,f11648,f11649,f11650,f11651,f11652,f11655,f11656,f11657,f588,f6031,f8419,f5726,f8415,f5723,f922,f2281,f1350,f1045,f2282,f1351,f3188,f2518,f2517,f2515,f2440,f2410,f2408,f2395,f2393,f2366,f2363,f2349,f2347,f2283,f2280,f1448,f1438,f1436,f1416,f1413,f1399,f1397,f1368,f1367,f1365,f1363,f1349,f1348,f1345,f1271,f1116,f928,f921,f869,f5764,f5765,f10597,f10599,f5742,f5741,f900,f621,f10603,f1402,f10576,f10566,f5860,f2521,f571,f10580,f10579,f10578,f10577,f10567,f8473,f5873,f5859,f5802,f5801,f5725,f5712,f5711,f5654,f5641,f5640,f2520,f1441,f1440,f1415,f1401,f898,f895,f893,f673,f671,f576,f573,f760,f788,f10558,f1048,f1070,f10559,f781,f1060,f10666,f10658,f10560,f10557,f9568,f8282,f6030,f5582,f5555,f5501,f5168,f4438,f1346,f1143,f1135,f1119,f912,f735,f733,f683,f681,f585,f568,f158,f157,f156,f10620,f9569,f8283,f5583,f5556,f5502,f4439,f2479,f1347,f1062,f1061,f1032,f1018,f1016,f2478,f2477,f5867,f5719,f5648,f2359,f1409,f936,f934,f11312,f11165,f11123,f11001,f10971,f10963,f10959,f10958,f10955,f10945,f10944,f10943,f10942,f10927,f10926,f10925,f10924,f10923,f10912,f10911,f10910,f10906,f10905,f10904,f10895,f10894,f10846,f10656,f10655,f10636,f10634,f10604,f10602,f10601,f10596,f10595,f10582,f10565,f10204,f10202,f10201,f10196,f10195,f10194,f10192,f10188,f10185,f10184,f10171,f10170,f10168,f10152,f10145,f10143,f10142,f10137,f10136,f10124,f10098,f10091,f10089,f10088,f10083,f10082,f10069,f10052,f9894,f9852,f9851,f9847,f9846,f9845,f9843,f9836,f9831,f9825,f9822,f9818,f9816,f9815,f9810,f9809,f9797,f9740,f9739,f9738,f9737,f9644,f9514,f9494,f9493,f9492,f9491,f9476,f9474,f9465,f9464,f9462,f9460,f9392,f9391,f9389,f9387,f9381,f9371,f9370,f9368,f9364,f9357,f9356,f9354,f9352,f8734,f8675,f8556,f8548,f8518,f8517,f8516,f8515,f7694,f7681,f6207,f6206,f6204,f6202,f6195,f6180,f6179,f6177,f6175,f6169,f5777,f5420,f5200,f5199,f5197,f5196,f5188,f5187,f5184,f5177,f5176,f5175,f5173,f5109,f5068,f5023,f4955,f4949,f4882,f4881,f4863,f4858,f4851,f4845,f4783,f4753,f4716,f4578,f4520,f4486,f4452,f4280,f3337,f3300,f3299,f3298,f3297,f3279,f3278,f3276,f3274,f3267,f3253,f3252,f3250,f3248,f3242,f3231,f3230,f3229,f3099,f3070,f3069,f3068,f3067,f2516,f2252,f2251,f2250,f2249,f2248,f2217,f2216,f2214,f2212,f2206,f2205,f2194,f2193,f2191,f2189,f2183,f2179,f2177,f2176,f2175,f2167,f2161,f1999,f1997,f1996,f1995,f1984,f1983,f1982,f1981,f1980,f1335,f1331,f1308,f1285,f1284,f1283,f1253,f1204,f1203,f1201,f1199,f1197,f1196,f1195,f1176,f1093,f1020,f965,f938,f933,f897,f892,f737,f675,f610,f575,f570,f537,f532,f497,f492,f239,f10198,f10133,f10132,f10131,f9806,f9805,f9804,f9803,f8308,f8307,f8306,f8232,f8018,f2178,f1998,f11659,f11673,f17893,f17891,f17895,f17958,f18062,f18063,f18064,f18065,f17971,f17972,f17973,f17974,f17975,f17976,f17977,f17978,f17979,f18066,f18067,f18068,f17999,f18000,f18001,f18002,f18003,f18004,f18005,f18006,f18007,f18008,f18009,f18010,f18011,f18012,f18013,f18069,f18070,f18071,f18072,f18073,f18019,f18022,f18024,f18074,f18075,f18030,f18032,f18076,f18036,f18038,f18042,f18044,f18046,f18047,f18077,f18079,f18080,f18083,f18088,f18411,f18448])).
fof(f18448,plain,(
  ( ! [X103,X101,X102] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X102,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool))) ) | ($spl0 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18340,f18088])).
fof(f18340,plain,(
  ( ! [X103,X101,X102] : (s(bool,'F_0') != s(bool,X103) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X102,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool))) ) | ($spl0 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f18088,f2349])).
fof(f9644,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl33 | $spl36)),
  inference(condensation,[],[f9643])).
fof(f9643,plain,(
  ( ! [X28,X29,X27] : (s(bool,X28) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f9578,f4267])).
fof(f9578,plain,(
  ( ! [X28,X29,X27] : (s(bool,X28) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X27)),bool)) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f4347,f4280])).
fof(f9737,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'F_0') != s(bool,X78) | s(bool,'T_0') = s(bool,X80) | s(bool,X80) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.NONE_0'),s(bool,X78),s(fun(X77,bool),X79)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f9636,f163])).
fof(f9738,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,'F_0') != s(bool,happ(s(fun(X81,bool),X84),s(X81,X82))) | s(bool,'T_0') = s(bool,X85) | s(bool,X85) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X81),'const.option.SOME_1'(s(X81,X82))),s(bool,X83),s(fun(X81,bool),X84)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f9636,f165])).
fof(f9739,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X86),X87),s(bool,X88),s(fun(X86,bool),X89))) | s(bool,'T_0') = s(bool,X90) | s(bool,X90) = s(bool,happ(s(fun(fun(X86,bool),bool),happ(s(fun(bool,fun(fun(X86,bool),bool)),happ(s(fun('type.option.option'(X86),fun(bool,fun(fun(X86,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X86),X87))),s(bool,X88))),s(fun(X86,bool),X89)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f9636,f166])).
fof(f9740,plain,(
  ( ! [X92,X91] : (s(bool,'F_0') != s(bool,X91) | s(bool,'T_0') = s(bool,X92) | s(bool,X92) = s(bool,sK7('const.option.SOME_1'(s(bool,X91)),bool))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f9636,f946])).
fof(f9894,plain,(
  ( ! [X92,X91] : (s(bool,X91) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X92) | s(bool,X92) = s(bool,sK7('const.option.SOME_1'(s(bool,X91)),bool))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f9644,f946])).
fof(f10846,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl33 | $spl36)),
  inference(condensation,[],[f10845])).
fof(f10845,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'T_0') = s(bool,X70) | s(bool,X70) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X71),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X71,bool),X72))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl33 | $spl36)),
  inference(trivial_inequality_removal,[],[f10825])).
fof(f10825,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X69) != s(bool,X69) | s(bool,'T_0') = s(bool,X70) | s(bool,X70) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X71),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X71,bool),X72))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f9737,f114])).
fof(f11001,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,X122) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X124),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X124,bool),X125))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X121),'const.option.NONE_0'),s(bool,X122),s(fun(X121,bool),X123)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f10846,f163])).
fof(f11123,plain,(
  ( ! [X424,X426,X425,X427] : (s(bool,X427) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X425),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X425,bool),X426))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X424))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X427))) | s(bool,'T_0') = s(bool,X424)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f155,f10846])).
fof(f11165,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,happ(s(fun(bool,bool),X80),s(bool,X79)))) | p(s(bool,happ(s(fun(bool,bool),X80),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8('const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X77,bool),X78)),X80,bool)) | s(bool,'T_0') = s(bool,X79)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f675,f10846])).
fof(f11312,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X2),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X2,bool),X3))) | s(bool,X0) = s(bool,X1)) ) | ($spl33 | $spl36)),
  inference(condensation,[],[f11311])).
fof(f11311,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,X111) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X109),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X109,bool),X110))) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,X108)) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f11066,f4267])).
fof(f11066,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,X111) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X109),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X109,bool),X110))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X108)),bool)) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,X108)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f4347,f10846])).
fof(f9569,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (s(bool,X129) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X125),X126),s(bool,X127),s(fun(X125,bool),X128))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X129)),bool)) | s(bool,X130) = s(bool,happ(s(fun(fun(X125,bool),bool),happ(s(fun(bool,fun(fun(X125,bool),bool)),happ(s(fun('type.option.option'(X125),fun(bool,fun(fun(X125,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X125),X126))),s(bool,X127))),s(fun(X125,bool),X128))) | s(bool,'T_0') = s(bool,X130)) ) | $spl36),
  inference(superposition,[],[f4347,f166])).
fof(f9568,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s(bool,X123) != s(bool,happ(s(fun(X119,bool),X122),s(X119,X120))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X123)),bool)) | s(bool,X124) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X119),'const.option.SOME_1'(s(X119,X120))),s(bool,X121),s(fun(X119,bool),X122))) | s(bool,'T_0') = s(bool,X124)) ) | $spl36),
  inference(superposition,[],[f4347,f165])).
fof(f11224,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X1,bool),X2))) | s(bool,X0) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X1,bool),X2)))) ) | ($spl33 | $spl36)),
  inference(equality_factoring,[],[f10846])).
fof(f10245,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f9894,f114])).
fof(f9974,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f9740,f114])).
fof(f9570,plain,(
  ( ! [X132,X133,X131] : (s(bool,X131) != s(bool,X132) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X132)),bool)) | s(bool,X133) = s(bool,sK7('const.option.SOME_1'(s(bool,X131)),bool)) | s(bool,'T_0') = s(bool,X133)) ) | $spl36),
  inference(superposition,[],[f4347,f946])).
fof(f9567,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X115) != s(bool,X117) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X117)),bool)) | s(bool,X118) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X114),'const.option.NONE_0'),s(bool,X115),s(fun(X114,bool),X116))) | s(bool,'T_0') = s(bool,X118)) ) | $spl36),
  inference(superposition,[],[f4347,f163])).
fof(f5775,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X12)),bool)) | s(bool,X11) = s(bool,X13) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f2954,f4991])).
fof(f18419,plain,(
  $false | ($spl33 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18123,f12624])).
fof(f18123,plain,(
  s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f18088,f17893])).
fof(f18418,plain,(
  $false | ($spl96 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18122,f11308])).
fof(f18122,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X0,bool),X1)))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f18088,f17891])).
fof(f18415,plain,(
  ( ! [X142,X143,X141] : (s(fun(X141,bool),X142) = s(fun(X141,bool),X143)) ) | $spl162),
  inference(cnf_transformation,[],[f18415_D])).
fof(f18415_D,plain,(
  ( ! [X142,X143,X141] : (s(fun(X141,bool),X142) = s(fun(X141,bool),X143)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f18414,plain,(
  ( ! [X144,X142,X143,X141] : (s(fun(X141,bool),X142) = s(fun(X141,bool),X143) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X144,sK0))) ) | ($spl0 | $spl33 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f18413,f12624])).
fof(f18413,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(fun(X141,bool),X142) = s(fun(X141,bool),X143) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X144,sK0))) ) | ($spl0 | $spl101 | $spl118)),
  inference(forward_demodulation,[],[f18412,f18088])).
fof(f18412,plain,(
  ( ! [X144,X142,X143,X141] : (s(fun(X141,bool),X142) = s(fun(X141,bool),X143) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X144,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(X141,sK9(X142,X143,bool,X141))))),bool))) ) | ($spl0 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18112,f18088])).
fof(f18112,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X144),s(sK0,sK2))) | s(fun(X141,bool),X142) = s(fun(X141,bool),X143) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X144,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(X141,sK9(X142,X143,bool,X141))))),bool))) ) | ($spl0 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f18088,f2359])).
fof(f4245,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(bool,X32),s(fun(X30,bool),X33)))) ) | $spl74),
  inference(cnf_transformation,[],[f4245_D])).
fof(f4245_D,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(bool,X32),s(fun(X30,bool),X33)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f198,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f18083,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(bool,X32),s(fun(X30,bool),X33)))) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18059,f17958])).
fof(f18059,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(bool,X32),s(fun(X30,bool),X33)))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f1981])).
fof(f18081,plain,(
  ( ! [X39,X38] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X38,sK0)) | s(bool,'F_0') = s(bool,X39)) ) | ($spl26 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18054,f17958])).
fof(f18054,plain,(
  ( ! [X39,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X38,sK0)) | s(bool,'F_0') = s(bool,X39)) ) | ($spl26 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f2154])).
fof(f4249,plain,(
  ( ! [X137] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(bool,'T_0') = s(bool,sK7(X137,bool))) ) | $spl76),
  inference(cnf_transformation,[],[f4249_D])).
fof(f4249_D,plain,(
  ( ! [X137] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(bool,'T_0') = s(bool,sK7(X137,bool))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f18080,plain,(
  ( ! [X138,X137] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X138,sK0)) | s(bool,'T_0') = s(bool,sK7(X137,bool))) ) | ($spl0 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18052,f611])).
fof(f18052,plain,(
  ( ! [X138,X137] : (s('type.option.option'(bool),X137) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X138,sK0)) | s(bool,'T_0') = s(bool,sK7(X137,bool))) ) | ($spl0 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f1413])).
fof(f18079,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0))) ) | ($spl0 | $spl33 | $spl101 | $spl118)),
  inference(condensation,[],[f18078])).
fof(f18078,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) = s(bool,X75) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X74,sK0)) | s(bool,'T_0') = s(bool,X73)) ) | ($spl0 | $spl33 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18051,f4991])).
fof(f18051,plain,(
  ( ! [X74,X75,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75))) | s(bool,X73) = s(bool,X75) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X74,sK0)) | s(bool,'T_0') = s(bool,X73)) ) | ($spl0 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f1397])).
fof(f4243,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X35,bool),bool),happ(s(fun(bool,fun(fun(X35,bool),bool)),happ(s(fun('type.option.option'(X35),fun(bool,fun(fun(X35,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X35),X36))),s(bool,X37))),s(fun(X35,bool),X38)))) ) | $spl72),
  inference(cnf_transformation,[],[f4243_D])).
fof(f4243_D,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X35,bool),bool),happ(s(fun(bool,fun(fun(X35,bool),bool)),happ(s(fun('type.option.option'(X35),fun(bool,fun(fun(X35,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X35),X36))),s(bool,X37))),s(fun(X35,bool),X38)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f18077,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X35,bool),bool),happ(s(fun(bool,fun(fun(X35,bool),bool)),happ(s(fun('type.option.option'(X35),fun(bool,fun(fun(X35,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X35),X36))),s(bool,X37))),s(fun(X35,bool),X38)))) ) | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f18049,f17958])).
fof(f18049,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X35,bool),bool),happ(s(fun(bool,fun(fun(X35,bool),bool)),happ(s(fun('type.option.option'(X35),fun(bool,fun(fun(X35,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X35),X36))),s(bool,X37))),s(fun(X35,bool),X38)))) ) | ($spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f1982])).
fof(f18061,plain,(
  $false | ($spl67 | $spl101 | $spl118)),
  inference(trivial_inequality_removal,[],[f17985])).
fof(f17985,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | ($spl67 | $spl101 | $spl118)),
  inference(backward_demodulation,[],[f17958,f6107])).
fof(f6107,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | $spl67),
  inference(unit_resulting_resolution,[],[f3917,f154])).
fof(f17956,plain,(
  $false | ($spl15 | $spl101 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f17955,f11659])).
fof(f17955,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | ($spl15 | $spl109 | $spl118)),
  inference(subsumption_resolution,[],[f17906,f12619])).
fof(f12619,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ($spl15 | $spl109)),
  inference(forward_demodulation,[],[f11831,f11817])).
fof(f11831,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,f(s(bool,'F_0'))))) | ($spl15 | $spl109)),
  inference(unit_resulting_resolution,[],[f4197,f11666,f492])).
fof(f17906,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl118),
  inference(superposition,[],[f11673,f114])).
fof(f17949,plain,(
  $false | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17948,f11659])).
fof(f17948,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl118),
  inference(forward_demodulation,[],[f17886,f163])).
fof(f17886,plain,(
  ( ! [X0,X1] : (s(bool,g(s(bool,'F_0'))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X0,bool),X1)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f11673,f491])).
fof(f17947,plain,(
  $false | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17887,f11659])).
fof(f17887,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl118),
  inference(unit_resulting_resolution,[],[f11673,f154])).
fof(f17946,plain,(
  $false | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17945,f11659])).
fof(f17945,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl118),
  inference(forward_demodulation,[],[f17888,f946])).
fof(f17888,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | $spl118),
  inference(unit_resulting_resolution,[],[f151,f11673,f611])).
fof(f17944,plain,(
  $false | ($spl15 | $spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17890,f11659])).
fof(f17890,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | ($spl15 | $spl118)),
  inference(unit_resulting_resolution,[],[f4197,f11673,f492])).
fof(f17943,plain,(
  $false | ($spl101 | $spl118)),
  inference(subsumption_resolution,[],[f17892,f11659])).
fof(f17892,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl118),
  inference(unit_resulting_resolution,[],[f11673,f154])).
fof(f17849,plain,(
  $false | ($spl101 | $spl126)),
  inference(subsumption_resolution,[],[f11659,f15444])).
fof(f15444,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl126),
  inference(equality_resolution,[],[f14979])).
fof(f17781,plain,(
  ( ! [X54,X56] : (s(bool,X56) = s(bool,g(s(bool,X54))) | s(bool,'T_0') = s(bool,X56)) ) | $spl152),
  inference(cnf_transformation,[],[f17781_D])).
fof(f17781_D,plain,(
  ( ! [X54,X56] : (s(bool,X56) = s(bool,g(s(bool,X54))) | s(bool,'T_0') = s(bool,X56)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f17532,plain,(
  ( ! [X144,X145,X143] : (s(bool,'F_0') != s(bool,X144) | s(bool,'T_0') = s(bool,X145) | s(bool,X145) = s(bool,g(s(bool,X143))) | s(bool,X144) = s(bool,f(s(bool,X144)))) ) | ($spl30 | $spl33 | $spl36 | $spl109 | $spl128)),
  inference(superposition,[],[f9636,f17195])).
fof(f17195,plain,(
  ( ! [X50,X51] : (s(bool,X51) = s(bool,g(s(bool,X50))) | s(bool,X51) = s(bool,f(s(bool,X51)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(subsumption_resolution,[],[f17159,f111])).
fof(f17159,plain,(
  ( ! [X50,X51] : (~p(s(bool,'T_0')) | s(bool,X51) = s(bool,g(s(bool,X50))) | s(bool,X51) = s(bool,f(s(bool,X51)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f17014,f16203])).
fof(f17803,plain,(
  ( ! [X130,X131] : (s(bool,X131) != s(bool,sK7('const.option.SOME_1'(s(bool,X130)),bool)) | s(bool,X130) = s(bool,f(s(bool,X130))) | s(bool,'T_0') = s(bool,X131)) ) | $spl160),
  inference(cnf_transformation,[],[f17803_D])).
fof(f17803_D,plain,(
  ( ! [X130,X131] : (s(bool,X131) != s(bool,sK7('const.option.SOME_1'(s(bool,X130)),bool)) | s(bool,X130) = s(bool,f(s(bool,X130))) | s(bool,'T_0') = s(bool,X131)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f17527,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X131) != s(bool,sK7('const.option.SOME_1'(s(bool,X130)),bool)) | s(bool,'T_0') = s(bool,X131) | s(bool,X132) = s(bool,g(s(bool,X129))) | s(bool,'T_0') = s(bool,X132) | s(bool,X130) = s(bool,f(s(bool,X130)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f7397,f17195])).
fof(f17792,plain,(
  ( ! [X88,X89] : (s(bool,X88) != s(bool,X89) | s(bool,X88) = s(bool,f(s(bool,X88))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X89)),bool))) ) | $spl158),
  inference(cnf_transformation,[],[f17792_D])).
fof(f17792_D,plain,(
  ( ! [X88,X89] : (s(bool,X88) != s(bool,X89) | s(bool,X88) = s(bool,f(s(bool,X88))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X89)),bool))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f17513,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X88) != s(bool,X89) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X89)),bool)) | s(bool,X90) = s(bool,g(s(bool,X87))) | s(bool,'T_0') = s(bool,X90) | s(bool,X88) = s(bool,f(s(bool,X88)))) ) | ($spl30 | $spl33 | $spl36 | $spl109 | $spl128)),
  inference(superposition,[],[f4347,f17195])).
fof(f17786,plain,(
  ( ! [X64,X65] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65))) | s(bool,'F_0') != s(bool,X64) | s(bool,X64) = s(bool,f(s(bool,X64)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(subsumption_resolution,[],[f17785,f155])).
fof(f17785,plain,(
  ( ! [X64,X65] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65))) | s(bool,'F_0') != s(bool,X64) | s(bool,'T_0') = s(bool,X65) | s(bool,X64) = s(bool,f(s(bool,X64)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(forward_demodulation,[],[f17504,f16203])).
fof(f17504,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') != s(bool,X64) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,X63))))) | s(bool,'T_0') = s(bool,X65) | s(bool,X64) = s(bool,f(s(bool,X64)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f537,f17195])).
fof(f17501,plain,(
  ( ! [X54,X56,X55] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X55))) | s(bool,X56) = s(bool,g(s(bool,X54))) | s(bool,'T_0') = s(bool,X56) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f497,f17195])).
fof(f17775,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X21) | p(s(bool,X20)) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(subsumption_resolution,[],[f17774,f192])).
fof(f17774,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X21) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(forward_demodulation,[],[f17489,f16203])).
fof(f17489,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) = s(bool,g(s(bool,X19))) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f159,f17195])).
fof(f17772,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,X18) | p(s(bool,X17)) | s(bool,X17) = s(bool,f(s(bool,X17)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(subsumption_resolution,[],[f17771,f192])).
fof(f17771,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,X18) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) = s(bool,f(s(bool,X17)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(forward_demodulation,[],[f17488,f16203])).
fof(f17488,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) = s(bool,g(s(bool,X16))) | s(bool,X17) = s(bool,f(s(bool,X17)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f159,f17195])).
fof(f17713,plain,(
  ( ! [X45,X44] : (s(bool,X45) != s(bool,g(s(bool,X44))) | s(bool,'T_0') = s(bool,X45)) ) | $spl142),
  inference(cnf_transformation,[],[f17713_D])).
fof(f17713_D,plain,(
  ( ! [X45,X44] : (s(bool,X45) != s(bool,g(s(bool,X44))) | s(bool,'T_0') = s(bool,X45)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f17718,plain,(
  ( ! [X76,X73] : (s(bool,X73) = s(bool,X76) | s(bool,X73) = s(bool,f(s(bool,X73))) | s(bool,'T_0') = s(bool,X76)) ) | $spl144),
  inference(cnf_transformation,[],[f17718_D])).
fof(f17718_D,plain,(
  ( ! [X76,X73] : (s(bool,X73) = s(bool,X76) | s(bool,X73) = s(bool,f(s(bool,X73))) | s(bool,'T_0') = s(bool,X76)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f17747,plain,(
  ( ! [X261,X259,X260,X258] : (s(bool,X258) = s(bool,f(s(bool,X258))) | s(bool,X260) != s(bool,g(s(bool,X259))) | s(bool,'T_0') = s(bool,X260) | s(bool,X258) = s(bool,X261) | s(bool,'T_0') = s(bool,X261)) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(forward_demodulation,[],[f17450,f946])).
fof(f17450,plain,(
  ( ! [X261,X259,X260,X258] : (s(bool,X260) != s(bool,g(s(bool,X259))) | s(bool,'T_0') = s(bool,X260) | s(bool,X258) = s(bool,X261) | s(bool,'T_0') = s(bool,X261) | s(bool,sK7('const.option.SOME_1'(s(bool,X258)),bool)) = s(bool,f(s(bool,sK7('const.option.SOME_1'(s(bool,X258)),bool))))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f7397,f17195])).
fof(f17741,plain,(
  ( ! [X241,X240] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X241) != s(bool,g(s(bool,X240))) | s(bool,'T_0') = s(bool,X241)) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(forward_demodulation,[],[f17443,f946])).
fof(f17443,plain,(
  ( ! [X241,X240] : (s(bool,X241) != s(bool,g(s(bool,X240))) | s(bool,'T_0') = s(bool,X241) | s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) = s(bool,f(s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool))))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f7298,f17195])).
fof(f17719,plain,(
  ( ! [X74,X75] : (s(bool,X75) != s(bool,g(s(bool,X74))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X75)),bool))) ) | $spl146),
  inference(cnf_transformation,[],[f17719_D])).
fof(f17719_D,plain,(
  ( ! [X74,X75] : (s(bool,X75) != s(bool,g(s(bool,X74))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X75)),bool))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f17740,plain,(
  ( ! [X239,X238] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X239) != s(bool,g(s(bool,X238))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X239)),bool))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(forward_demodulation,[],[f17442,f946])).
fof(f17442,plain,(
  ( ! [X239,X238] : (s(bool,X239) != s(bool,g(s(bool,X238))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X239)),bool)) | s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) = s(bool,f(s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool))))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f7441,f17195])).
fof(f7441,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool))) ) | ($spl30 | $spl33)),
  inference(condensation,[],[f7440])).
fof(f7440,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X10)),bool))) ) | ($spl30 | $spl33)),
  inference(subsumption_resolution,[],[f7374,f4267])).
fof(f7374,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X9)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X10)),bool))) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f7298,f2476])).
fof(f17739,plain,(
  ( ! [X237,X236] : (s(bool,X237) != s(bool,g(s(bool,X236))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X237)))) ) | $spl148),
  inference(cnf_transformation,[],[f17739_D])).
fof(f17739_D,plain,(
  ( ! [X237,X236] : (s(bool,X237) != s(bool,g(s(bool,X236))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X237)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f17738,plain,(
  ( ! [X237,X236] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X237) != s(bool,g(s(bool,X236))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X237)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(forward_demodulation,[],[f17441,f946])).
fof(f17441,plain,(
  ( ! [X237,X236] : (s(bool,X237) != s(bool,g(s(bool,X236))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X237))) | s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) = s(bool,f(s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool))))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f7453,f17195])).
fof(f7453,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29)))) ) | ($spl30 | $spl33)),
  inference(subsumption_resolution,[],[f7386,f2948])).
fof(f7386,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29)))) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f7298,f4881])).
fof(f17392,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X75) != s(bool,g(s(bool,X74))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X75)),bool)) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,X73) = s(bool,f(s(bool,X73)))) ) | ($spl30 | $spl33 | $spl36 | $spl109 | $spl128)),
  inference(superposition,[],[f4347,f17195])).
fof(f17712,plain,(
  ( ! [X43] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,X43) = s(bool,f(s(bool,X43)))) ) | $spl140),
  inference(cnf_transformation,[],[f17712_D])).
fof(f17712_D,plain,(
  ( ! [X43] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,X43) = s(bool,f(s(bool,X43)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f17381,plain,(
  ( ! [X45,X43,X44] : (s(bool,X45) != s(bool,g(s(bool,X44))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'T_0') = s(bool,X45) | s(bool,X43) = s(bool,f(s(bool,X43)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f532,f17195])).
fof(f17681,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,g(s(bool,X26))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s(bool,X25) = s(bool,X27)) ) | $spl138),
  inference(cnf_transformation,[],[f17681_D])).
fof(f17681_D,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,g(s(bool,X26))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s(bool,X25) = s(bool,X27)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f17318,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X25) = s(bool,g(s(bool,X26))) | s(bool,X25) = s(bool,X27) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f17195,f4991])).
fof(f17676,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X2) = s(bool,X4)) ) | $spl134),
  inference(cnf_transformation,[],[f17676_D])).
fof(f17676_D,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X2) = s(bool,X4)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f17313,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X11) = s(bool,g(s(bool,X12))) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f17195,f190])).
fof(f17310,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X2) = s(bool,g(s(bool,X3))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | ($spl30 | $spl33 | $spl109 | $spl128)),
  inference(superposition,[],[f17195,f190])).
fof(f16203,plain,(
  ( ! [X34] : (s(bool,'T_0') = s(bool,g(s(bool,X34)))) ) | $spl128),
  inference(cnf_transformation,[],[f16203_D])).
fof(f16203_D,plain,(
  ( ! [X34] : (s(bool,'T_0') = s(bool,g(s(bool,X34)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f16205,plain,(
  ( ! [X37,X36] : (s(bool,X37) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,g(s(bool,X36)))) ) | ($spl30 | $spl33 | $spl100 | $spl124)),
  inference(subsumption_resolution,[],[f16152,f15398])).
fof(f15398,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(condensation,[],[f15397])).
fof(f15397,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X43) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,g(s(bool,X42)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(duplicate_literal_removal,[],[f15396])).
fof(f15396,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X43) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,g(s(bool,X42)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(forward_demodulation,[],[f15351,f11660])).
fof(f15351,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) | s(bool,X43) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,g(s(bool,X42)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(superposition,[],[f497,f14522])).
fof(f14522,plain,(
  ( ! [X46] : (s(bool,X46) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X46)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(subsumption_resolution,[],[f14493,f111])).
fof(f14493,plain,(
  ( ! [X46] : (~p(s(bool,'T_0')) | s(bool,X46) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X46)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(superposition,[],[f14273,f11660])).
fof(f14273,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,g(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(forward_demodulation,[],[f14272,f946])).
fof(f14272,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | s(bool,X21) = s(bool,sK7('const.option.SOME_1'(s(bool,X20)),bool)) | s(bool,'T_0') = s(bool,g(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(subsumption_resolution,[],[f14167,f111])).
fof(f14167,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | s(bool,X21) = s(bool,sK7('const.option.SOME_1'(s(bool,X20)),bool)) | s(bool,'T_0') = s(bool,g(s(bool,X20)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(superposition,[],[f162,f11780])).
fof(f11780,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X5)),bool)) | s(bool,'T_0') = s(bool,g(s(bool,X5)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(subsumption_resolution,[],[f11689,f2948])).
fof(f11689,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X5)),bool))) ) | ($spl30 | $spl100)),
  inference(superposition,[],[f11660,f2476])).
fof(f16152,plain,(
  ( ! [X37,X36] : (s(bool,X37) != s(bool,g(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,g(s(bool,X36)))) ) | ($spl30 | $spl33 | $spl100 | $spl124)),
  inference(superposition,[],[f532,f15257])).
fof(f15257,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X30)))) ) | ($spl30 | $spl33 | $spl100 | $spl124)),
  inference(subsumption_resolution,[],[f15192,f111])).
fof(f15192,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X30)))) ) | ($spl30 | $spl33 | $spl100 | $spl124)),
  inference(superposition,[],[f14273,f14863])).
fof(f16204,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X35)) ) | $spl130),
  inference(cnf_transformation,[],[f16204_D])).
fof(f16204_D,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X35)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f16202,plain,(
  ( ! [X35,X34] : (s(bool,X35) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,g(s(bool,X34)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(subsumption_resolution,[],[f16151,f15398])).
fof(f16151,plain,(
  ( ! [X35,X34] : (s(bool,X35) != s(bool,g(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,g(s(bool,X34)))) ) | ($spl30 | $spl33 | $spl100)),
  inference(superposition,[],[f532,f14522])).
fof(f11664,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl106),
  inference(cnf_transformation,[],[f11664_D])).
fof(f11664_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f14979,plain,(
  ( ! [X70] : (s(bool,'F_0') != s(bool,X70) | s(bool,'T_0') = s(bool,g(s(bool,X70)))) ) | $spl126),
  inference(cnf_transformation,[],[f14979_D])).
fof(f14979_D,plain,(
  ( ! [X70] : (s(bool,'F_0') != s(bool,X70) | s(bool,'T_0') = s(bool,g(s(bool,X70)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f14749,plain,(
  ( ! [X70,X71] : (s(bool,'F_0') != s(bool,X70) | s(bool,'T_0') = s(bool,X71) | s(bool,X71) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X70)))) ) | ($spl30 | $spl33 | $spl36 | $spl100)),
  inference(superposition,[],[f9636,f14522])).
fof(f14861,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl122),
  inference(cnf_transformation,[],[f14861_D])).
fof(f14861_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f14863,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl124),
  inference(cnf_transformation,[],[f14863_D])).
fof(f14863_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f14588,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X7) = s(bool,g(s(bool,'F_0'))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl30 | $spl33 | $spl100)),
  inference(superposition,[],[f14522,f190])).
fof(f14585,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'F_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl30 | $spl33 | $spl100)),
  inference(superposition,[],[f14522,f190])).
fof(f12601,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f12600,f111])).
fof(f12600,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(forward_demodulation,[],[f11828,f12012])).
fof(f12012,plain,(
  ( ! [X43] : (s(bool,'T_0') = s(bool,X43)) ) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(global_subsumption,[],[f12011,f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f588,f3186,f2263,f3242,f3099,f2252,f933,f610,f532,f497,f113,f2260,f3078,f3920,f2948,f938,f936,f934,f1016,f1018,f1020,f1032,f1061,f1062,f1197,f1285,f156,f157,f158,f568,f585,f681,f683,f733,f735,f737,f912,f1119,f1135,f1143,f1196,f1284,f1060,f781,f1048,f1070,f760,f788,f570,f573,f576,f671,f673,f675,f892,f893,f895,f897,f898,f1203,f3297,f3300,f571,f1204,f1335,f575,f3298,f3299,f900,f621,f2194,f2193,f2217,f2216,f3337,f4196,f239,f492,f537,f1093,f1176,f1195,f1199,f1200,f1201,f1253,f1283,f1308,f1331,f1980,f1984,f1995,f1999,f2161,f2167,f2175,f2179,f2183,f2189,f2191,f2205,f2206,f2212,f2214,f2954,f1981,f1996,f2176,f1982,f1997,f2177,f3279,f3278,f3276,f3274,f3267,f3253,f3252,f3250,f3248,f3231,f3230,f3229,f3070,f3069,f3068,f3067,f2251,f2250,f2249,f2248,f965,f2178,f1998,f1983,f4267,f4280,f4347,f4437,f4438,f4439,f4452,f4486,f4520,f4578,f4652,f4845,f4716,f4851,f4858,f4863,f4753,f4881,f4783,f4882,f4903,f4949,f4955,f4983,f4989,f4990,f4991,f5023,f5158,f5168,f5173,f5068,f5175,f5176,f5177,f5184,f5187,f5188,f5194,f5195,f5196,f5109,f5197,f5199,f5200,f5420,f5500,f5501,f5502,f5528,f5529,f5554,f5555,f5556,f5559,f5560,f5581,f5582,f5583,f5636,f5638,f5640,f5641,f5648,f5652,f5654,f5655,f5707,f5709,f5711,f5712,f5719,f5723,f5725,f5726,f5731,f5737,f5739,f5741,f5742,f5753,f5760,f5762,f5764,f5765,f5775,f5776,f5777,f5799,f5801,f5802,f6026,f6030,f6031,f5855,f5857,f5859,f5860,f5867,f5871,f5873,f5874,f6126,f6169,f6175,f6177,f6179,f6180,f6195,f6202,f6204,f6206,f6207,f6302,f6313,f6314,f7662,f7681,f7694,f8232,f8224,f8282,f8283,f8206,f8306,f8307,f8308,f8406,f8008,f8018,f8415,f8419,f8085,f8473,f8491,f8148,f8508,f8515,f8516,f8517,f8518,f8548,f8556,f8675,f8734,f9491,f9492,f9493,f9494,f9311,f9352,f9354,f9356,f9357,f9364,f9368,f9370,f9371,f9374,f9381,f9387,f9389,f9391,f9392,f9514,f9476,f9474,f9460,f9462,f9464,f9465,f9567,f9568,f9569,f9570,f9636,f9644,f9607,f9737,f9738,f9739,f9740,f9831,f9832,f9833,f9836,f9837,f9838,f9839,f9840,f9841,f9775,f9776,f9777,f9778,f9780,f9781,f9843,f9844,f9845,f9790,f9791,f9792,f9793,f9797,f9846,f9847,f9803,f9804,f9805,f9806,f9809,f9810,f9811,f9851,f9813,f9852,f9815,f9816,f9817,f9818,f9822,f9825,f9827,f9894,f9974,f10052,f10053,f10054,f10168,f10063,f10064,f10065,f10066,f10067,f10069,f10170,f10171,f10178,f10082,f10083,f10084,f10184,f10086,f10185,f10088,f10089,f10090,f10091,f10188,f10098,f10100,f10102,f10103,f10104,f10105,f10107,f10108,f10192,f10193,f10194,f10117,f10118,f10119,f10120,f10121,f10124,f10195,f10196,f10198,f10131,f10132,f10133,f10136,f10137,f10138,f10201,f10140,f10202,f10142,f10143,f10144,f10145,f10204,f10152,f10245,f10557,f10558,f10559,f10560,f10565,f10566,f10567,f10634,f10576,f10577,f10578,f10579,f10580,f10582,f10636,f10595,f10596,f10597,f10655,f10599,f10656,f10601,f10602,f10603,f10604,f10658,f10666,f10620,f10846,f10942,f10943,f10944,f10945,f10894,f10895,f10904,f10905,f10906,f10910,f10911,f10912,f10955,f10958,f10959,f10923,f10924,f10925,f10926,f10927,f10963,f10971,f11259,f11001,f11253,f11245,f11307,f11312,f11313,f11239,f11350,f11353,f11123,f11165,f11226,f11398,f11224,f11648,f11649,f11650,f11651,f11652,f11655,f11656,f11657,f11670,f11666,f11919,f11910,f11956,f11960])).
fof(f11960,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X14)),bool))) ) | ($spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11959,f1093])).
fof(f11959,plain,(
  ( ! [X14] : (s(bool,'F_0') != s(bool,X14) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X14)),bool))) ) | ($spl33 | $spl36 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f11956,f11670])).
fof(f11956,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl33 | $spl36 | $spl109)),
  inference(forward_demodulation,[],[f11895,f163])).
fof(f11895,plain,(
  ( ! [X0,X1] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X0,bool),X1)))) ) | ($spl33 | $spl36 | $spl109)),
  inference(unit_resulting_resolution,[],[f11666,f10846])).
fof(f11910,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl109),
  inference(unit_resulting_resolution,[],[f111,f11666,f162])).
fof(f11919,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,f(s(bool,'F_0'))))) | $spl109),
  inference(unit_resulting_resolution,[],[f11666,f154])).
fof(f12011,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11961,f4197])).
fof(f11961,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,'F_0') | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl33 | $spl36 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f11960,f1998])).
fof(f12598,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11831,f12476])).
fof(f12476,plain,(
  ( ! [X8] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8)))) ) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f12029,f5195])).
fof(f12029,plain,(
  ( ! [X8] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X8)),bool))) ) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f12012,f965])).
fof(f12597,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11832,f12476])).
fof(f11832,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl109),
  inference(unit_resulting_resolution,[],[f11666,f11666,f497])).
fof(f12594,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl67 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11835,f12012])).
fof(f12591,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11838,f12012])).
fof(f12582,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11849,f12489])).
fof(f12489,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool))) ) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(condensation,[],[f12488])).
fof(f12488,plain,(
  ( ! [X6,X5] : (s(bool,X5) = s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,X5) = s(bool,X6)) ) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f12130,f12476])).
fof(f12130,plain,(
  ( ! [X6,X5] : (s(bool,X5) = s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,X5) = s(bool,X6)) ) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f12012,f5023])).
fof(f11849,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl30 | $spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f11666,f7298])).
fof(f12581,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11850,f12489])).
fof(f11850,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl15 | $spl30 | $spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f4197,f4197,f11666,f7397])).
fof(f12580,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11851,f12489])).
fof(f11851,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl30 | $spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f11666,f11666,f11666,f7397])).
fof(f12579,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl96 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11852,f12489])).
fof(f11852,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl30 | $spl33 | $spl96 | $spl109)),
  inference(unit_resulting_resolution,[],[f11308,f11308,f11666,f7397])).
fof(f12578,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11853,f12489])).
fof(f11853,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl30 | $spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f2948,f2948,f11666,f7397])).
fof(f12572,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11859,f12489])).
fof(f11859,plain,(
  s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl15 | $spl30 | $spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f4197,f11666,f11666,f7397])).
fof(f12571,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11860,f12489])).
fof(f11860,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl30 | $spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f11666,f11666,f11666,f7397])).
fof(f12570,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl96 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11861,f12489])).
fof(f11861,plain,(
  ( ! [X0,X1] : (s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X0,bool),X1)))) ) | ($spl30 | $spl33 | $spl96 | $spl109)),
  inference(unit_resulting_resolution,[],[f11308,f11666,f11666,f7397])).
fof(f12569,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11862,f12489])).
fof(f11862,plain,(
  s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | ($spl30 | $spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f2948,f11666,f11666,f7397])).
fof(f12542,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11882,f12012])).
fof(f11882,plain,(
  s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | ($spl33 | $spl36 | $spl109)),
  inference(unit_resulting_resolution,[],[f11666,f11666,f9644])).
fof(f12530,plain,(
  $false | ($spl15 | $spl30 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f11893,f12489])).
fof(f11893,plain,(
  s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl30 | $spl33 | $spl109)),
  inference(unit_resulting_resolution,[],[f11666,f11666,f10463])).
fof(f12505,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f12289,f12012])).
fof(f12289,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f12012,f11666])).
fof(f12504,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl96 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f12286,f12012])).
fof(f12286,plain,(
  ( ! [X95,X96] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X95,bool),X96)))) ) | ($spl15 | $spl33 | $spl36 | $spl96 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f12012,f11308])).
fof(f12484,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f12086,f12012])).
fof(f12086,plain,(
  s(bool,'T_0') != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f12012,f2948])).
fof(f12475,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(subsumption_resolution,[],[f12013,f111])).
fof(f12013,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl33 | $spl36 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f12012,f113])).
fof(f12473,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl67 | $spl109 | $spl114)),
  inference(trivial_inequality_removal,[],[f12113])).
fof(f12113,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ($spl15 | $spl33 | $spl36 | $spl67 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f12012,f3917])).
fof(f12471,plain,(
  $false | ($spl15 | $spl33 | $spl36 | $spl67 | $spl109 | $spl114)),
  inference(trivial_inequality_removal,[],[f12146])).
fof(f12146,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | ($spl15 | $spl33 | $spl36 | $spl67 | $spl109 | $spl114)),
  inference(backward_demodulation,[],[f12012,f6107])).
fof(f11660,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl100),
  inference(cnf_transformation,[],[f11660_D])).
fof(f11660_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f11661,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK9(X0,X1,bool,bool))) ) | $spl102),
  inference(cnf_transformation,[],[f11661_D])).
fof(f11661_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK9(X0,X1,bool,bool))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f11674,plain,(
  ( ! [X27] : (s(bool,X27) != s(bool,f(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) | $spl120),
  inference(cnf_transformation,[],[f11674_D])).
fof(f11674_D,plain,(
  ( ! [X27] : (s(bool,X27) != s(bool,f(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11657,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,'T_0') = s(bool,sK9(X28,X29,bool,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl33),
  inference(superposition,[],[f933,f4991])).
fof(f11673,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,'F_0'))))) | $spl118),
  inference(cnf_transformation,[],[f11673_D])).
fof(f11673_D,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,'F_0'))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f11667,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl110),
  inference(cnf_transformation,[],[f11667_D])).
fof(f11667_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f11656,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,'T_0') = s(bool,sK9(X25,X26,bool,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,g(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X24)) ) | $spl33),
  inference(superposition,[],[f933,f4991])).
fof(f11671,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,g(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23)))) ) | $spl116),
  inference(cnf_transformation,[],[f11671_D])).
fof(f11671_D,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,g(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f11666,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl109),
  inference(cnf_transformation,[],[f11666_D])).
fof(f11666_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f11655,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK9(X21,X22,bool,bool)) | s(bool,X23) = s(bool,g(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23)))) ) | $spl33),
  inference(superposition,[],[f933,f4991])).
fof(f11669,plain,(
  s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(bool,'F_0')))),bool)) | $spl112),
  inference(cnf_transformation,[],[f11669_D])).
fof(f11669_D,plain,(
  s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(bool,'F_0')))),bool)) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f11670,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X14)),bool))) ) | $spl114),
  inference(cnf_transformation,[],[f11670_D])).
fof(f11670_D,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X14)),bool))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f11654,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK9(X18,X19,bool,bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X17)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(bool,'F_0')))),bool))) ) | $spl30),
  inference(superposition,[],[f933,f2476])).
fof(f11653,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') = s(bool,sK9(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(bool,'F_0')))),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X14)),bool))) ) | $spl30),
  inference(superposition,[],[f933,f2476])).
fof(f11652,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK9(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f933,f190])).
fof(f11651,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK9(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f933,f190])).
fof(f11650,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK9(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f933,f190])).
fof(f11649,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK9(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f933,f190])).
fof(f11663,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl105),
  inference(cnf_transformation,[],[f11663_D])).
fof(f11663_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f11648,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK9(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f933,f114])).
fof(f11279,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X50,bool),X51))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49)))) ) | $spl92),
  inference(cnf_transformation,[],[f11279_D])).
fof(f11279_D,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X50,bool),X51))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f11398,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,X164) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X161),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X161,bool),X162))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X164))) | s(bool,'T_0') = s(bool,X163)) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f11187,f192])).
fof(f11187,plain,(
  ( ! [X163,X161,X164,X162] : (p(s(bool,X163)) | s(bool,X164) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X161),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X161,bool),X162))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X164))) | s(bool,'T_0') = s(bool,X163)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5158,f10846])).
fof(f11226,plain,(
  ( ! [X152,X151,X149,X150] : (s(bool,'T_0') = s(bool,X151) | s(bool,X152) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X149),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X149,bool),X150))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X152)))) ) | ($spl33 | $spl36)),
  inference(duplicate_literal_removal,[],[f11184])).
fof(f11184,plain,(
  ( ! [X152,X151,X149,X150] : (s(bool,'T_0') = s(bool,X151) | s(bool,X152) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X149),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X149,bool),X150))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X152))) | s(bool,'T_0') = s(bool,X151)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f4991,f10846])).
fof(f11308,plain,(
  ( ! [X95,X96] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X95,bool),X96)))) ) | $spl96),
  inference(cnf_transformation,[],[f11308_D])).
fof(f11308_D,plain,(
  ( ! [X95,X96] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X95,bool),X96)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f11353,plain,(
  ( ! [X383,X382,X381] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X382),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X382,bool),X383))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X381)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f11113,f5195])).
fof(f11113,plain,(
  ( ! [X383,X382,X381] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X382),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X382,bool),X383))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X381))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X381)),bool))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5195,f10846])).
fof(f11350,plain,(
  ( ! [X368,X367,X369] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X368),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X368,bool),X369))) | s(bool,'T_0') = s(bool,X367)) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f11109,f4267])).
fof(f11109,plain,(
  ( ! [X368,X367,X369] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X368),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X368,bool),X369))) | s(bool,'T_0') = s(bool,X367) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X367)),bool))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f4267,f10846])).
fof(f11239,plain,(
  ( ! [X156,X154,X157,X155] : (s(bool,X157) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X155),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X155,bool),X156))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X157))) | s(bool,'T_0') = s(bool,X154)) ) | ($spl33 | $spl36)),
  inference(duplicate_literal_removal,[],[f11078])).
fof(f11078,plain,(
  ( ! [X156,X154,X157,X155] : (s(bool,X157) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X155),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X155,bool),X156))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X157))) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f4991,f10846])).
fof(f11314,plain,(
  ( ! [X116,X117,X113] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X116),'const.option.NONE_0'),s(bool,X113),s(fun(X116,bool),X117)))) ) | $spl98),
  inference(cnf_transformation,[],[f11314_D])).
fof(f11314_D,plain,(
  ( ! [X116,X117,X113] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X116),'const.option.NONE_0'),s(bool,X113),s(fun(X116,bool),X117)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f11313,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X114),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X114,bool),X115))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X116),'const.option.NONE_0'),s(bool,X113),s(fun(X116,bool),X117)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f11067,f4437])).
fof(f11067,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X114),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X114,bool),X115))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X116),'const.option.NONE_0'),s(bool,X113),s(fun(X116,bool),X117))) | s(bool,'T_0') = s(bool,X113)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f4437,f10846])).
fof(f11307,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X95,bool),X96))) | s(bool,'T_0') = s(bool,X94)) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f11062,f4267])).
fof(f11062,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X95,bool),X96))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X94)),bool)) | s(bool,'T_0') = s(bool,X94)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f2954,f10846])).
fof(f11300,plain,(
  ( ! [X64,X65,X63] : (s(bool,X65) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X63,bool),X64))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65)))) ) | $spl94),
  inference(cnf_transformation,[],[f11300_D])).
fof(f11300_D,plain,(
  ( ! [X64,X65,X63] : (s(bool,X65) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X63,bool),X64))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f11245,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X65) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X63,bool),X64))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65))) | s(bool,'T_0') = s(bool,X62)) ) | ($spl33 | $spl36)),
  inference(duplicate_literal_removal,[],[f11054])).
fof(f11054,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X65) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X63,bool),X64))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X65))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f532,f10846])).
fof(f11253,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X47,bool),X48))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X50)))) ) | ($spl33 | $spl36)),
  inference(duplicate_literal_removal,[],[f11033])).
fof(f11033,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X50) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X47,bool),X48))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X50)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f10846,f4991])).
fof(f11259,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X49) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X50,bool),X51))) | s(bool,'T_0') = s(bool,X48) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49)))) ) | ($spl33 | $spl36)),
  inference(duplicate_literal_removal,[],[f10987])).
fof(f10987,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X49) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X50,bool),X51))) | s(bool,'T_0') = s(bool,X48) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49))) | s(bool,'T_0') = s(bool,X48)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f10846,f4991])).
fof(f10759,plain,(
  ( ! [X38,X40] : (s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X38)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40))) | s(bool,X38) = s(bool,X40)) ) | $spl90),
  inference(cnf_transformation,[],[f10759_D])).
fof(f10759_D,plain,(
  ( ! [X38,X40] : (s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X38)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40))) | s(bool,X38) = s(bool,X40)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f10740,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X38)),bool)) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40)))) ) | ($spl30 | $spl33)),
  inference(duplicate_literal_removal,[],[f10686])).
fof(f10686,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,sK7('const.option.SOME_1'(s(bool,X38)),bool)) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40)))) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f10463,f4991])).
fof(f9645,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) != s(bool,sK7('const.option.SOME_1'(s(bool,X34)),bool)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36)) ) | $spl88),
  inference(cnf_transformation,[],[f9645_D])).
fof(f9645_D,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) != s(bool,sK7('const.option.SOME_1'(s(bool,X34)),bool)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X36) | s(bool,X35) = s(bool,X36)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f9607,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X35) != s(bool,sK7('const.option.SOME_1'(s(bool,X34)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X33)),bool)) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) ) | ($spl33 | $spl36)),
  inference(duplicate_literal_removal,[],[f9580])).
fof(f9580,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X35) != s(bool,sK7('const.option.SOME_1'(s(bool,X34)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X33)),bool)) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X33)),bool)) | s(bool,'T_0') = s(bool,X34)) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f4347,f4652])).
fof(f8241,plain,(
  ( ! [X26,X25] : (s(bool,sK7('const.option.SOME_1'(s(bool,X25)),bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X26)),bool)) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X26)) ) | $spl86),
  inference(cnf_transformation,[],[f8241_D])).
fof(f8241_D,plain,(
  ( ! [X26,X25] : (s(bool,sK7('const.option.SOME_1'(s(bool,X25)),bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X26)),bool)) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8206,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X63)),bool)) | s(bool,sK7('const.option.SOME_1'(s(bool,X62)),bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X64)),bool)) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X62)) ) | $spl33),
  inference(duplicate_literal_removal,[],[f7919])).
fof(f7919,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X63)),bool)) | s(bool,sK7('const.option.SOME_1'(s(bool,X62)),bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X64)),bool)) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X63)),bool)) | s(bool,'T_0') = s(bool,X62)) ) | $spl33),
  inference(superposition,[],[f4652,f4652])).
fof(f8224,plain,(
  ( ! [X26,X24,X25] : (s(bool,sK7('const.option.SOME_1'(s(bool,X25)),bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X26)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X24)),bool)) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) ) | $spl33),
  inference(duplicate_literal_removal,[],[f7824])).
fof(f7824,plain,(
  ( ! [X26,X24,X25] : (s(bool,sK7('const.option.SOME_1'(s(bool,X25)),bool)) = s(bool,sK7('const.option.SOME_1'(s(bool,X26)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X24)),bool)) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X24)),bool)) | s(bool,'T_0') = s(bool,X25)) ) | $spl33),
  inference(superposition,[],[f4652,f4652])).
fof(f7680,plain,(
  ( ! [X33,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X33)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'F_0') = s(bool,sK8('F_0',X34,bool))) ) | $spl84),
  inference(cnf_transformation,[],[f7680_D])).
fof(f7680_D,plain,(
  ( ! [X33,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X33)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'F_0') = s(bool,sK8('F_0',X34,bool))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7662,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X33)))) | s(bool,'F_0') = s(bool,sK8('F_0',X34,bool)) | s(bool,'T_0') = s(bool,X32) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33)))) ) | $spl33),
  inference(duplicate_literal_removal,[],[f7633])).
fof(f7633,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X33)))) | s(bool,'F_0') = s(bool,sK8('F_0',X34,bool)) | s(bool,'T_0') = s(bool,X32) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X32)) ) | $spl33),
  inference(superposition,[],[f570,f4991])).
fof(f7452,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21)))) ) | $spl82),
  inference(cnf_transformation,[],[f7452_D])).
fof(f7452_D,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f7427,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X20) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21)))) ) | ($spl30 | $spl33)),
  inference(duplicate_literal_removal,[],[f7381])).
fof(f7381,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X20) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f7298,f4991])).
fof(f6312,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15)))) ) | $spl80),
  inference(cnf_transformation,[],[f6312_D])).
fof(f6312_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f6314,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18)))) ) | $spl33),
  inference(subsumption_resolution,[],[f6259,f155])).
fof(f6259,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,'T_0') = s(bool,X18) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18)))) ) | $spl33),
  inference(superposition,[],[f532,f4881])).
fof(f6313,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X16)) ) | $spl33),
  inference(subsumption_resolution,[],[f6258,f4267])).
fof(f6258,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X16)),bool))) ) | $spl33),
  inference(superposition,[],[f532,f4452])).
fof(f6302,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | $spl33),
  inference(duplicate_literal_removal,[],[f6257])).
fof(f6257,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl33),
  inference(superposition,[],[f532,f4280])).
fof(f5193,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) | $spl78),
  inference(cnf_transformation,[],[f5193_D])).
fof(f5193_D,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5192,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,X13) | ~p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) | ($spl18 | $spl33)),
  inference(subsumption_resolution,[],[f5191,f220])).
fof(f5191,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,X13) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) | $spl33),
  inference(forward_demodulation,[],[f5094,f946])).
fof(f5094,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12)))) ) | $spl33),
  inference(superposition,[],[f162,f4881])).
fof(f4250,plain,(
  ( ! [X39,X38] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X38,sK0)) | s(bool,'F_0') = s(bool,X39)) ) | ($spl15 | $spl26 | $spl60)),
  inference(subsumption_resolution,[],[f4233,f1200])).
fof(f4233,plain,(
  ( ! [X39,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X38,sK0)) | s(bool,'F_0') = s(bool,X39)) ) | ($spl15 | $spl26 | $spl60)),
  inference(backward_demodulation,[],[f4202,f2154])).
fof(f4202,plain,(
  ( ! [X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29)))) ) | ($spl15 | $spl60)),
  inference(subsumption_resolution,[],[f4200,f4197])).
fof(f4200,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,'F_0') | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29)))) ) | $spl60),
  inference(backward_demodulation,[],[f3578,f2178])).
fof(f4248,plain,(
  ( ! [X138,X137] : (s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X138,sK0)) | s(bool,'T_0') = s(bool,sK7(X137,bool))) ) | ($spl0 | $spl15 | $spl60)),
  inference(subsumption_resolution,[],[f4231,f611])).
fof(f4231,plain,(
  ( ! [X138,X137] : (s('type.option.option'(bool),X137) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X138,sK0)) | s(bool,'T_0') = s(bool,sK7(X137,bool))) ) | ($spl0 | $spl15 | $spl60)),
  inference(backward_demodulation,[],[f4202,f1413])).
fof(f4247,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0))) ) | ($spl0 | $spl15 | $spl60)),
  inference(condensation,[],[f4246])).
fof(f4246,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) = s(bool,X75) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X74,sK0)) | s(bool,'T_0') = s(bool,X73)) ) | ($spl0 | $spl15 | $spl60)),
  inference(subsumption_resolution,[],[f4230,f4202])).
fof(f4230,plain,(
  ( ! [X74,X75,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X75))) | s(bool,X73) = s(bool,X75) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X74,sK0)) | s(bool,'T_0') = s(bool,X73)) ) | ($spl0 | $spl15 | $spl60)),
  inference(backward_demodulation,[],[f4202,f1397])).
fof(f4244,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(bool,X32),s(fun(X30,bool),X33)))) ) | ($spl15 | $spl60)),
  inference(subsumption_resolution,[],[f4229,f1200])).
fof(f4229,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(bool,X32),s(fun(X30,bool),X33)))) ) | ($spl15 | $spl60)),
  inference(backward_demodulation,[],[f4202,f1981])).
fof(f4242,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X35,bool),bool),happ(s(fun(bool,fun(fun(X35,bool),bool)),happ(s(fun('type.option.option'(X35),fun(bool,fun(fun(X35,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X35),X36))),s(bool,X37))),s(fun(X35,bool),X38)))) ) | ($spl15 | $spl60)),
  inference(subsumption_resolution,[],[f4227,f1200])).
fof(f4227,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X35,bool),bool),happ(s(fun(bool,fun(fun(X35,bool),bool)),happ(s(fun('type.option.option'(X35),fun(bool,fun(fun(X35,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X35),X36))),s(bool,X37))),s(fun(X35,bool),X38)))) ) | ($spl15 | $spl60)),
  inference(backward_demodulation,[],[f4202,f1982])).
fof(f4198,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X40)),bool))) ) | $spl68),
  inference(subsumption_resolution,[],[f1983,f3939])).
fof(f4195,plain,(
  ( ! [X70,X72] : (s(bool,X72) = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X70,sK0))) ) | $spl70),
  inference(cnf_transformation,[],[f4195_D])).
fof(f4195_D,plain,(
  ( ! [X70,X72] : (s(bool,X72) = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X70,sK0))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4194,plain,(
  ( ! [X70,X72,X71] : (s(bool,X72) = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X70,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X71)),bool))) ) | ($spl0 | $spl68)),
  inference(subsumption_resolution,[],[f2393,f3939])).
fof(f4171,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f111,f4087])).
fof(f4087,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl12),
  inference(superposition,[],[f113,f198])).
fof(f3939,plain,(
  ( ! [X103,X101] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool))) ) | $spl68),
  inference(cnf_transformation,[],[f3939_D])).
fof(f3939_D,plain,(
  ( ! [X103,X101] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3938,plain,(
  ( ! [X103,X101,X102] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X102,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f3937,f198])).
fof(f3937,plain,(
  ( ! [X103,X101,X102] : (s(bool,'F_0') != s(bool,X103) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X102,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X101)),bool))) ) | ($spl0 | $spl12)),
  inference(forward_demodulation,[],[f2349,f198])).
fof(f3917,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl67),
  inference(cnf_transformation,[],[f3917_D])).
fof(f3917_D,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f2260,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))) )),
  inference(condensation,[],[f2259])).
fof(f2259,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f2258])).
fof(f2258,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f2240])).
fof(f2240,plain,(
  ( ! [X6,X8,X7] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f497,f190])).
fof(f2263,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))) )),
  inference(condensation,[],[f2262])).
fof(f2262,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f2261])).
fof(f2261,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f2239])).
fof(f2239,plain,(
  ( ! [X4,X5,X3] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f497,f190])).
fof(f3909,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f3855,f111])).
fof(f3855,plain,(
  ~p(s(bool,'T_0')) | $spl32),
  inference(backward_demodulation,[],[f3854,f113])).
fof(f3854,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl32),
  inference(subsumption_resolution,[],[f2260,f3838])).
fof(f3838,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f3837,f3099])).
fof(f3837,plain,(
  ( ! [X21] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21))) | s(bool,'T_0') != s(bool,X21)) ) | $spl32),
  inference(forward_demodulation,[],[f3642,f946])).
fof(f3642,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))))) ) | $spl32),
  inference(superposition,[],[f155,f2949])).
fof(f3836,plain,(
  $false | ($spl15 | $spl32)),
  inference(subsumption_resolution,[],[f3835,f151])).
fof(f3835,plain,(
  s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | ($spl15 | $spl32)),
  inference(subsumption_resolution,[],[f3638,f490])).
fof(f3638,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl32),
  inference(superposition,[],[f146,f2949])).
fof(f3664,plain,(
  $false | ($spl15 | $spl32)),
  inference(subsumption_resolution,[],[f3615,f200])).
fof(f3615,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl32),
  inference(superposition,[],[f2949,f946])).
fof(f3663,plain,(
  $false | ($spl15 | $spl32)),
  inference(subsumption_resolution,[],[f3662,f490])).
fof(f3662,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl32),
  inference(forward_demodulation,[],[f3603,f946])).
fof(f3603,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | $spl32),
  inference(unit_resulting_resolution,[],[f2949,f155])).
fof(f3661,plain,(
  $false | ($spl15 | $spl32)),
  inference(subsumption_resolution,[],[f3660,f490])).
fof(f3660,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl32),
  inference(forward_demodulation,[],[f3605,f946])).
fof(f3605,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)))) | $spl32),
  inference(unit_resulting_resolution,[],[f2949,f155])).
fof(f3586,plain,(
  ( ! [X19,X18] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,X19) != s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))) ) | $spl64),
  inference(cnf_transformation,[],[f3586_D])).
fof(f3586_D,plain,(
  ( ! [X19,X18] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s(bool,X19) != s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3091,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) != s(bool,happ(s(fun(sK0,bool),X15),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X14) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X15,sK0))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f3057])).
fof(f3057,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) != s(bool,happ(s(fun(sK0,bool),X15),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X14) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X15,sK0)) | s(bool,'T_0') = s(bool,X14)) ) | $spl0),
  inference(superposition,[],[f532,f928])).
fof(f3585,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X17)),bool)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | s(bool,X19) != s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19)))) ) | ($spl0 | $spl30)),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f170,f813,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f1472,f1467,f1460,f1479,f1481,f1432,f1485,f1433,f1486,f588,f1045,f1351,f922,f1350,f1448,f1438,f1436,f1416,f1413,f1399,f1397,f1368,f1367,f1365,f1363,f1349,f1348,f1345,f1271,f1116,f928,f921,f869,f621,f900,f1402,f571,f1441,f1440,f1415,f1401,f898,f895,f893,f673,f671,f576,f573,f1070,f1048,f788,f760,f781,f1060,f1346,f1143,f1135,f1119,f912,f735,f733,f683,f681,f585,f568,f158,f157,f156,f1347,f1062,f1061,f1032,f1018,f1016,f934,f936,f1409,f1335,f1331,f1308,f1285,f1284,f1283,f1253,f1204,f1203,f1201,f1200,f1199,f1197,f1196,f1195,f1176,f1093,f1020,f965,f938,f933,f897,f892,f737,f675,f610,f575,f570,f537,f532,f497,f492,f239,f113,f1980,f1981,f1982,f1983,f1984,f1995,f1996,f1997,f1998,f1999,f2131,f2161,f2167,f2175,f2176,f2177,f2178,f2179,f2183,f2189,f2191,f2193,f2194,f2205,f2206,f2212,f2214,f2216,f2217,f2248,f2249,f2250,f2251,f2252,f2280,f2281,f2282,f2283,f2466,f2477,f2478,f2479,f2347,f2349,f2359,f2363,f2366,f2393,f2395,f2408,f2410,f2515,f2516,f2517,f2440,f2518,f2520,f2521,f2476,f2918,f2919,f2920,f2656,f2672,f2674,f2684,f2688,f2691,f2718,f3018,f3020,f2776,f3021,f3026,f2822,f3039,f3067,f3068,f3069,f3070,f3099,f3186,f3229,f3230,f3231,f3242,f3248,f3250,f3252,f3253,f3267,f3274,f3276,f3278,f3279,f3297,f3298,f3299,f3300,f3337,f3570,f3325,f3575,f3327,f3223,f3577,f3290,f3188,f3102,f3581,f3058])).
fof(f3039,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') != s(bool,X0)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f3038])).
fof(f3038,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool))) ) | $spl30),
  inference(forward_demodulation,[],[f2828,f946])).
fof(f2828,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool))),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool))) ) | $spl30),
  inference(equality_factoring,[],[f2476])).
fof(f2822,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK7('const.option.SOME_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool))) ) | $spl30),
  inference(equality_factoring,[],[f2476])).
fof(f3026,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') = s(bool,X83) | s(bool,X81) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X81)),bool),X82,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'T_0')))) | s(bool,X81) = s(bool,X83)) ) | $spl30),
  inference(forward_demodulation,[],[f3025,f946])).
fof(f3025,plain,(
  ( ! [X83,X81,X82] : (s(bool,X81) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X81)),bool),X82,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'T_0')))) | s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X83)),bool))) ) | $spl30),
  inference(forward_demodulation,[],[f2781,f946])).
fof(f2781,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'T_0')))) | s(bool,sK7('const.option.SOME_1'(s(bool,X81)),bool)) = s(bool,sK8(sK7('const.option.SOME_1'(s(bool,X81)),bool),X82,bool)) | s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X83)),bool))) ) | $spl30),
  inference(superposition,[],[f157,f2476])).
fof(f3021,plain,(
  ( ! [X72,X71,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X71))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X72))) | s(bool,'T_0') != s(bool,X72) | s(bool,X71) = s(bool,X73) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X73)),bool))) ) | $spl30),
  inference(forward_demodulation,[],[f2778,f946])).
fof(f2778,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X72) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X72))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK7('const.option.SOME_1'(s(bool,X71)),bool)))) | s(bool,X71) = s(bool,X73) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X73)),bool))) ) | $spl30),
  inference(superposition,[],[f155,f2476])).
fof(f2776,plain,(
  ( ! [X66,X67,X65] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X66))) | s(bool,X66) = s(bool,sK7('const.option.SOME_1'(s(bool,X65)),bool)) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X67)),bool))) ) | $spl30),
  inference(superposition,[],[f154,f2476])).
fof(f3020,plain,(
  ( ! [X61,X62] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X61))) | s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X62)),bool))) ) | $spl30),
  inference(subsumption_resolution,[],[f2774,f151])).
fof(f2774,plain,(
  ( ! [X61,X62] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X61))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X61))) | s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X62)),bool))) ) | $spl30),
  inference(superposition,[],[f146,f2476])).
fof(f3018,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) = s(bool,sK7('const.option.SOME_1'(s(bool,X46)),bool)) | s(bool,'F_0') = s(bool,X47) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X48)),bool))) ) | $spl30),
  inference(subsumption_resolution,[],[f2769,f111])).
fof(f2769,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'T_0')) | s(bool,X47) = s(bool,sK7('const.option.SOME_1'(s(bool,X46)),bool)) | s(bool,'F_0') = s(bool,X47) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X48)),bool))) ) | $spl30),
  inference(superposition,[],[f1199,f2476])).
fof(f2691,plain,(
  ( ! [X189,X187,X188,X186] : (s(X188,f(s(bool,X189))) != s(X188,g(s(bool,X189))) | s(fun(bool,X188),X186) = s(fun(bool,X188),X187) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK9(X186,X187,X188,bool))),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X189)),bool))) ) | $spl30),
  inference(superposition,[],[f164,f2476])).
fof(f2688,plain,(
  ( ! [X180,X179] : (s('type.option.option'(bool),X179) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X180))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X179) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,sK7(X179,bool))),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X180)),bool))) ) | $spl30),
  inference(superposition,[],[f146,f2476])).
fof(f2684,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,X164) != s(bool,f(s(X161,sK9(X162,X163,bool,X161)))) | s(fun(X161,bool),X162) = s(fun(X161,bool),X163) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,g(s(X161,sK9(X162,X163,bool,X161))))),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X164)),bool))) ) | $spl30),
  inference(superposition,[],[f164,f2476])).
fof(f2674,plain,(
  ( ! [X123,X121,X122] : (s(bool,X122) != s(bool,X123) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X121))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X121)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X122)),bool))) ) | $spl30),
  inference(superposition,[],[f155,f2476])).
fof(f2672,plain,(
  ( ! [X116,X117,X115] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X116))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X117))) | s(bool,X115) = s(bool,X117) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X115)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X116)),bool))) ) | $spl30),
  inference(superposition,[],[f154,f2476])).
fof(f2920,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X64),X65),s(bool,X66),s(fun(X64,bool),X67))) | s(bool,X68) = s(bool,happ(s(fun(fun(X64,bool),bool),happ(s(fun(bool,fun(fun(X64,bool),bool)),happ(s(fun('type.option.option'(X64),fun(bool,fun(fun(X64,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X64),X65))),s(bool,X66))),s(fun(X64,bool),X67))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X68)),bool))) ) | $spl30),
  inference(forward_demodulation,[],[f2589,f946])).
fof(f2589,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X64),X65),s(bool,X66),s(fun(X64,bool),X67)))),bool)) | s(bool,X68) = s(bool,happ(s(fun(fun(X64,bool),bool),happ(s(fun(bool,fun(fun(X64,bool),bool)),happ(s(fun('type.option.option'(X64),fun(bool,fun(fun(X64,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X64),X65))),s(bool,X66))),s(fun(X64,bool),X67))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X68)),bool))) ) | $spl30),
  inference(superposition,[],[f2476,f166])).
fof(f2919,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),X62),s(X59,X60))) | s(bool,X63) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X59),'const.option.SOME_1'(s(X59,X60))),s(bool,X61),s(fun(X59,bool),X62))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X63)),bool))) ) | $spl30),
  inference(forward_demodulation,[],[f2588,f946])).
fof(f2588,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,happ(s(fun(X59,bool),X62),s(X59,X60)))),bool)) | s(bool,X63) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X59),'const.option.SOME_1'(s(X59,X60))),s(bool,X61),s(fun(X59,bool),X62))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X63)),bool))) ) | $spl30),
  inference(superposition,[],[f2476,f165])).
fof(f2918,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,X56) | s(bool,X58) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X55),'const.option.NONE_0'),s(bool,X56),s(fun(X55,bool),X57))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X58)),bool))) ) | $spl30),
  inference(forward_demodulation,[],[f2587,f946])).
fof(f2587,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X56)),bool)) | s(bool,X58) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X55),'const.option.NONE_0'),s(bool,X56),s(fun(X55,bool),X57))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X58)),bool))) ) | $spl30),
  inference(superposition,[],[f2476,f163])).
fof(f3579,plain,(
  ( ! [X12] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X12,sK0))) ) | $spl62),
  inference(cnf_transformation,[],[f3579_D])).
fof(f3579_D,plain,(
  ( ! [X12] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X12,sK0))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3581,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X0)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0))) ) | $spl0),
  inference(condensation,[],[f3580])).
fof(f3580,plain,(
  ( ! [X33,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X32),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X34) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X32,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X33)),bool))) ) | $spl0),
  inference(subsumption_resolution,[],[f3082,f1091])).
fof(f3082,plain,(
  ( ! [X33,X34,X32] : (s(bool,X33) != s(bool,X34) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X32),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X34) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X32,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X33)),bool))) ) | $spl0),
  inference(superposition,[],[f532,f1348])).
fof(f3102,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0))) ) | $spl0),
  inference(condensation,[],[f3101])).
fof(f3101,plain,(
  ( ! [X37,X35,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X35),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X37) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X35,sK0)) | s(bool,'T_0') = s(bool,X36)) ) | $spl0),
  inference(subsumption_resolution,[],[f3083,f190])).
fof(f3083,plain,(
  ( ! [X37,X35,X36] : (s(bool,X36) != s(bool,X37) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X35),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X37) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X35,sK0)) | s(bool,'T_0') = s(bool,X36)) ) | $spl0),
  inference(superposition,[],[f532,f928])).
fof(f3290,plain,(
  ( ! [X10,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X10),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X9) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X10,sK0))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f3222])).
fof(f3222,plain,(
  ( ! [X10,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X10),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X9) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X10,sK0)) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f3099,f928])).
fof(f3577,plain,(
  ( ! [X12,X11] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X11)),bool)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X12,sK0))) ) | ($spl0 | $spl30)),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f170,f813,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f1472,f1467,f1460,f1479,f1481,f1432,f1485,f1433,f1486,f588,f1045,f1351,f922,f1350,f1448,f1438,f1436,f1416,f1413,f1399,f1397,f1368,f1367,f1365,f1363,f1349,f1348,f1345,f1271,f1116,f928,f921,f869,f621,f900,f1402,f571,f1441,f1440,f1415,f1401,f898,f895,f893,f673,f671,f576,f573,f1070,f1048,f788,f760,f781,f1060,f1346,f1143,f1135,f1119,f912,f735,f733,f683,f681,f585,f568,f158,f157,f156,f1347,f1062,f1061,f1032,f1018,f1016,f934,f936,f1409,f1335,f1331,f1308,f1285,f1284,f1283,f1253,f1204,f1203,f1201,f1200,f1199,f1197,f1196,f1195,f1176,f1093,f1020,f965,f938,f933,f897,f892,f737,f675,f610,f575,f570,f537,f532,f497,f492,f239,f113,f1980,f1981,f1982,f1983,f1984,f1995,f1996,f1997,f1998,f1999,f2131,f2161,f2167,f2175,f2176,f2177,f2178,f2179,f2183,f2189,f2191,f2193,f2194,f2205,f2206,f2212,f2214,f2216,f2217,f2248,f2249,f2250,f2251,f2252,f2280,f2281,f2282,f2283,f2466,f2477,f2478,f2479,f2347,f2349,f2359,f2363,f2366,f2393,f2395,f2408,f2410,f2515,f2516,f2517,f2440,f2518,f2520,f2521,f2476,f2918,f2919,f2920,f2656,f2672,f2674,f2684,f2688,f2691,f2718,f3018,f3020,f2776,f3021,f3026,f2822,f3039,f3067,f3068,f3069,f3070,f3099,f3186,f3229,f3230,f3231,f3242,f3248,f3250,f3252,f3253,f3267,f3274,f3276,f3278,f3279,f3297,f3298,f3299,f3300,f3337,f3570,f3325,f3575,f3327,f3223])).
fof(f3576,plain,(
  ( ! [X19,X18] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | s(bool,'F_0') = s(bool,sK8('F_0',X19,bool))) ) | $spl58),
  inference(cnf_transformation,[],[f3576_D])).
fof(f3576_D,plain,(
  ( ! [X19,X18] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | s(bool,'F_0') = s(bool,sK8('F_0',X19,bool))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3327,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,happ(s(fun(sK0,bool),X15),s(sK0,sK2)))))) | s(bool,'F_0') = s(bool,sK8('F_0',X16,bool)) | s(bool,'T_0') = s(bool,X14) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X15,sK0))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f3306])).
fof(f3306,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,happ(s(fun(sK0,bool),X15),s(sK0,sK2)))))) | s(bool,'F_0') = s(bool,sK8('F_0',X16,bool)) | s(bool,'T_0') = s(bool,X14) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X15,sK0)) | s(bool,'T_0') = s(bool,X14)) ) | $spl0),
  inference(superposition,[],[f570,f928])).
fof(f3575,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X17) | ~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))))) | s(bool,'F_0') = s(bool,sK8('F_0',X19,bool)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f3574])).
fof(f3574,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X17) | ~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))))) | s(bool,'F_0') = s(bool,sK8('F_0',X19,bool)) | s(bool,'T_0') = s(bool,X17) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f3307,f946])).
fof(f3307,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))))) | s(bool,'F_0') = s(bool,sK8('F_0',X19,bool)) | s(bool,'T_0') = s(bool,X17) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X18,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X17)),bool))) ) | $spl0),
  inference(superposition,[],[f570,f1348])).
fof(f3572,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,sK8('F_0',X17,bool))) ) | $spl56),
  inference(cnf_transformation,[],[f3572_D])).
fof(f3572_D,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,sK8('F_0',X17,bool))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f3571,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X19,sK0))) ) | $spl54),
  inference(cnf_transformation,[],[f3571_D])).
fof(f3571_D,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X19,sK0))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3325,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,happ(s(fun(sK0,bool),X16),s(sK0,sK2)))) | s(bool,'F_0') = s(bool,sK8('F_0',X14,bool)) | s(bool,'T_0') = s(bool,X15) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X16,sK0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) ) | $spl0),
  inference(superposition,[],[f570,f928])).
fof(f3570,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | ~p(s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK2)))) | s(bool,'F_0') = s(bool,sK8('F_0',X17,bool)) | s(bool,'T_0') = s(bool,X18) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X19,sK0))) ) | $spl0),
  inference(forward_demodulation,[],[f3326,f946])).
fof(f3326,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK2)))) | s(bool,'F_0') = s(bool,sK8('F_0',X17,bool)) | s(bool,'T_0') = s(bool,X18) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X19,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))),bool))) ) | $spl0),
  inference(superposition,[],[f570,f1348])).
fof(f3540,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0))) ) | $spl38),
  inference(cnf_transformation,[],[f3540_D])).
fof(f3540_D,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f3568,plain,(
  ( ! [X109,X110,X108] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X109))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X108),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X110,sK0),X108,sK0))) ) | $spl28),
  inference(subsumption_resolution,[],[f3491,f155])).
fof(f3491,plain,(
  ( ! [X109,X110,X108] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X109))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X108),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X110,sK0),X108,sK0)) | s(bool,'T_0') = s(bool,X109)) ) | $spl28),
  inference(superposition,[],[f1200,f3206])).
fof(f3206,plain,(
  ( ! [X10,X8,X9] : (s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X9,sK0),X8,sK0)) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X10)) ) | $spl28),
  inference(subsumption_resolution,[],[f3200,f111])).
fof(f3200,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X9,sK0),X8,sK0)) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X10)) ) | $spl28),
  inference(superposition,[],[f3130,f190])).
fof(f3130,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X17,sK0),X18,sK0))) ) | $spl28),
  inference(forward_demodulation,[],[f3117,f2475])).
fof(f3117,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK2)))) | s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X17,sK0)) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X17,sK0),X18,sK0))) ) | $spl28),
  inference(superposition,[],[f157,f2475])).
fof(f3567,plain,(
  ( ! [X103,X101,X102] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X101),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X102) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X103,sK0),X101,sK0))) ) | $spl28),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f588,f621,f900,f571,f898,f895,f893,f673,f671,f576,f573,f1070,f1048,f788,f760,f781,f1060,f1143,f1135,f1119,f912,f735,f733,f683,f681,f585,f568,f158,f157,f156,f1062,f1061,f1032,f1018,f1016,f934,f936,f1335,f1331,f1308,f1285,f1284,f1283,f1253,f1204,f1203,f1201,f1200,f1199,f1197,f1196,f1195,f1176,f1093,f1020,f965,f938,f933,f897,f892,f737,f675,f610,f575,f570,f537,f532,f497,f492,f239,f113,f1980,f1981,f1982,f1983,f1984,f1995,f1996,f1997,f1998,f1999,f2161,f2167,f2175,f2176,f2177,f2178,f2179,f2183,f2189,f2191,f2193,f2194,f2205,f2206,f2212,f2214,f2216,f2217,f2248,f2249,f2250,f2251,f2252,f2475,f3067,f3068,f3069,f3070,f3099,f3112,f3130,f3186,f3193,f3194,f3195,f3196,f3206,f3229,f3230,f3231,f3242,f3248,f3250,f3252,f3253,f3267,f3274,f3276,f3278,f3279,f3297,f3298,f3299,f3300,f3337,f3355,f3539,f3533,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3393,f3395,f3397,f3398,f3525,f3519,f3518,f3517,f3515,f3443,f3445,f3447,f3448,f3455,f3459,f3461,f3462,f3552,f3554,f3480,f3558,f3560,f3564,f3566,f3489])).
fof(f3489,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X102)),bool)) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X101),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X103,sK0),X101,sK0)) | s(bool,'T_0') = s(bool,X102)) ) | $spl28),
  inference(superposition,[],[f1093,f3206])).
fof(f3566,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0))) ) | $spl28),
  inference(condensation,[],[f3565])).
fof(f3565,plain,(
  ( ! [X92,X90,X91,X89] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X91))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X89),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X91) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X92,sK0),X89,sK0)) | s(bool,'T_0') = s(bool,X90)) ) | $spl28),
  inference(subsumption_resolution,[],[f3486,f114])).
fof(f3486,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') != s(bool,X90) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X91))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X89),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X91) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X92,sK0),X89,sK0)) | s(bool,'T_0') = s(bool,X90)) ) | $spl28),
  inference(superposition,[],[f537,f3206])).
fof(f3480,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X66),'const.option.NONE_0'),s(bool,X65),s(fun(X66,bool),X67))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X64),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X68,sK0),X64,sK0)) | s(bool,'T_0') = s(bool,X65)) ) | $spl28),
  inference(superposition,[],[f239,f3206])).
fof(f3462,plain,(
  ( ! [X255,X253,X257,X254,X256] : (s(X255,f(s(bool,happ(s(fun(sK0,bool),X256),s(sK0,sK2))))) != s(X255,g(s(bool,happ(s(fun(sK0,bool),X256),s(sK0,sK2))))) | s(fun(bool,X255),X253) = s(fun(bool,X255),X254) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X257,sK0),X256,sK0)) | s(bool,'T_0') = s(bool,sK9(X253,X254,X255,bool))) ) | $spl28),
  inference(superposition,[],[f164,f3206])).
fof(f3461,plain,(
  ( ! [X251,X249,X252,X250] : (p(s(bool,happ(s(fun(bool,bool),X250),s(bool,happ(s(fun(sK0,bool),X251),s(sK0,sK2)))))) | ~p(s(bool,happ(s(fun(bool,bool),X250),s(bool,X249)))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X252,sK0),X251,sK0)) | s(bool,'T_0') = s(bool,sK8(X249,X250,bool))) ) | $spl28),
  inference(superposition,[],[f158,f3206])).
fof(f3459,plain,(
  ( ! [X243,X244,X245] : (s('type.option.option'(bool),X243) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X244),s(sK0,sK2))))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X243) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X245,sK0),X244,sK0)) | s(bool,'T_0') = s(bool,sK7(X243,bool))) ) | $spl28),
  inference(superposition,[],[f146,f3206])).
fof(f3455,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s(bool,f(s(X221,sK9(X222,X223,bool,X221)))) != s(bool,happ(s(fun(sK0,bool),X224),s(sK0,sK2))) | s(fun(X221,bool),X222) = s(fun(X221,bool),X223) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X225,sK0),X224,sK0)) | s(bool,'T_0') = s(bool,g(s(X221,sK9(X222,X223,bool,X221))))) ) | $spl28),
  inference(superposition,[],[f164,f3206])).
fof(f3448,plain,(
  ( ! [X187,X185,X186,X184] : (~p(s(bool,happ(s(fun(bool,bool),X186),s(bool,happ(s(fun(sK0,bool),X185),s(sK0,sK2)))))) | s(bool,X184) = s(bool,sK8(X184,X186,bool)) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X187,sK0),X185,sK0)) | s(bool,'T_0') = s(bool,X184)) ) | $spl28),
  inference(superposition,[],[f157,f3206])).
fof(f3447,plain,(
  ( ! [X182,X180,X183,X181,X179] : (~p(s(bool,happ(s(fun(bool,bool),X181),s(bool,happ(s(fun(sK0,bool),X180),s(sK0,sK2)))))) | p(s(bool,happ(s(fun(bool,bool),X181),s(bool,X182)))) | s(bool,X179) != s(bool,X182) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X183,sK0),X180,sK0)) | s(bool,'T_0') = s(bool,X179)) ) | $spl28),
  inference(superposition,[],[f156,f3206])).
fof(f3445,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,X173) != s(bool,happ(s(fun(sK0,bool),X172),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X171))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X173))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X174,sK0),X172,sK0)) | s(bool,'T_0') = s(bool,X171)) ) | $spl28),
  inference(superposition,[],[f155,f3206])).
fof(f3443,plain,(
  ( ! [X165,X163,X166,X164] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X165))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X164),s(sK0,sK2))))) | s(bool,X163) = s(bool,X165) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X166,sK0),X164,sK0)) | s(bool,'T_0') = s(bool,X163)) ) | $spl28),
  inference(superposition,[],[f154,f3206])).
fof(f3398,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,happ(s(fun(sK0,bool),X36),s(sK0,sK2)))) | s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X34,sK0),X35,sK0)) = s(sK0,sK8(sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X34,sK0),X35,sK0),X36,sK0)) | s(bool,X37) = s(bool,happ(s(fun(sK0,bool),X35),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X37)) ) | $spl28),
  inference(superposition,[],[f157,f3206])).
fof(f3397,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK2)))) | p(s(bool,happ(s(fun(sK0,bool),X31),s(sK0,X32)))) | s(sK0,X32) != s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X29,sK0),X30,sK0)) | s(bool,X33) = s(bool,happ(s(fun(sK0,bool),X30),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X33)) ) | $spl28),
  inference(superposition,[],[f156,f3206])).
fof(f3395,plain,(
  ( ! [X24,X23,X21,X22] : (s(sK0,sK2) != s(sK0,X23) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X23))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X21,sK0),X22,sK0)))) | s(bool,X24) = s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X24)) ) | $spl28),
  inference(superposition,[],[f155,f3206])).
fof(f3393,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X15))) | s(sK0,X15) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X13,sK0),X14,sK0)) | s(bool,X16) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X16)) ) | $spl28),
  inference(superposition,[],[f154,f3206])).
fof(f3376,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,X15) = s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X11),X12),s(fun(sK0,bool),X13),s(fun(X11,fun(sK0,bool)),X14))),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X16,sK0),happ(s(fun(fun(X11,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X11,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X11),fun(fun(sK0,bool),fun(fun(X11,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X11),X12))),s(fun(sK0,bool),X13))),s(fun(X11,fun(sK0,bool)),X14)),sK0)) | s(bool,'T_0') = s(bool,X15)) ) | $spl28),
  inference(superposition,[],[f3206,f166])).
fof(f3375,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,X9) = s(bool,happ(s(fun(sK0,bool),happ(s(fun(X5,fun(sK0,bool)),X8),s(X5,X6))),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X10,sK0),'const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.SOME_1'(s(X5,X6))),s(fun(sK0,bool),X7),s(fun(X5,fun(sK0,bool)),X8)),sK0)) | s(bool,'T_0') = s(bool,X9)) ) | $spl28),
  inference(superposition,[],[f3206,f165])).
fof(f3374,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X3) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X4,sK0),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0)) | s(bool,'T_0') = s(bool,X3)) ) | $spl28),
  inference(superposition,[],[f3206,f163])).
fof(f3373,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,happ(s(fun(sK0,bool),X81),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X82,sK0),X81,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X80)),bool))) ) | $spl28),
  inference(superposition,[],[f3206,f946])).
fof(f3372,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (s(bool,happ(s(fun(sK0,bool),X78),s(sK0,sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X74),X75),s(bool,X76),s(fun(X74,bool),X77))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X79,sK0),X78,sK0)) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X74,bool),bool),happ(s(fun(bool,fun(fun(X74,bool),bool)),happ(s(fun('type.option.option'(X74),fun(bool,fun(fun(X74,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X74),X75))),s(bool,X76))),s(fun(X74,bool),X77)))) ) | $spl28),
  inference(superposition,[],[f3206,f166])).
fof(f3371,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (s(bool,happ(s(fun(sK0,bool),X72),s(sK0,sK2))) = s(bool,happ(s(fun(X68,bool),X71),s(X68,X69))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X73,sK0),X72,sK0)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X68),'const.option.SOME_1'(s(X68,X69))),s(bool,X70),s(fun(X68,bool),X71)))) ) | $spl28),
  inference(superposition,[],[f3206,f165])).
fof(f3370,plain,(
  ( ! [X66,X64,X67,X65,X63] : (s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X67,sK0),X66,sK0)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.NONE_0'),s(bool,X64),s(fun(X63,bool),X65)))) ) | $spl28),
  inference(superposition,[],[f3206,f163])).
fof(f3355,plain,(
  ( ! [X10,X8,X9] : (s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X9,sK0),sK7('const.option.SOME_1'(s(fun(sK0,bool),X8)),fun(sK0,bool)),sK0)) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X10)) ) | $spl28),
  inference(subsumption_resolution,[],[f3345,f111])).
fof(f3345,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X9,sK0),sK7('const.option.SOME_1'(s(fun(sK0,bool),X8)),fun(sK0,bool)),sK0)) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X10)) ) | $spl28),
  inference(superposition,[],[f3196,f190])).
fof(f3196,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X15,sK0),sK7('const.option.SOME_1'(s(fun(sK0,bool),X14)),fun(sK0,bool)),sK0))) ) | $spl28),
  inference(superposition,[],[f3130,f946])).
fof(f3195,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(sK0,bool),X11),s(fun(X9,fun(sK0,bool)),X12))),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X13,sK0),happ(s(fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X9),fun(fun(sK0,bool),fun(fun(X9,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(sK0,bool),X11))),s(fun(X9,fun(sK0,bool)),X12)),sK0))) ) | $spl28),
  inference(superposition,[],[f3130,f166])).
fof(f3194,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(X4,fun(sK0,bool)),X7),s(X4,X5))),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X8,sK0),'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(sK0,bool),X6),s(fun(X4,fun(sK0,bool)),X7)),sK0))) ) | $spl28),
  inference(superposition,[],[f3130,f165])).
fof(f3193,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X3,sK0),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0))) ) | $spl28),
  inference(superposition,[],[f3130,f163])).
fof(f3112,plain,(
  ( ! [X6,X7] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X7))) | s(sK0,X7) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X6,sK0))) ) | $spl28),
  inference(superposition,[],[f154,f2475])).
fof(f3550,plain,(
  ( ! [X97,X98] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X97),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X98,sK0),X97,sK0))) ) | $spl48),
  inference(cnf_transformation,[],[f3550_D])).
fof(f3550_D,plain,(
  ( ! [X97,X98] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X97),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X98,sK0),X97,sK0))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3564,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0))) ) | $spl28),
  inference(condensation,[],[f3563])).
fof(f3563,plain,(
  ( ! [X83,X81,X84,X82] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X81),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X83) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X84,sK0),X81,sK0)) | s(bool,'T_0') = s(bool,X82)) ) | $spl28),
  inference(subsumption_resolution,[],[f3484,f190])).
fof(f3484,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,X82) != s(bool,X83) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X81),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X83) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X84,sK0),X81,sK0)) | s(bool,'T_0') = s(bool,X82)) ) | $spl28),
  inference(superposition,[],[f532,f3206])).
fof(f3560,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0))) ) | $spl28),
  inference(condensation,[],[f3559])).
fof(f3559,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X69),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X71) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X72,sK0),X69,sK0)) | s(bool,'T_0') = s(bool,X70)) ) | $spl28),
  inference(subsumption_resolution,[],[f3481,f3186])).
fof(f3481,plain,(
  ( ! [X70,X72,X71,X69] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X70))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X71))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X69),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X71) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X72,sK0),X69,sK0)) | s(bool,'T_0') = s(bool,X70)) ) | $spl28),
  inference(superposition,[],[f492,f3206])).
fof(f3558,plain,(
  ( ! [X68,X64,X65] : (s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X64),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X68,sK0),X64,sK0))) ) | $spl28),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f114,f146,f151,f155,f154,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f270,f280,f290,f482,f488,f491,f493,f495,f533,f535,f538,f603,f606,f608,f611,f617,f619,f716,f739,f929,f931,f946,f961,f963,f967,f1022,f1089,f1091,f1208,f1246,f588,f621,f900,f571,f898,f895,f893,f673,f671,f576,f573,f1070,f1048,f788,f760,f781,f1060,f1143,f1135,f1119,f912,f735,f733,f683,f681,f585,f568,f158,f157,f156,f1062,f1061,f1032,f1018,f1016,f934,f936,f1335,f1331,f1308,f1285,f1284,f1283,f1253,f1204,f1203,f1201,f1200,f1199,f1197,f1196,f1195,f1176,f1093,f1020,f965,f938,f933,f897,f892,f737,f675,f610,f575,f570,f537,f532,f497,f492,f239,f113,f1980,f1981,f1982,f1983,f1984,f1995,f1996,f1997,f1998,f1999,f2161,f2167,f2175,f2176,f2177,f2178,f2179,f2183,f2189,f2191,f2193,f2194,f2205,f2206,f2212,f2214,f2216,f2217,f2248,f2249,f2250,f2251,f2252,f2475,f3067,f3068,f3069,f3070,f3099,f3112,f3130,f3186,f3193,f3194,f3195,f3196,f3206,f3229,f3230,f3231,f3242,f3248,f3250,f3252,f3253,f3267,f3274,f3276,f3278,f3279,f3297,f3298,f3299,f3300,f3337,f3355,f3539,f3533,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3393,f3395,f3397,f3398,f3525,f3519,f3518,f3517,f3515,f3443,f3445,f3447,f3448,f3455,f3459,f3461,f3462,f3552,f3554,f3480])).
fof(f3555,plain,(
  ( ! [X59,X57,X58] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X57),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X59,sK0),X57,sK0)) | s(bool,'T_0') = s(bool,X58)) ) | ($spl18 | $spl28)),
  inference(subsumption_resolution,[],[f3478,f192])).
fof(f3478,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,X58)) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X57),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X59,sK0),X57,sK0)) | s(bool,'T_0') = s(bool,X58)) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f220,f3206])).
fof(f3553,plain,(
  ( ! [X28,X26,X29] : (p(s(bool,X28)) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X29,sK0),X26,sK0)) | s(bool,X28) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK2)))) ) | $spl50),
  inference(cnf_transformation,[],[f3553_D])).
fof(f3553_D,plain,(
  ( ! [X28,X26,X29] : (p(s(bool,X28)) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X29,sK0),X26,sK0)) | s(bool,X28) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK2)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3554,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X32)) | s(bool,X32) = s(bool,happ(s(fun(sK0,bool),X30),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X33,sK0),X30,sK0)) | s(bool,'T_0') = s(bool,X31)) ) | $spl28),
  inference(subsumption_resolution,[],[f3471,f192])).
fof(f3471,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) = s(bool,happ(s(fun(sK0,bool),X30),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X33,sK0),X30,sK0)) | s(bool,'T_0') = s(bool,X31)) ) | $spl28),
  inference(superposition,[],[f159,f3206])).
fof(f3552,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X28)) | s(bool,X28) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X29,sK0),X26,sK0)) | s(bool,'T_0') = s(bool,X27)) ) | $spl28),
  inference(subsumption_resolution,[],[f3470,f192])).
fof(f3470,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X29,sK0),X26,sK0)) | s(bool,'T_0') = s(bool,X27)) ) | $spl28),
  inference(superposition,[],[f159,f3206])).
fof(f3515,plain,(
  ( ! [X97,X98,X96] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X97),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X96) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X98,sK0),X97,sK0))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f3430])).
fof(f3430,plain,(
  ( ! [X97,X98,X96] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X97),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X96) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X98,sK0),X97,sK0)) | s(bool,'T_0') = s(bool,X96)) ) | $spl28),
  inference(superposition,[],[f3099,f3206])).
fof(f3549,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK2)))))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X73,sK0),X71,sK0)) | s(bool,'F_0') = s(bool,sK8('F_0',X72,bool))) ) | $spl46),
  inference(cnf_transformation,[],[f3549_D])).
fof(f3549_D,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK2)))))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X73,sK0),X71,sK0)) | s(bool,'F_0') = s(bool,sK8('F_0',X72,bool))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3517,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK2)))))) | s(bool,'F_0') = s(bool,sK8('F_0',X72,bool)) | s(bool,'T_0') = s(bool,X70) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X73,sK0),X71,sK0))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f3423])).
fof(f3423,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,happ(s(fun(bool,bool),X72),s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK2)))))) | s(bool,'F_0') = s(bool,sK8('F_0',X72,bool)) | s(bool,'T_0') = s(bool,X70) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X73,sK0),X71,sK0)) | s(bool,'T_0') = s(bool,X70)) ) | $spl28),
  inference(superposition,[],[f570,f3206])).
fof(f3548,plain,(
  ( ! [X64,X65,X63] : (s(bool,X64) != s(bool,happ(s(fun(sK0,bool),X63),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X65,sK0),X63,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64)))) ) | $spl44),
  inference(cnf_transformation,[],[f3548_D])).
fof(f3548_D,plain,(
  ( ! [X64,X65,X63] : (s(bool,X64) != s(bool,happ(s(fun(sK0,bool),X63),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X65,sK0),X63,sK0)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3518,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) != s(bool,happ(s(fun(sK0,bool),X63),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) | s(bool,'T_0') = s(bool,X62) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X65,sK0),X63,sK0))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f3421])).
fof(f3421,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) != s(bool,happ(s(fun(sK0,bool),X63),s(sK0,sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X64))) | s(bool,'T_0') = s(bool,X62) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X65,sK0),X63,sK0)) | s(bool,'T_0') = s(bool,X62)) ) | $spl28),
  inference(superposition,[],[f532,f3206])).
fof(f3547,plain,(
  ( ! [X52,X53,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X53,sK0),X51,sK0)) | s(bool,'F_0') = s(bool,X52)) ) | $spl42),
  inference(cnf_transformation,[],[f3547_D])).
fof(f3547_D,plain,(
  ( ! [X52,X53,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X53,sK0),X51,sK0)) | s(bool,'F_0') = s(bool,X52)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3519,plain,(
  ( ! [X52,X50,X53,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X50) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X53,sK0),X51,sK0))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f3418])).
fof(f3418,plain,(
  ( ! [X52,X50,X53,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X50) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X53,sK0),X51,sK0)) | s(bool,'T_0') = s(bool,X50)) ) | $spl28),
  inference(superposition,[],[f492,f3206])).
fof(f3525,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X0) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f3405])).
fof(f3405,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X0) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(superposition,[],[f114,f3206])).
fof(f3541,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK2))) = s(bool,happ(s(fun(sK0,bool),X29),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X31,sK0),X28,sK0)) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X30,sK0),X29,sK0))) ) | $spl40),
  inference(cnf_transformation,[],[f3541_D])).
fof(f3541_D,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK2))) = s(bool,happ(s(fun(sK0,bool),X29),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X31,sK0),X28,sK0)) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X30,sK0),X29,sK0))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3533,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK2))) = s(bool,happ(s(fun(sK0,bool),X29),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X30,sK0),X29,sK0)) | s(bool,'T_0') = s(bool,X27) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X31,sK0),X28,sK0))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f3363])).
fof(f3363,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK2))) = s(bool,happ(s(fun(sK0,bool),X29),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X30,sK0),X29,sK0)) | s(bool,'T_0') = s(bool,X27) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X31,sK0),X28,sK0)) | s(bool,'T_0') = s(bool,X27)) ) | $spl28),
  inference(superposition,[],[f3206,f3206])).
fof(f3539,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(duplicate_literal_removal,[],[f3356])).
fof(f3356,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8(sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0),X1,sK0)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(superposition,[],[f3206,f114])).
fof(f2949,plain,(
  s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | $spl32),
  inference(cnf_transformation,[],[f2949_D])).
fof(f2949_D,plain,(
  s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2954,plain,(
  ( ! [X78] : (s(bool,'T_0') != s(bool,X78) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X78)),bool))) ) | $spl36),
  inference(cnf_transformation,[],[f2954_D])).
fof(f2954_D,plain,(
  ( ! [X78] : (s(bool,'T_0') != s(bool,X78) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X78)),bool))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2721,plain,(
  ( ! [X78] : (s(bool,'T_0') != s(bool,X78) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X78)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | ($spl15 | $spl30)),
  inference(superposition,[],[f200,f2476])).
fof(f2950,plain,(
  ( ! [X74,X75,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X73))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X74))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X73)),bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,X74) = s(bool,X75)) ) | $spl34),
  inference(cnf_transformation,[],[f2950_D])).
fof(f2950_D,plain,(
  ( ! [X74,X75,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X73))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X74))) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X73)),bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,X74) = s(bool,X75)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2718,plain,(
  ( ! [X74,X75,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X73))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X74))) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X73)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool))) ) | $spl30),
  inference(superposition,[],[f497,f2476])).
fof(f2659,plain,(
  ( ! [X78] : (s(bool,'T_0') != s(bool,X78) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X78)),bool))) ) | ($spl15 | $spl30)),
  inference(superposition,[],[f200,f2476])).
fof(f2656,plain,(
  ( ! [X74,X75,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X73))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X74))) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X73)),bool))) ) | $spl30),
  inference(superposition,[],[f497,f2476])).
fof(f2476,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X2)),bool))) ) | $spl30),
  inference(cnf_transformation,[],[f2476_D])).
fof(f2476_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X2)),bool))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2466,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X0,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X2)),bool)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f2284])).
fof(f2284,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X0,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X2)),bool)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X0,sK0)) | s(bool,'T_0') = s(bool,sK7('const.option.SOME_1'(s(bool,X1)),bool))) ) | $spl0),
  inference(superposition,[],[f1348,f1348])).
fof(f2154,plain,(
  ( ! [X39,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X38,sK0)) | s(bool,'F_0') = s(bool,X39)) ) | $spl26),
  inference(cnf_transformation,[],[f2154_D])).
fof(f2154_D,plain,(
  ( ! [X39,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK2))))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X38,sK0)) | s(bool,'F_0') = s(bool,X39)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2131,plain,(
  ( ! [X39,X37,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X37) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X38,sK0))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f2058])).
fof(f2058,plain,(
  ( ! [X39,X37,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X37) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X38,sK0)) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(superposition,[],[f492,f928])).
fof(f1612,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1487,f111])).
fof(f1487,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f219,f113])).
fof(f1611,plain,(
  $false | ($spl15 | $spl16)),
  inference(trivial_inequality_removal,[],[f1490])).
fof(f1490,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f219,f490])).
fof(f1610,plain,(
  $false | ($spl15 | $spl16)),
  inference(trivial_inequality_removal,[],[f1500])).
fof(f1500,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f219,f879])).
fof(f1474,plain,(
  ( ! [X16] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X16,sK0)) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X16),s(sK0,sK2)))) ) | $spl22),
  inference(cnf_transformation,[],[f1474_D])).
fof(f1474_D,plain,(
  ( ! [X16] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X16,sK0)) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X16),s(sK0,sK2)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1486,plain,(
  ( ! [X43,X44] : (s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X43,sK0)) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X43),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X44)) ) | $spl0),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f113,f114,f146,f151,f155,f154,f158,f157,f156,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f239,f270,f280,f290,f482,f488,f491,f492,f493,f495,f497,f532,f533,f535,f537,f538,f568,f585,f570,f571,f573,f575,f576,f588,f603,f606,f608,f610,f611,f617,f619,f621,f681,f671,f673,f675,f683,f716,f733,f735,f737,f739,f781,f788,f760,f170,f813,f869,f892,f893,f895,f897,f898,f900,f912,f921,f922,f928,f929,f931,f933,f934,f936,f938,f946,f961,f963,f965,f967,f1016,f1018,f1020,f1022,f1060,f1061,f1032,f1062,f1045,f1070,f1048,f1089,f1091,f1093,f1135,f1116,f1143,f1119,f1195,f1196,f1197,f1199,f1200,f1176,f1201,f1203,f1204,f1208,f1246,f1253,f1271,f1283,f1284,f1285,f1331,f1308,f1335,f1472,f1467,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1363,f1365,f1367,f1368,f1460,f1397,f1399,f1401,f1402,f1409,f1413,f1415,f1416,f1479,f1481,f1432,f1485,f1433])).
fof(f1485,plain,(
  ( ! [X39,X40] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X39,sK0)) | s(bool,'T_0') = s(bool,X40)) ) | $spl0),
  inference(global_subsumption,[],[f110,f109,f108,f107,f106,f105,f111,f113,f114,f146,f151,f155,f154,f158,f157,f156,f162,f159,f163,f164,f165,f166,f190,f192,f197,f218,f221,f222,f239,f270,f280,f290,f482,f488,f491,f492,f493,f495,f497,f532,f533,f535,f537,f538,f568,f585,f570,f571,f573,f575,f576,f588,f603,f606,f608,f610,f611,f617,f619,f621,f681,f671,f673,f675,f683,f716,f733,f735,f737,f739,f781,f788,f760,f170,f813,f869,f892,f893,f895,f897,f898,f900,f912,f921,f922,f928,f929,f931,f933,f934,f936,f938,f946,f961,f963,f965,f967,f1016,f1018,f1020,f1022,f1060,f1061,f1032,f1062,f1045,f1070,f1048,f1089,f1091,f1093,f1135,f1116,f1143,f1119,f1195,f1196,f1197,f1199,f1200,f1176,f1201,f1203,f1204,f1208,f1246,f1253,f1271,f1283,f1284,f1285,f1331,f1308,f1335,f1472,f1467,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1363,f1365,f1367,f1368,f1460,f1397,f1399,f1401,f1402,f1409,f1413,f1415,f1416,f1479,f1481,f1432])).
fof(f1482,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X34),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X34,sK0)) | s(bool,'T_0') = s(bool,X35)) ) | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f1430,f192])).
fof(f1430,plain,(
  ( ! [X35,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X34),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X34,sK0)) | s(bool,'T_0') = s(bool,X35)) ) | ($spl0 | $spl18)),
  inference(superposition,[],[f220,f928])).
fof(f1480,plain,(
  ( ! [X13,X11] : (p(s(bool,X13)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X11,sK0)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK2)))) ) | $spl24),
  inference(cnf_transformation,[],[f1480_D])).
fof(f1480_D,plain,(
  ( ! [X13,X11] : (p(s(bool,X13)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X11,sK0)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK2)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1481,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,X16) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X14,sK0)) | s(bool,'T_0') = s(bool,X15)) ) | $spl0),
  inference(subsumption_resolution,[],[f1423,f192])).
fof(f1423,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X14,sK0)) | s(bool,'T_0') = s(bool,X15)) ) | $spl0),
  inference(superposition,[],[f159,f928])).
fof(f1479,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X11,sK0)) | s(bool,'T_0') = s(bool,X12)) ) | $spl0),
  inference(subsumption_resolution,[],[f1422,f192])).
fof(f1422,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X11,sK0)) | s(bool,'T_0') = s(bool,X12)) ) | $spl0),
  inference(superposition,[],[f159,f928])).
fof(f1460,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X0) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1375])).
fof(f1375,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X0) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f114,f928])).
fof(f1467,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X16),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X16,sK0)) | s(bool,'T_0') = s(bool,X15)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1341])).
fof(f1341,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X16),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X16,sK0)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl0),
  inference(superposition,[],[f928,f114])).
fof(f1473,plain,(
  ( ! [X2,X1] : (s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) = s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0))) ) | $spl20),
  inference(cnf_transformation,[],[f1473_D])).
fof(f1473_D,plain,(
  ( ! [X2,X1] : (s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) = s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0)) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1472,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) = s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0)) | s(bool,'T_0') = s(bool,X0) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1336])).
fof(f1336,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK2))) = s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK2))) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X2,sK0)) | s(bool,'T_0') = s(bool,X0) | s(sK0,sK2) = s(sK0,sK8('const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)),X1,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f928,f928])).
fof(f866,plain,(
  $false | ($spl0 | $spl2 | $spl10)),
  inference(subsumption_resolution,[],[f865,f180])).
fof(f865,plain,(
  s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK6))) | ($spl0 | $spl2 | $spl10)),
  inference(forward_demodulation,[],[f864,f691])).
fof(f691,plain,(
  s(sK1,sK6) = s(sK1,sK7('const.option.SOME_1'(s(sK1,sK6)),sK1)) | $spl10),
  inference(unit_resulting_resolution,[],[f151,f180,f615])).
fof(f615,plain,(
  ( ! [X9] : (s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),X9) | s(sK1,sK6) = s(sK1,sK7(X9,sK1)) | s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),X9)) ) | $spl10),
  inference(superposition,[],[f502,f146])).
fof(f502,plain,(
  ( ! [X0] : (s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,X0))) | s(sK1,sK6) = s(sK1,X0)) ) | $spl10),
  inference(superposition,[],[f154,f180])).
fof(f864,plain,(
  s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK7('const.option.SOME_1'(s(sK1,sK6)),sK1)))) | ($spl0 | $spl2 | $spl10)),
  inference(subsumption_resolution,[],[f857,f832])).
fof(f832,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))) | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f821,f154])).
fof(f821,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))))) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f811,f715])).
fof(f715,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK6))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(X0,X2),s(fun(sK1,X0),X1)))) ) | $spl10),
  inference(superposition,[],[f165,f180])).
fof(f811,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))))) | $spl0),
  inference(unit_resulting_resolution,[],[f170,f155])).
fof(f857,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))) | s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK7('const.option.SOME_1'(s(sK1,sK6)),sK1)))) | ($spl2 | $spl10)),
  inference(superposition,[],[f172,f691])).
fof(f863,plain,(
  $false | ($spl0 | $spl2 | $spl10)),
  inference(subsumption_resolution,[],[f862,f180])).
fof(f862,plain,(
  s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK6))) | ($spl0 | $spl2 | $spl10)),
  inference(forward_demodulation,[],[f861,f624])).
fof(f624,plain,(
  s(sK1,sK6) = s(sK1,sK7(sK4,sK1)) | $spl10),
  inference(unit_resulting_resolution,[],[f604,f502])).
fof(f604,plain,(
  s('type.option.option'(sK1),sK4) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK7(sK4,sK1)))) | $spl10),
  inference(unit_resulting_resolution,[],[f189,f146])).
fof(f189,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') != s('type.option.option'(sK1),sK4) | $spl10),
  inference(superposition,[],[f151,f180])).
fof(f861,plain,(
  s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK7(sK4,sK1)))) | ($spl0 | $spl2 | $spl10)),
  inference(subsumption_resolution,[],[f856,f832])).
fof(f856,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))) | s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK7(sK4,sK1)))) | ($spl2 | $spl10)),
  inference(superposition,[],[f172,f624])).
fof(f860,plain,(
  $false | ($spl0 | $spl2 | $spl10)),
  inference(subsumption_resolution,[],[f855,f832])).
fof(f855,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f180,f172])).
fof(f778,plain,(
  $false | ($spl1 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f777])).
fof(f777,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) | ($spl1 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f774,f178])).
fof(f774,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))))) | ($spl1 | $spl10)),
  inference(backward_demodulation,[],[f715,f487])).
fof(f487,plain,(
  s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))))) | $spl1),
  inference(unit_resulting_resolution,[],[f171,f154])).
fof(f776,plain,(
  $false | ($spl1 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f775])).
fof(f775,plain,(
  s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))))) != s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))))) | ($spl1 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f773,f178])).
fof(f773,plain,(
  s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))))) != s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))))))) | ($spl1 | $spl10)),
  inference(backward_demodulation,[],[f715,f563])).
fof(f563,plain,(
  s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK2))))) != s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f487,f154])).
fof(f220,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl18),
  inference(cnf_transformation,[],[f220_D])).
fof(f220_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f222,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f216,f192])).
fof(f216,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f159,f114])).
fof(f218,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f214,f192])).
fof(f214,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f159,f114])).
fof(f200,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl15),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f197,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f114])).
fof(f186,plain,(
  $false | ($spl4 | $spl10)),
  inference(subsumption_resolution,[],[f185,f151])).
fof(f185,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK6))) | ($spl4 | $spl10)),
  inference(forward_demodulation,[],[f180,f173])).
fof(f183,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f182,f163])).
fof(f182,plain,(
  s(sK0,sK2) != s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),'const.option.NONE_0'),s(sK0,sK2),s(fun(sK1,sK0),sK3))) | ($spl1 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f181,f173])).
fof(f181,plain,(
  s(sK0,sK2) != s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK2),s(fun(sK1,sK0),sK3))) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f171,f175])).
fof(f170,plain,(
  s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))) | $spl0),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f180,plain,(
  s('type.option.option'(sK1),sK4) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK6))) | $spl10),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  s('type.option.option'(sK1),sK4) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK6))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f105,plain,(
  s('type.option.option'(sK1),sK4) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK6))) | s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK4) | s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  (s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))) | (s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK4) & s(sK0,sK2) = s(sK0,sK5)) | (s('type.option.option'(sK1),sK4) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK6))) & s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))))) & (s(sK0,sK2) != s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))) | ((s('type.option.option'(sK1),'const.option.NONE_0') != s('type.option.option'(sK1),sK4) | s(sK0,sK2) != s(sK0,sK5)) & ! [X6] : (s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,X6))) | s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f73])).
fof(f73,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X0,X2) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) | (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) & s(X0,X2) = s(X0,X5)) | ? [X7] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X7))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X7))))) & (s(X0,X2) != s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) | ((s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4) | s(X0,X2) != s(X0,X5)) & ! [X6] : (s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | s(X0,X2) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6)))))))),
  inference(rectify,[],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X0,X2) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) | (s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) & s(X0,X2) = s(X0,X5)) | ? [X6] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & (s(X0,X2) != s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) | ((s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4) | s(X0,X2) != s(X0,X5)) & ! [X6] : (s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | s(X0,X2) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6)))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X0,X2) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) | ((s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) & s(X0,X2) = s(X0,X5)) | ? [X6] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6)))))) & (s(X0,X2) != s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) | ((s('type.option.option'(X1),'const.option.NONE_0') != s('type.option.option'(X1),X4) | s(X0,X2) != s(X0,X5)) & ! [X6] : (s('type.option.option'(X1),X4) != s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) | s(X0,X2) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6)))))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(X0,X2) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) <~> ((s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) & s(X0,X2) = s(X0,X5)) | ? [X6] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s(X0,X2) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) <=> ((s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) & s(X0,X2) = s(X0,X5)) | ? [X6] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (s(X0,X2) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X4),s(X0,X5),s(fun(X1,X0),X3))) <=> ((s('type.option.option'(X1),'const.option.NONE_0') = s('type.option.option'(X1),X4) & s(X0,X2) = s(X0,X5)) | ? [X6] : (s('type.option.option'(X1),X4) = s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X6))) & s(X0,X2) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (s(X6,X8) = s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X10),s(X6,X11),s(fun(X7,X6),X9))) <=> ((s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X10) & s(X6,X8) = s(X6,X11)) | ? [X12] : (s('type.option.option'(X7),X10) = s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X12))) & s(X6,X8) = s(X6,happ(s(fun(X7,X6),X9),s(X7,X12))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (s(X6,X8) = s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X10),s(X6,X11),s(fun(X7,X6),X9))) <=> ((s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X10) & s(X6,X8) = s(X6,X11)) | ? [X12] : (s('type.option.option'(X7),X10) = s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X12))) & s(X6,X8) = s(X6,happ(s(fun(X7,X6),X9),s(X7,X12))))))),
  file('thm.clos_numberProof.clos_numberProof_|48|_prove_ltb','thm.clos_numberProof.clos_numberProof_|48|_prove_ltb')).
fof(f175,plain,(
  s(sK0,sK2) = s(sK0,sK5) | $spl6),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  s(sK0,sK2) = s(sK0,sK5) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f106,plain,(
  s('type.option.option'(sK1),sK4) = s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,sK6))) | s(sK0,sK2) = s(sK0,sK5) | s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)))),
  inference(cnf_transformation,[],[f74])).
fof(f178,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))) | $spl8),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f107,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))) | s('type.option.option'(sK1),'const.option.NONE_0') = s('type.option.option'(sK1),sK4) | s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)))),
  inference(cnf_transformation,[],[f74])).
fof(f108,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,sK6))) | s(sK0,sK2) = s(sK0,sK5) | s(sK0,sK2) = s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)))),
  inference(cnf_transformation,[],[f74])).
fof(f171,plain,(
  s(sK0,sK2) != s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))) | $spl1),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  s(sK0,sK2) != s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f174,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') != s('type.option.option'(sK1),sK4) | $spl5),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  s('type.option.option'(sK1),'const.option.NONE_0') != s('type.option.option'(sK1),sK4) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f176,plain,(
  s(sK0,sK2) != s(sK0,sK5) | $spl7),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  s(sK0,sK2) != s(sK0,sK5) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f109,plain,(
  s(sK0,sK2) != s(sK0,sK5) | s('type.option.option'(sK1),'const.option.NONE_0') != s('type.option.option'(sK1),sK4) | s(sK0,sK2) != s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)))),
  inference(cnf_transformation,[],[f74])).
fof(f172,plain,(
  ( ! [X6] : (s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,X6))) | s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,X6)))) ) | $spl2),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  ( ! [X6] : (s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,X6))) | s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,X6)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f110,plain,(
  ( ! [X6] : (s(sK0,sK2) != s(sK0,happ(s(fun(sK1,sK0),sK3),s(sK1,X6))) | s('type.option.option'(sK1),sK4) != s('type.option.option'(sK1),'const.option.SOME_1'(s(sK1,X6))) | s(sK0,sK2) != s(sK0,'const.option.option_CASE_3'(s('type.option.option'(sK1),sK4),s(sK0,sK5),s(fun(sK1,sK0),sK3)))) )),
  inference(cnf_transformation,[],[f74])).
% SZS output end Proof for thm.clos_numberProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 236
% Passive clauses: 3227
% Generated clauses: 30440
% Final active clauses: 88
% Final passive clauses: 1049
% Input formulas: 28
% Initial clauses: 62
% 
% Duplicate literals: 1002
% Trivial inequalities: 353
% Fw subsumption resolutions: 2161
% Fw demodulations: 5210
% Bw demodulations: 1596
% Condensations: 1799
% Global subsumptions: 27
% 
% Simple tautologies: 314
% Equational tautologies: 309
% Forward subsumptions: 15621
% Fw demodulations to eq. taut.: 224
% Bw demodulations to eq. taut.: 122
% 
% Unit resulting resolution: 1389
% Factoring: 1
% Forward superposition: 7511
% Backward superposition: 7033
% Self superposition: 257
% Equality factoring: 84
% Equality resolution: 45
% 
% Split clauses: 136
% Split components: 287
% SAT solver clauses: 3594
% SAT solver unit clauses: 15
% SAT solver binary clauses: 341
% 
% Sat splits: 136
% Sat splitting refutations: 66
% 
% Memory used [KB]: 17014
% Time elapsed: 1.887 s
% ------------------------------
% ------------------------------
% Success in time 3.32 s
