Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ImplicitQueue
% (116490)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 99
% Passive clauses: 2357
% Generated clauses: 11745
% Final active clauses: 91
% Final passive clauses: 2125
% Input formulas: 37
% Initial clauses: 71
% 
% Duplicate literals: 512
% Trivial inequalities: 220
% Fw subsumption resolutions: 1285
% Fw demodulations: 815
% Bw demodulations: 3
% 
% Simple tautologies: 285
% Equational tautologies: 66
% Forward subsumptions: 4435
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 291
% Factoring: 3
% Forward superposition: 3873
% Backward superposition: 4156
% Self superposition: 238
% Equality factoring: 164
% Equality resolution: 22
% 
% Split clauses: 274
% Split components: 560
% SAT solver clauses: 581
% SAT solver unit clauses: 30
% SAT solver binary clauses: 148
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 274
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 579
% TWLsolver calls for satisfiability: 133
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ImplicitQueue
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ImplicitQueue
% SZS output start Proof for thm.ImplicitQueue
fof(f4683,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4680,f249,f135,f212,f210,f164,f227,f4589,f4590,f172,f235,f3917,f3918,f134,f211,f209,f186,f245])).
fof(f245,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s('type.ImplicitQueue.digit'(X1),X2))) = s(X0,'const.ImplicitQueue.queue_CASE_3'(s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X1),X2))),s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s(fun('type.ImplicitQueue.digit'(X1),fun('type.ImplicitQueue.queue'(X1),fun('type.ImplicitQueue.digit'(X1),X0))),X4)))) ) | $spl56),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s('type.ImplicitQueue.digit'(X1),X2))) = s(X0,'const.ImplicitQueue.queue_CASE_3'(s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X1),X2))),s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s(fun('type.ImplicitQueue.digit'(X1),fun('type.ImplicitQueue.queue'(X1),fun('type.ImplicitQueue.digit'(X1),X0))),X4)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f186,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s('type.ImplicitQueue.digit'(X1),X2))) = s(X0,'const.ImplicitQueue.queue_CASE_3'(s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X1),X2))),s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s(fun('type.ImplicitQueue.digit'(X1),fun('type.ImplicitQueue.queue'(X1),fun('type.ImplicitQueue.digit'(X1),X0))),X4)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s('type.ImplicitQueue.digit'(X1),X2))) = s(X0,'const.ImplicitQueue.queue_CASE_3'(s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X1),X2))),s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s(fun('type.ImplicitQueue.digit'(X1),fun('type.ImplicitQueue.queue'(X1),fun('type.ImplicitQueue.digit'(X1),X0))),X4)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,happ(s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s('type.ImplicitQueue.digit'(X1),X2))) = s(X0,'const.ImplicitQueue.queue_CASE_3'(s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X1),X2))),s(fun('type.ImplicitQueue.digit'(X1),X0),X3),s(fun('type.ImplicitQueue.digit'(X1),fun('type.ImplicitQueue.queue'(X1),fun('type.ImplicitQueue.digit'(X1),X0))),X4)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X6] : ! [X7,X10,X12] : s(X11,happ(s(fun('type.ImplicitQueue.digit'(X6),X11),X10),s('type.ImplicitQueue.digit'(X6),X7))) = s(X11,'const.ImplicitQueue.queue_CASE_3'(s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X6),X7))),s(fun('type.ImplicitQueue.digit'(X6),X11),X10),s(fun('type.ImplicitQueue.digit'(X6),fun('type.ImplicitQueue.queue'(X6),fun('type.ImplicitQueue.digit'(X6),X11))),X12)))),
  file('thm.ImplicitQueue.queue_|49||49||split|0','thm.ImplicitQueue.queue_case_def|split|0')).
fof(f209,plain,(
  s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) = s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK2))) | $spl0),
  inference(cnf_transformation,[],[f209_D])).
fof(f209_D,plain,(
  s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) = s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f211,plain,(
  s('type.ImplicitQueue.digit'(sK0),sK1) = s('type.ImplicitQueue.digit'(sK0),sK2) | $spl2),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  s('type.ImplicitQueue.digit'(sK0),sK1) = s('type.ImplicitQueue.digit'(sK0),sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f134,plain,(
  s('type.ImplicitQueue.digit'(sK0),sK1) = s('type.ImplicitQueue.digit'(sK0),sK2) | s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) = s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK2)))),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  (s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) = s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK2))) | s('type.ImplicitQueue.digit'(sK0),sK1) = s('type.ImplicitQueue.digit'(sK0),sK2)) & (s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) != s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK2))) | s('type.ImplicitQueue.digit'(sK0),sK1) != s('type.ImplicitQueue.digit'(sK0),sK2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f113])).
fof(f113,plain,(
  ? [X0,X1,X2] : ((s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X1))) = s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2))) | s('type.ImplicitQueue.digit'(X0),X1) = s('type.ImplicitQueue.digit'(X0),X2)) & (s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X1))) != s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2))) | s('type.ImplicitQueue.digit'(X0),X1) != s('type.ImplicitQueue.digit'(X0),X2)))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ? [X0,X1,X2] : (s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X1))) = s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2))) <~> s('type.ImplicitQueue.digit'(X0),X1) = s('type.ImplicitQueue.digit'(X0),X2))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2] : (s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X1))) = s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2))) <=> s('type.ImplicitQueue.digit'(X0),X1) = s('type.ImplicitQueue.digit'(X0),X2))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X1))) = s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2))) <=> s('type.ImplicitQueue.digit'(X0),X1) = s('type.ImplicitQueue.digit'(X0),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X6),X7))) = s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X6),X8))) <=> s('type.ImplicitQueue.digit'(X6),X7) = s('type.ImplicitQueue.digit'(X6),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X6),X7))) = s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X6),X8))) <=> s('type.ImplicitQueue.digit'(X6),X7) = s('type.ImplicitQueue.digit'(X6),X8))),
  file('thm.ImplicitQueue.queue_|49||49||split|0','thm.ImplicitQueue.queue_|49||49||split|0')).
fof(f3918,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.digit'(sK0),X0),X1),s('type.ImplicitQueue.digit'(sK0),sK1))) = s(X0,happ(s(fun('type.ImplicitQueue.digit'(sK0),X0),X1),s('type.ImplicitQueue.digit'(sK0),sK2)))) ) | $spl616),
  inference(cnf_transformation,[],[f3918_D])).
fof(f3918_D,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.digit'(sK0),X0),X1),s('type.ImplicitQueue.digit'(sK0),sK1))) = s(X0,happ(s(fun('type.ImplicitQueue.digit'(sK0),X0),X1),s('type.ImplicitQueue.digit'(sK0),sK2)))) ) <=> ~$spl616),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl616])])).
fof(f3917,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.digit'(sK0),X0),X1),s('type.ImplicitQueue.digit'(sK0),sK1))) = s(X0,happ(s(fun('type.ImplicitQueue.digit'(sK0),X0),X1),s('type.ImplicitQueue.digit'(sK0),sK2)))) ) | ($spl0 | $spl56)),
  inference(forward_demodulation,[],[f3832,f245])).
fof(f3832,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.digit'(sK0),X0),X1),s('type.ImplicitQueue.digit'(sK0),sK2))) = s(X0,'const.ImplicitQueue.queue_CASE_3'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))),s(fun('type.ImplicitQueue.digit'(sK0),X0),X1),s(fun('type.ImplicitQueue.digit'(sK0),fun('type.ImplicitQueue.queue'(sK0),fun('type.ImplicitQueue.digit'(sK0),X0))),X2)))) ) | ($spl0 | $spl56)),
  inference(superposition,[],[f245,f209])).
fof(f235,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) | $spl36),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f172,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X9,X13,X14] : s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X13),s('type.num.num',X14))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X9))),s(fun('type.num.num',X6),X13))),s('type.num.num',X14)))),
  file('thm.ImplicitQueue.queue_|49||49||split|0','ARITY_0_3_const.ind_type.FCONS')).
fof(f4590,plain,(
  ( ! [X4,X5] : (s('type.ImplicitQueue.digit'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.digit'(sK0),sK1),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X4),s('type.num.num',X5))) = s('type.ImplicitQueue.digit'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.digit'(sK0),sK2),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X4),s('type.num.num',X5)))) ) | $spl708),
  inference(cnf_transformation,[],[f4590_D])).
fof(f4590_D,plain,(
  ( ! [X4,X5] : (s('type.ImplicitQueue.digit'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.digit'(sK0),sK1),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X4),s('type.num.num',X5))) = s('type.ImplicitQueue.digit'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.digit'(sK0),sK2),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X4),s('type.num.num',X5)))) ) <=> ~$spl708),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl708])])).
fof(f4589,plain,(
  ( ! [X4,X5] : (s('type.ImplicitQueue.digit'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.digit'(sK0),sK1),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X4),s('type.num.num',X5))) = s('type.ImplicitQueue.digit'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.digit'(sK0),sK2),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X4),s('type.num.num',X5)))) ) | ($spl36 | $spl616)),
  inference(forward_demodulation,[],[f4472,f235])).
fof(f4472,plain,(
  ( ! [X4,X5] : (s('type.ImplicitQueue.digit'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.digit'(sK0),sK2),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X4),s('type.num.num',X5))) = s('type.ImplicitQueue.digit'(sK0),happ(s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),happ(s(fun(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),fun('type.num.num','type.ImplicitQueue.digit'(sK0))),happ(s(fun('type.ImplicitQueue.digit'(sK0),fun(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),fun('type.num.num','type.ImplicitQueue.digit'(sK0)))),'const.ind_type.FCONS_0'),s('type.ImplicitQueue.digit'(sK0),sK1))),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X4))),s('type.num.num',X5)))) ) | ($spl36 | $spl616)),
  inference(superposition,[],[f235,f3918])).
fof(f227,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f164,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7,X10] : s(X6,X7) = s(X6,'const.ind_type.FCONS_3'(s(X6,X7),s(fun('type.num.num',X6),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.ImplicitQueue.queue_|49||49||split|0','thm.ind_type.FCONS|split|0')).
fof(f210,plain,(
  s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) != s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK2))) | $spl1),
  inference(cnf_transformation,[],[f210_D])).
fof(f210_D,plain,(
  s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) != s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f212,plain,(
  s('type.ImplicitQueue.digit'(sK0),sK1) != s('type.ImplicitQueue.digit'(sK0),sK2) | $spl3),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  s('type.ImplicitQueue.digit'(sK0),sK1) != s('type.ImplicitQueue.digit'(sK0),sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f135,plain,(
  s('type.ImplicitQueue.digit'(sK0),sK1) != s('type.ImplicitQueue.digit'(sK0),sK2) | s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) != s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK2)))),
  inference(cnf_transformation,[],[f114])).
fof(f249,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f248])).
fof(f248,plain,(
  s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) != s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(sK0),sK1))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f210,f211])).
fof(f4680,plain,(
  $false | ($spl3 | $spl20 | $spl708)),
  inference(subsumption_resolution,[],[f4679,f212])).
fof(f4679,plain,(
  s('type.ImplicitQueue.digit'(sK0),sK1) = s('type.ImplicitQueue.digit'(sK0),sK2) | ($spl20 | $spl708)),
  inference(forward_demodulation,[],[f4662,f227])).
fof(f4662,plain,(
  ( ! [X2] : (s('type.ImplicitQueue.digit'(sK0),sK2) = s('type.ImplicitQueue.digit'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.digit'(sK0),sK1),s(fun('type.num.num','type.ImplicitQueue.digit'(sK0)),X2),s('type.num.num','const.num.|48|_0')))) ) | ($spl20 | $spl708)),
  inference(superposition,[],[f4590,f227])).
% SZS output end Proof for thm.ImplicitQueue
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 60
% Passive clauses: 383
% Generated clauses: 4504
% Final active clauses: 52
% Final passive clauses: 311
% Input formulas: 37
% Initial clauses: 71
% 
% Duplicate literals: 237
% Trivial inequalities: 46
% Fw subsumption resolutions: 249
% Fw demodulations: 324
% Bw demodulations: 11
% Condensations: 492
% Global subsumptions: 4
% 
% Simple tautologies: 108
% Equational tautologies: 35
% Forward subsumptions: 1894
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 64
% Factoring: 1
% Forward superposition: 1135
% Backward superposition: 1305
% Self superposition: 94
% Equality factoring: 68
% Equality resolution: 9
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 792
% SAT solver unit clauses: 77
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 10
% SAT solver learnt literals: 1
% 
% Sat splits: 14
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 773
% TWLsolver calls for satisfiability: 221
% 
% Memory used [KB]: 5884
% Time elapsed: 0.216 s
% ------------------------------
% ------------------------------
% Success in time 0.84 s
