Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closSem
% (8901)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 183
% Passive clauses: 3055
% Generated clauses: 21319
% Final active clauses: 175
% Final passive clauses: 2636
% Input formulas: 69
% Initial clauses: 108
% 
% Duplicate literals: 680
% Trivial inequalities: 382
% Fw subsumption resolutions: 2108
% Fw demodulations: 2536
% Bw demodulations: 6
% 
% Simple tautologies: 412
% Equational tautologies: 47
% Forward subsumptions: 11674
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 582
% Factoring: 2
% Forward superposition: 8679
% Backward superposition: 6186
% Self superposition: 161
% Equality factoring: 86
% Equality resolution: 37
% 
% Split clauses: 257
% Split components: 528
% SAT solver clauses: 653
% SAT solver unit clauses: 36
% SAT solver binary clauses: 165
% 
% Sat splits: 257
% 
% TWLsolver clauses: 652
% TWLsolver calls for satisfiability: 184
% 
% Memory used [KB]: 14711
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closSem
% SZS output start Proof for thm.closSem
fof(f7254,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7244,f5947,f433,f223,f335,f224,f337,f225,f339,f226,f341,f227,f343,f333,f287,f384,f302,f396,f5985,f6241,f6242,f6192,f6194,f6362,f6164,f6165,f5989,f5992,f5986,f5988,f6189,f228,f344,f342,f340,f338,f336,f334,f284,f387,f301,f395,f286,f385,f298,f392,f6964,f6736,f6737,f6190,f6686,f6568,f6741,f6420,f283,f388,f300,f394,f285,f386,f299,f393])).
fof(f393,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) | $spl96),
  inference(cnf_transformation,[],[f393_D])).
fof(f393_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f299,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X43,X27,X44,X28,X45] : s('type.ffi.ffi_state'(X6),X44) = s('type.ffi.ffi_state'(X6),'const.closSem.state_ffi_1'(s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X43),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X27),s('type.ffi.ffi_state'(X6),X44),s('type.num.num',X28),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X45)))))),
  file('thm.closSem.state_|49||49|','thm.closSem.state_ffi')).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2)))) ) | $spl82),
  inference(cnf_transformation,[],[f386_D])).
fof(f386_D,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))))) & (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) != s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) != s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))) | s('type.closSem.state'(X0),X1) = s('type.closSem.state'(X0),X2)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))))) & ((s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) != s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) != s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))) | s('type.closSem.state'(X0),X1) = s('type.closSem.state'(X0),X2)))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : (s('type.closSem.state'(X0),X1) = s('type.closSem.state'(X0),X2) <=> (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ! [X1,X2] : (s('type.closSem.state'(X0),X1) = s('type.closSem.state'(X0),X2) <=> (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))) & s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.closSem.state_ffi_1'(s('type.closSem.state'(X0),X2))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X18,X17] : (s('type.closSem.state'(X6),X17) = s('type.closSem.state'(X6),X18) <=> (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X6),X17))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X6),X18))) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X6),X17))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X6),X18))) & s('type.ffi.ffi_state'(X6),'const.closSem.state_ffi_1'(s('type.closSem.state'(X6),X17))) = s('type.ffi.ffi_state'(X6),'const.closSem.state_ffi_1'(s('type.closSem.state'(X6),X18))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),X17))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),X18))) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X6),X17))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X6),X18)))))),
  file('thm.closSem.state_|49||49|','thm.closSem.state_component_equality')).
fof(f394,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) | $spl98),
  inference(cnf_transformation,[],[f394_D])).
fof(f394_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f300,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X43,X27,X44,X28,X45] : s('type.list.list'('type.option.option'('type.closSem.v')),X43) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X43),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X27),s('type.ffi.ffi_state'(X6),X44),s('type.num.num',X28),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X45)))))),
  file('thm.closSem.state_|49||49|','thm.closSem.state_globals')).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))) ) | $spl86),
  inference(cnf_transformation,[],[f388_D])).
fof(f388_D,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))) | s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f6420,plain,(
  ( ! [X17] : (s('type.closSem.state'(sK0),X17) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.list.list'('type.option.option'('type.closSem.v')),sK6) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),X17)))) ) | $spl760),
  inference(cnf_transformation,[],[f6420_D])).
fof(f6420_D,plain,(
  ( ! [X17] : (s('type.closSem.state'(sK0),X17) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.list.list'('type.option.option'('type.closSem.v')),sK6) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),X17)))) ) <=> ~$spl760),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl760])])).
fof(f6741,plain,(
  ( ! [X17] : (s('type.list.list'('type.option.option'('type.closSem.v')),sK6) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),X17))) | s('type.closSem.state'(sK0),X17) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))) ) | ($spl86 | $spl98 | $spl768)),
  inference(forward_demodulation,[],[f6709,f394])).
fof(f6709,plain,(
  ( ! [X17] : (s('type.closSem.state'(sK0),X17) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),X17))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))))) ) | ($spl86 | $spl768)),
  inference(superposition,[],[f388,f6686])).
fof(f6568,plain,(
  $false | ($spl3 | $spl98 | $spl760)),
  inference(subsumption_resolution,[],[f6567,f336])).
fof(f6567,plain,(
  s('type.list.list'('type.option.option'('type.closSem.v')),sK1) = s('type.list.list'('type.option.option'('type.closSem.v')),sK6) | ($spl98 | $spl760)),
  inference(forward_demodulation,[],[f6557,f394])).
fof(f6557,plain,(
  s('type.list.list'('type.option.option'('type.closSem.v')),sK6) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))))) | $spl760),
  inference(equality_resolution,[],[f6420])).
fof(f6686,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | $spl768),
  inference(cnf_transformation,[],[f6686_D])).
fof(f6686_D,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) <=> ~$spl768),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl768])])).
fof(f6190,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f333,f343])).
fof(f6737,plain,(
  ( ! [X13] : (s('type.closSem.state'(sK0),X13) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X13)))) ) | $spl770),
  inference(cnf_transformation,[],[f6737_D])).
fof(f6737_D,plain,(
  ( ! [X13] : (s('type.closSem.state'(sK0),X13) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X13)))) ) <=> ~$spl770),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl770])])).
fof(f6736,plain,(
  ( ! [X13] : (s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X13))) | s('type.closSem.state'(sK0),X13) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))) ) | ($spl82 | $spl96 | $spl768)),
  inference(forward_demodulation,[],[f6705,f393])).
fof(f6705,plain,(
  ( ! [X13] : (s('type.closSem.state'(sK0),X13) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),X13))) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))))) ) | ($spl82 | $spl768)),
  inference(superposition,[],[f386,f6686])).
fof(f6964,plain,(
  $false | ($spl7 | $spl96 | $spl770)),
  inference(subsumption_resolution,[],[f6963,f340])).
fof(f6963,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK8) | ($spl96 | $spl770)),
  inference(forward_demodulation,[],[f6953,f393])).
fof(f6953,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),'const.closSem.state_ffi_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))))) | $spl770),
  inference(equality_resolution,[],[f6737])).
fof(f392,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num',X4) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) | $spl94),
  inference(cnf_transformation,[],[f392_D])).
fof(f392_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num',X4) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f298,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num',X4) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.num.num',X4) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.num.num',X4) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X43,X27,X44,X28,X45] : s('type.num.num',X28) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X43),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X27),s('type.ffi.ffi_state'(X6),X44),s('type.num.num',X28),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X45)))))),
  file('thm.closSem.state_|49||49|','thm.closSem.state_clock')).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2)))) ) | $spl80),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X1))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))) | s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f395,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) | $spl100),
  inference(cnf_transformation,[],[f395_D])).
fof(f395_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f301,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X43,X27,X44,X28,X45] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X27) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X43),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X27),s('type.ffi.ffi_state'(X6),X44),s('type.num.num',X28),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X45)))))),
  file('thm.closSem.state_|49||49|','thm.closSem.state_refs')).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))) ) | $spl84),
  inference(cnf_transformation,[],[f387_D])).
fof(f387_D,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(X0),X2))) | s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f334,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) | $spl1),
  inference(cnf_transformation,[],[f334_D])).
fof(f334_D,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f336,plain,(
  s('type.list.list'('type.option.option'('type.closSem.v')),sK1) != s('type.list.list'('type.option.option'('type.closSem.v')),sK6) | $spl3),
  inference(cnf_transformation,[],[f336_D])).
fof(f336_D,plain,(
  s('type.list.list'('type.option.option'('type.closSem.v')),sK1) != s('type.list.list'('type.option.option'('type.closSem.v')),sK6) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f338,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) | $spl5),
  inference(cnf_transformation,[],[f338_D])).
fof(f338_D,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f340,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK8) | $spl7),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK8) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f342,plain,(
  s('type.num.num',sK4) != s('type.num.num',sK9) | $spl9),
  inference(cnf_transformation,[],[f342_D])).
fof(f342_D,plain,(
  s('type.num.num',sK4) != s('type.num.num',sK9) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f344,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) | $spl11),
  inference(cnf_transformation,[],[f344_D])).
fof(f344_D,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f228,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) | s('type.num.num',sK4) != s('type.num.num',sK9) | s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK8) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) | s('type.list.list'('type.option.option'('type.closSem.v')),sK1) != s('type.list.list'('type.option.option'('type.closSem.v')),sK6) | s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10)))),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  (s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) | (s('type.list.list'('type.option.option'('type.closSem.v')),sK1) = s('type.list.list'('type.option.option'('type.closSem.v')),sK6) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) & s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK8) & s('type.num.num',sK4) = s('type.num.num',sK9) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) & (s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) | s('type.list.list'('type.option.option'('type.closSem.v')),sK1) != s('type.list.list'('type.option.option'('type.closSem.v')),sK6) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) | s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK8) | s('type.num.num',sK4) != s('type.num.num',sK9) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f197])).
fof(f197,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5))) = s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) | (s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),X6) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X8) & s('type.num.num',X4) = s('type.num.num',X9) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) & (s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5))) != s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) | s('type.list.list'('type.option.option'('type.closSem.v')),X1) != s('type.list.list'('type.option.option'('type.closSem.v')),X6) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7) | s('type.ffi.ffi_state'(X0),X3) != s('type.ffi.ffi_state'(X0),X8) | s('type.num.num',X4) != s('type.num.num',X9) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10)))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5))) = s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) | (s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),X6) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X8) & s('type.num.num',X4) = s('type.num.num',X9) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) & (s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5))) != s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) | (s('type.list.list'('type.option.option'('type.closSem.v')),X1) != s('type.list.list'('type.option.option'('type.closSem.v')),X6) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) != s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7) | s('type.ffi.ffi_state'(X0),X3) != s('type.ffi.ffi_state'(X0),X8) | s('type.num.num',X4) != s('type.num.num',X9) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5))) = s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) <~> (s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),X6) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X8) & s('type.num.num',X4) = s('type.num.num',X9) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5))) = s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) <=> (s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),X6) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X8) & s('type.num.num',X4) = s('type.num.num',X9) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5))) = s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10))) <=> (s('type.list.list'('type.option.option'('type.closSem.v')),X1) = s('type.list.list'('type.option.option'('type.closSem.v')),X6) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X7) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X8) & s('type.num.num',X4) = s('type.num.num',X9) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X10)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X7),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X8),s('type.ffi.ffi_state'(X6),X9),s('type.num.num',X10),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X11))) = s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X13),s('type.ffi.ffi_state'(X6),X14),s('type.num.num',X15),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X16))) <=> (s('type.list.list'('type.option.option'('type.closSem.v')),X7) = s('type.list.list'('type.option.option'('type.closSem.v')),X12) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X8) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X13) & s('type.ffi.ffi_state'(X6),X9) = s('type.ffi.ffi_state'(X6),X14) & s('type.num.num',X10) = s('type.num.num',X15) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X11) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X16)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X7),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X8),s('type.ffi.ffi_state'(X6),X9),s('type.num.num',X10),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X11))) = s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X13),s('type.ffi.ffi_state'(X6),X14),s('type.num.num',X15),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X16))) <=> (s('type.list.list'('type.option.option'('type.closSem.v')),X7) = s('type.list.list'('type.option.option'('type.closSem.v')),X12) & s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X8) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X13) & s('type.ffi.ffi_state'(X6),X9) = s('type.ffi.ffi_state'(X6),X14) & s('type.num.num',X10) = s('type.num.num',X15) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X11) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X16)))),
  file('thm.closSem.state_|49||49|','thm.closSem.state_|49||49|')).
fof(f6189,plain,(
  $false | ($spl5 | $spl100 | $spl750)),
  inference(subsumption_resolution,[],[f6188,f338])).
fof(f6188,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) | ($spl100 | $spl750)),
  inference(forward_demodulation,[],[f6178,f395])).
fof(f6178,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))))) | $spl750),
  inference(equality_resolution,[],[f6165])).
fof(f5988,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) | $spl738),
  inference(cnf_transformation,[],[f5988_D])).
fof(f5988_D,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) <=> ~$spl738),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl738])])).
fof(f5986,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) | ($spl0 | $spl2 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f435,f335])).
fof(f435,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) | ($spl0 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f434,f339])).
fof(f434,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f333,f341])).
fof(f5992,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | $spl740),
  inference(cnf_transformation,[],[f5992_D])).
fof(f5992_D,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) <=> ~$spl740),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl740])])).
fof(f5989,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl10 | $spl738)),
  inference(backward_demodulation,[],[f343,f5988])).
fof(f6165,plain,(
  ( ! [X15] : (s('type.closSem.state'(sK0),X15) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),X15)))) ) | $spl750),
  inference(cnf_transformation,[],[f6165_D])).
fof(f6165_D,plain,(
  ( ! [X15] : (s('type.closSem.state'(sK0),X15) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),X15)))) ) <=> ~$spl750),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl750])])).
fof(f6164,plain,(
  ( ! [X15] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),X15))) | s('type.closSem.state'(sK0),X15) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))) ) | ($spl84 | $spl100 | $spl740)),
  inference(forward_demodulation,[],[f6136,f395])).
fof(f6136,plain,(
  ( ! [X15] : (s('type.closSem.state'(sK0),X15) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),X15))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),'const.closSem.state_refs_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))))) ) | ($spl84 | $spl740)),
  inference(superposition,[],[f387,f5992])).
fof(f6362,plain,(
  $false | ($spl9 | $spl94 | $spl754)),
  inference(subsumption_resolution,[],[f6361,f342])).
fof(f6361,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK9) | ($spl94 | $spl754)),
  inference(forward_demodulation,[],[f6351,f392])).
fof(f6351,plain,(
  s('type.num.num',sK9) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))))) | $spl754),
  inference(equality_resolution,[],[f6242])).
fof(f6194,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | $spl752),
  inference(cnf_transformation,[],[f6194_D])).
fof(f6194_D,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) <=> ~$spl752),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl752])])).
fof(f6192,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl0 | $spl2 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f6191,f335])).
fof(f6191,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f6190,f339])).
fof(f6242,plain,(
  ( ! [X11] : (s('type.closSem.state'(sK0),X11) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.num.num',sK9) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X11)))) ) | $spl754),
  inference(cnf_transformation,[],[f6242_D])).
fof(f6242_D,plain,(
  ( ! [X11] : (s('type.closSem.state'(sK0),X11) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.num.num',sK9) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X11)))) ) <=> ~$spl754),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl754])])).
fof(f6241,plain,(
  ( ! [X11] : (s('type.num.num',sK9) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X11))) | s('type.closSem.state'(sK0),X11) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))) ) | ($spl80 | $spl94 | $spl752)),
  inference(forward_demodulation,[],[f6212,f392])).
fof(f6212,plain,(
  ( ! [X11] : (s('type.closSem.state'(sK0),X11) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),X11))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))))) ) | ($spl80 | $spl752)),
  inference(superposition,[],[f385,f6194])).
fof(f5985,plain,(
  $false | ($spl11 | $spl102 | $spl736)),
  inference(subsumption_resolution,[],[f5984,f344])).
fof(f5984,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) | ($spl102 | $spl736)),
  inference(forward_demodulation,[],[f5974,f396])).
fof(f5974,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))))) | $spl736),
  inference(equality_resolution,[],[f5947])).
fof(f396,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) | $spl102),
  inference(cnf_transformation,[],[f396_D])).
fof(f396_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f302,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X5)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X43,X27,X44,X28,X45] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X45) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X6),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),X43),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),X27),s('type.ffi.ffi_state'(X6),X44),s('type.num.num',X28),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X45)))))),
  file('thm.closSem.state_|49||49|','thm.closSem.state_code')).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))) ) | $spl78),
  inference(cnf_transformation,[],[f384_D])).
fof(f384_D,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X2))) | s('type.closSem.state'(X0),X1) != s('type.closSem.state'(X0),X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f333,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) | $spl0),
  inference(cnf_transformation,[],[f333_D])).
fof(f333_D,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f343,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) | $spl10),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f227,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) | s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10)))),
  inference(cnf_transformation,[],[f198])).
fof(f341,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK9) | $spl8),
  inference(cnf_transformation,[],[f341_D])).
fof(f341_D,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK9) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f226,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK9) | s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10)))),
  inference(cnf_transformation,[],[f198])).
fof(f339,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK8) | $spl6),
  inference(cnf_transformation,[],[f339_D])).
fof(f339_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK8) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f225,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK8) | s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10)))),
  inference(cnf_transformation,[],[f198])).
fof(f337,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) | $spl4),
  inference(cnf_transformation,[],[f337_D])).
fof(f337_D,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f224,plain,(
  s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7) | s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10)))),
  inference(cnf_transformation,[],[f198])).
fof(f335,plain,(
  s('type.list.list'('type.option.option'('type.closSem.v')),sK1) = s('type.list.list'('type.option.option'('type.closSem.v')),sK6) | $spl2),
  inference(cnf_transformation,[],[f335_D])).
fof(f335_D,plain,(
  s('type.list.list'('type.option.option'('type.closSem.v')),sK1) = s('type.list.list'('type.option.option'('type.closSem.v')),sK6) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f223,plain,(
  s('type.list.list'('type.option.option'('type.closSem.v')),sK1) = s('type.list.list'('type.option.option'('type.closSem.v')),sK6) | s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) = s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10)))),
  inference(cnf_transformation,[],[f198])).
fof(f433,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f432])).
fof(f432,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f431,f335])).
fof(f431,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl1 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f430,f337])).
fof(f430,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl1 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f429,f339])).
fof(f429,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl1 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f428,f341])).
fof(f428,plain,(
  s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | ($spl1 | $spl10)),
  inference(forward_demodulation,[],[f334,f343])).
fof(f5947,plain,(
  ( ! [X0] : (s('type.closSem.state'(sK0),X0) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),X0)))) ) | $spl736),
  inference(cnf_transformation,[],[f5947_D])).
fof(f5947_D,plain,(
  ( ! [X0] : (s('type.closSem.state'(sK0),X0) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),X0)))) ) <=> ~$spl736),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl736])])).
fof(f7244,plain,(
  ( ! [X9] : (s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),X9))) | s('type.closSem.state'(sK0),X9) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5)))) ) | ($spl0 | $spl78 | $spl102)),
  inference(forward_demodulation,[],[f7212,f396])).
fof(f7212,plain,(
  ( ! [X9] : (s('type.closSem.state'(sK0),X9) != s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.num.num',sK4),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK5))) | s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),X9))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),'const.closSem.recordtype|46|state_5'(s('type.list.list'('type.option.option'('type.closSem.v')),sK6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.closSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8),s('type.num.num',sK9),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),sK10)))))) ) | ($spl0 | $spl78)),
  inference(superposition,[],[f384,f333])).
% SZS output end Proof for thm.closSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 155
% Passive clauses: 450
% Generated clauses: 7382
% Final active clauses: 104
% Final passive clauses: 275
% Input formulas: 69
% Initial clauses: 108
% 
% Duplicate literals: 268
% Trivial inequalities: 245
% Fw subsumption resolutions: 376
% Fw demodulations: 959
% Bw demodulations: 26
% Condensations: 380
% Global subsumptions: 4
% 
% Simple tautologies: 130
% Equational tautologies: 37
% Forward subsumptions: 3698
% Fw demodulations to eq. taut.: 7
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 70
% Factoring: 1
% Forward superposition: 1778
% Backward superposition: 2157
% Self superposition: 71
% Equality factoring: 54
% Equality resolution: 28
% 
% Split clauses: 20
% Split components: 44
% SAT solver clauses: 923
% SAT solver unit clauses: 154
% SAT solver binary clauses: 58
% SAT solver learnt clauses: 13
% SAT solver learnt literals: 9
% 
% Sat splits: 20
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 893
% TWLsolver calls for satisfiability: 607
% 
% Memory used [KB]: 6268
% Time elapsed: 0.227 s
% ------------------------------
% ------------------------------
% Success in time 0.87 s
