Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.typeSysProps
% (110411)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1292
% Passive clauses: 7556
% Generated clauses: 29420
% Final active clauses: 1148
% Final passive clauses: 5900
% Input formulas: 261
% Initial clauses: 684
% 
% Duplicate literals: 254
% Trivial inequalities: 171
% Fw subsumption resolutions: 2151
% Fw demodulations: 2559
% Bw demodulations: 122
% 
% Simple tautologies: 315
% Equational tautologies: 27
% Forward subsumptions: 15857
% 
% Binary resolution: 7637
% Factoring: 821
% Forward superposition: 7995
% Backward superposition: 6703
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 24
% 
% Split clauses: 626
% Split components: 1254
% SAT solver clauses: 802
% SAT solver unit clauses: 11
% SAT solver binary clauses: 103
% 
% Sat splits: 626
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 802
% TWLsolver calls for satisfiability: 441
% 
% Memory used [KB]: 14456
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.typeSysProps
% (110426)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 821
% Passive clauses: 3219
% Generated clauses: 20702
% Final active clauses: 754
% Final passive clauses: 2031
% Input formulas: 261
% Initial clauses: 684
% 
% Duplicate literals: 119
% Trivial inequalities: 30
% Fw subsumption resolutions: 829
% Fw demodulations: 1925
% Bw demodulations: 413
% Condensations: 86
% Global subsumptions: 70
% 
% Simple tautologies: 254
% Equational tautologies: 11
% Forward subsumptions: 10532
% Fw demodulations to eq. taut.: 15
% 
% Binary resolution: 2752
% Factoring: 239
% Forward superposition: 5143
% Backward superposition: 5042
% Self superposition: 59
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 35
% Split components: 70
% SAT solver clauses: 4430
% SAT solver unit clauses: 982
% SAT solver binary clauses: 272
% SAT solver learnt clauses: 148
% SAT solver learnt literals: 124
% 
% Sat splits: 35
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 4134
% TWLsolver calls for satisfiability: 8276
% 
% Memory used [KB]: 21875
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.typeSysProps
% (110440)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1742
% Passive clauses: 10860
% Generated clauses: 77610
% Final active clauses: 1445
% Final passive clauses: 3644
% Input formulas: 261
% Initial clauses: 684
% Discarded non-redundant clauses: 706
% 
% Duplicate literals: 384
% Trivial inequalities: 56
% Fw subsumption resolutions: 2317
% Fw demodulations: 17971
% Bw demodulations: 4876
% Condensations: 407
% Global subsumptions: 1192
% 
% Simple tautologies: 129
% Equational tautologies: 81
% Forward subsumptions: 34477
% Fw demodulations to eq. taut.: 6647
% Bw demodulations to eq. taut.: 59
% 
% Unit resulting resolution: 19901
% Factoring: 1
% Forward superposition: 15944
% Backward superposition: 13520
% Self superposition: 126
% Equality factoring: 26
% Equality resolution: 29
% 
% Split clauses: 31
% Split components: 62
% SAT solver clauses: 13899
% SAT solver unit clauses: 8109
% SAT solver binary clauses: 3432
% 
% Sat splits: 31
% Sat splitting refutations: 1782
% 
% Memory used [KB]: 279270
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 19 next slice time: 19
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.typeSysProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.typeSysProps
% SZS output start Proof for thm.typeSysProps
fof(f21407,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1285,f1286,f1287,f1288,f1289,f1291,f1397,f1292,f1371,f1395,f1409,f1410,f1411,f1412,f1551,f1413,f1415,f1416,f1501,f1550,f1739,f1714,f1863,f1864,f1865,f1883,f1892,f1902,f1915,f1930,f2087,f1931,f2091,f1932,f2089,f1886,f1942,f1943,f1944,f1945,f1946,f1948,f2093,f1949,f2028,f2092,f2038,f2098,f2043,f2101,f2052,f2102,f2053,f2054,f2106,f2055,f2107,f2056,f2108,f2057,f2109,f2058,f2110,f2059,f2111,f2060,f2112,f2061,f2113,f2062,f2114,f2063,f2104,f2115,f2123,f2142,f2147,f2155,f2166,f2180,f2197,f1885,f1884,f2202,f1894,f1893,f2210,f1904,f1903,f2221,f1917,f1916,f2235,f1887,f2125,f2124,f2720,f2722,f2726,f2737,f2738,f2742,f1398,f2877,f2878,f3112,f3111,f3110,f3115,f3182,f3181,f2721,f3419,f3343,f3326,f3298,f3297,f3493,f3494,f3495,f3496,f3497,f3498,f3549,f3550,f3551,f3979,f4148,f4352,f4353,f4401,f4402,f4543,f4544,f4629,f4630,f4734,f4735,f4840,f4841,f4846,f4855,f4858,f4850,f4925,f4926,f5033,f5034,f5106,f5108,f5109,f5110,f5113,f5192,f5229,f5346,f5569,f5729,f5910,f7617,f7664,f8330,f8444,f8445,f8558,f8671,f8672,f9012,f9074,f9794,f9795,f9926,f10036,f10037,f10179,f10877,f10878,f11137,f11138,f11588,f11698,f11699,f11857,f11858,f12075,f12164,f12360,f12427,f13043,f13154,f13348,f13524,f13659,f13833,f14025,f14277,f14348,f14419,f14494,f14495,f15070,f15071,f15290,f15814,f15813,f15812,f16106,f16108,f16110,f16112,f16113,f16198,f16321,f16327,f16328,f16453,f16455,f16456,f16457,f16460,f16539,f16598,f17214,f17216,f17218,f17220,f17221,f17306,f17414,f17425,f17426,f17523,f17525,f17527,f17529,f17530,f17615,f17723,f17734,f17735,f18748,f18749,f19328,f19327,f15819,f19694,f19696,f19698,f19700,f19701,f19786,f19881,f20169,f20170,f20399,f20400,f20475,f20476,f20483,f20485,f20487,f1396,f20489,f1400,f20490,f1399,f20575,f1401,f20706,f20667,f20712,f20734,f20740,f20739,f20755,f20787,f20760,f20788,f20770,f20789,f21111,f21143,f2096,f21116,f21144,f2099,f21126,f21145,f2105,f21352,f21406])).
fof(f21406,plain,(
  $false | $spl164),
  inference(equality_resolution,[],[f20740])).
fof(f21352,plain,(
  $false | $spl162),
  inference(equality_resolution,[],[f20739])).
fof(f2105,plain,(
  ( ! [X12,X10,X13,X11] : (sP22(X10,X11,sK56,X12,X13) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X10),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X11),s('type.list.list'('type.ast.pat'),X13),s('type.list.list'('type.ast.t'),X12),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) ) | $spl58),
  inference(cnf_transformation,[],[f2105_D])).
fof(f2105_D,plain,(
  ( ! [X12,X10,X13,X11] : (sP22(X10,X11,sK56,X12,X13) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X10),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X11),s('type.list.list'('type.ast.pat'),X13),s('type.list.list'('type.ast.t'),X12),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f21145,plain,(
  ( ! [X15] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')),X15)) ) | $spl176),
  inference(cnf_transformation,[],[f21145_D])).
fof(f21145_D,plain,(
  ( ! [X15] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')),X15)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f21126,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')),X15) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f2011,f1081])).
fof(f1081,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP22(X0,X1,X4,X3,X2) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) )),
  inference(cnf_transformation,[],[f661])).
fof(f661,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) | (s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),X2) & s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | sP22(X0,X1,X4,X3,X2)) & (((s('type.list.list'('type.ast.pat'),'const.list.NIL_0') != s('type.list.list'('type.ast.pat'),X2) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) | ~p(s(bool,'T_0'))) & ~sP22(X0,X1,X4,X3,X2)) | p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))))),
  inference(flattening,[],[f660])).
fof(f660,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) | ((s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),X2) & s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | sP22(X0,X1,X4,X3,X2))) & (((s('type.list.list'('type.ast.pat'),'const.list.NIL_0') != s('type.list.list'('type.ast.pat'),X2) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) | ~p(s(bool,'T_0'))) & ~sP22(X0,X1,X4,X3,X2)) | p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))))),
  inference(nnf_transformation,[],[f479])).
fof(f479,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) <=> ((s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),X2) & s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | sP22(X0,X1,X4,X3,X2)))),
  inference(definition_folding,[],[f343,f478])).
fof(f478,plain,(
  ! [X0,X1,X4,X3,X2] : (sP22(X0,X1,X4,X3,X2) <=> ? [X5,X6,X7,X8,X9,X10] : (s('type.list.list'('type.ast.pat'),X2) = s('type.list.list'('type.ast.pat'),'const.list.CONS_2'(s('type.ast.pat',X5),s('type.list.list'('type.ast.pat'),X6))) & s('type.list.list'('type.ast.t'),X3) = s('type.list.list'('type.ast.t'),'const.list.CONS_2'(s('type.ast.t',X7),s('type.list.list'('type.ast.t'),X8))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X9))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X9)))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X6),s('type.list.list'('type.ast.t'),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP22])])).
fof(f343,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) <=> ((s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),X2) & s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | ? [X5,X6,X7,X8,X9,X10] : (s('type.list.list'('type.ast.pat'),X2) = s('type.list.list'('type.ast.pat'),'const.list.CONS_2'(s('type.ast.pat',X5),s('type.list.list'('type.ast.pat'),X6))) & s('type.list.list'('type.ast.t'),X3) = s('type.list.list'('type.ast.t'),'const.list.CONS_2'(s('type.ast.t',X7),s('type.list.list'('type.ast.t'),X8))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X9))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X9)))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X6),s('type.list.list'('type.ast.t'),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10)))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X16,X17,X18,X19,X20] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X18),s('type.list.list'('type.ast.t'),X19),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X20)))) <=> ((s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),X18) & s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X19) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X20) & p(s(bool,'T_0'))) | ? [X12,X8,X13,X9,X21,X22] : (s('type.list.list'('type.ast.pat'),X18) = s('type.list.list'('type.ast.pat'),'const.list.CONS_2'(s('type.ast.pat',X12),s('type.list.list'('type.ast.pat'),X8))) & s('type.list.list'('type.ast.t'),X19) = s('type.list.list'('type.ast.t'),'const.list.CONS_2'(s('type.ast.t',X13),s('type.list.list'('type.ast.t'),X9))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X20) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X21))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.ast.pat',X12),s('type.ast.t',X13),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X21)))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X8),s('type.list.list'('type.ast.t'),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22)))))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.typeSystem.type_p_cases|split|1')).
fof(f2011,plain,(
  ( ! [X157] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)),X157)) )),
  inference(resolution,[],[f767,f895])).
fof(f895,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f559])).
fof(f559,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP10(X0,X1)))),
  inference(flattening,[],[f558])).
fof(f558,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP10(X0,X1)))),
  inference(nnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1] : (sP10(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f767,plain,(
  p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))),
  inference(cnf_transformation,[],[f512])).
fof(f512,plain,(
  p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) & (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) & ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK52,sK53,sK54,sK55,sK56,sK57])],[f396])).
fof(f396,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) & ? [X5] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X5)))) & ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4),s('type.typeSystem.tenv_val',X5))))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) => ! [X5] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X5)))) => p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4),s('type.typeSystem.tenv_val',X5))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9,X10] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X6),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X7),s('type.list.list'('type.ast.pat'),X8),s('type.list.list'('type.ast.t'),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10)))) => ! [X11] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X11)))) => p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10),s('type.typeSystem.tenv_val',X11))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9,X10] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X6),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X7),s('type.list.list'('type.ast.pat'),X8),s('type.list.list'('type.ast.t'),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10)))) => ! [X11] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X11)))) => p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10),s('type.typeSystem.tenv_val',X11))))))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.typeSysProps.type_p_bvl|split|1')).
fof(f2099,plain,(
  ( ! [X2,X0,X3,X1] : (sP22(X0,X1,X2,sK55,X3) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X3),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) ) | $spl50),
  inference(cnf_transformation,[],[f2099_D])).
fof(f2099_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP22(X0,X1,X2,sK55,X3) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X3),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f21144,plain,(
  ( ! [X0] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)),X0)) ) | $spl174),
  inference(cnf_transformation,[],[f21144_D])).
fof(f21144_D,plain,(
  ( ! [X0] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)),X0)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f21116,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)),X0) | sP22(X1,X2,X3,sK55,X4) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X1),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X2),s('type.list.list'('type.ast.pat'),X4),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X3))))) )),
  inference(superposition,[],[f2011,f1080])).
fof(f1080,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP22(X0,X1,X4,X3,X2) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X3) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) )),
  inference(cnf_transformation,[],[f661])).
fof(f2096,plain,(
  ( ! [X2,X0,X3,X1] : (sP22(X0,X1,X2,X3,sK54) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) ) | $spl46),
  inference(cnf_transformation,[],[f2096_D])).
fof(f2096_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP22(X0,X1,X2,X3,sK54) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f21143,plain,(
  ( ! [X0] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)),X0)) ) | $spl172),
  inference(cnf_transformation,[],[f21143_D])).
fof(f21143_D,plain,(
  ( ! [X0] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)),X0)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f21111,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)),X0) | sP22(X1,X2,X3,X4,sK54) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X1),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X2),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X3))))) )),
  inference(superposition,[],[f2011,f1079])).
fof(f1079,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP22(X0,X1,X4,X3,X2) | s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),X2) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) )),
  inference(cnf_transformation,[],[f661])).
fof(f20789,plain,(
  ( ! [X15] : (sP10(X15,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) ) | $spl170),
  inference(cnf_transformation,[],[f20789_D])).
fof(f20789_D,plain,(
  ( ! [X15] : (sP10(X15,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f20770,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'))) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f2010,f1081])).
fof(f2010,plain,(
  ( ! [X156] : (sP10(X156,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f767,f894])).
fof(f894,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f559])).
fof(f20788,plain,(
  ( ! [X0] : (sP10(X0,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) | $spl168),
  inference(cnf_transformation,[],[f20788_D])).
fof(f20788_D,plain,(
  ( ! [X0] : (sP10(X0,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f20760,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))) | sP22(X1,X2,X3,sK55,X4) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X1),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X2),s('type.list.list'('type.ast.pat'),X4),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X3))))) )),
  inference(superposition,[],[f2010,f1080])).
fof(f20787,plain,(
  ( ! [X0] : (sP10(X0,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) | $spl166),
  inference(cnf_transformation,[],[f20787_D])).
fof(f20787_D,plain,(
  ( ! [X0] : (sP10(X0,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f20755,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))) | sP22(X1,X2,X3,X4,sK54) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X1),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X2),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),X4),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X3))))) )),
  inference(superposition,[],[f2010,f1079])).
fof(f20739,plain,(
  ( ! [X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),sK54)))) ) | $spl162),
  inference(cnf_transformation,[],[f20739_D])).
fof(f20739_D,plain,(
  ( ! [X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),sK54)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f20740,plain,(
  ( ! [X0] : (s('type.ast.t',X0) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK55),s('type.ast.tctor','const.ast.TC_tup_0')))) ) | $spl164),
  inference(cnf_transformation,[],[f20740_D])).
fof(f20740_D,plain,(
  ( ! [X0] : (s('type.ast.t',X0) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK55),s('type.ast.tctor','const.ast.TC_tup_0')))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f20734,plain,(
  ( ! [X0,X1] : (s('type.ast.t',X0) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK55),s('type.ast.tctor','const.ast.TC_tup_0'))) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),sK54)))) )),
  inference(resolution,[],[f20723,f1934])).
fof(f1934,plain,(
  ( ! [X10,X11] : (sP30(sK52,sK53,sK56,X10,X11) | s('type.ast.t',X10) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK55),s('type.ast.tctor','const.ast.TC_tup_0'))) | s('type.ast.pat',X11) != s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),sK54)))) )),
  inference(resolution,[],[f767,f1100])).
fof(f1100,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP30(X0,X1,X2,X3,X4) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X5),s('type.list.list'('type.ast.t'),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))) | s('type.ast.t',X3) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X6),s('type.ast.tctor','const.ast.TC_tup_0'))) | s('type.ast.pat',X4) != s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X5)))) )),
  inference(cnf_transformation,[],[f668])).
fof(f668,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP30(X0,X1,X2,X3,X4) | (s('type.ast.pat',X4) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),sK106(X4,X3,X2,X1,X0)))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK107(X4,X3,X2,X1,X0)),s('type.ast.tctor','const.ast.TC_tup_0'))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),sK106(X4,X3,X2,X1,X0)),s('type.list.list'('type.ast.t'),sK107(X4,X3,X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))) & (! [X5,X6] : (s('type.ast.pat',X4) != s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X5))) | s('type.ast.t',X3) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X6),s('type.ast.tctor','const.ast.TC_tup_0'))) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X5),s('type.list.list'('type.ast.t'),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) | sP30(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK106,sK107])],[f667])).
fof(f667,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP30(X0,X1,X2,X3,X4) | ? [X7,X8] : (s('type.ast.pat',X4) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X7))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X8),s('type.ast.tctor','const.ast.TC_tup_0'))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X7),s('type.list.list'('type.ast.t'),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))) & (! [X5,X6] : (s('type.ast.pat',X4) != s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X5))) | s('type.ast.t',X3) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X6),s('type.ast.tctor','const.ast.TC_tup_0'))) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X5),s('type.list.list'('type.ast.t'),X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) | sP30(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f666])).
fof(f666,plain,(
  ! [X0,X1,X4,X3,X2] : ((~sP30(X0,X1,X4,X3,X2) | ? [X7,X8] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X7))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X8),s('type.ast.tctor','const.ast.TC_tup_0'))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X7),s('type.list.list'('type.ast.t'),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))))) & (! [X7,X8] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X7))) | s('type.ast.t',X3) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X8),s('type.ast.tctor','const.ast.TC_tup_0'))) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X7),s('type.list.list'('type.ast.t'),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) | sP30(X0,X1,X4,X3,X2)))),
  inference(nnf_transformation,[],[f487])).
fof(f487,plain,(
  ! [X0,X1,X4,X3,X2] : (sP30(X0,X1,X4,X3,X2) <=> ? [X7,X8] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X7))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X8),s('type.ast.tctor','const.ast.TC_tup_0'))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X7),s('type.list.list'('type.ast.t'),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP30])])).
fof(f20723,plain,(
  ( ! [X17,X15,X16] : (~sP30(sK52,X15,sK56,X16,X17)) )),
  inference(resolution,[],[f20480,f1095])).
fof(f1095,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP31(X0,X1,X2,X3,X4) | ~sP30(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f665])).
fof(f665,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP31(X0,X1,X2,X3,X4) | sP29(X0,X3,X2,X4) | sP24(X2,X3,X4) | sP25(X2,X3,X4) | sP26(X2,X3,X4) | sP27(X2,X3,X4) | sP28(X2,X3,X4) | sP23(X1,X0,X2,X3,X4) | sP30(X0,X1,X2,X3,X4) | (s('type.ast.pat',X4) = s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',sK104(X4,X3,X2,X1,X0)))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',sK105(X4,X3,X2,X1,X0)))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',sK104(X4,X3,X2,X1,X0)),s('type.ast.t',sK105(X4,X3,X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))) & ((~sP29(X0,X3,X2,X4) & ~sP24(X2,X3,X4) & ~sP25(X2,X3,X4) & ~sP26(X2,X3,X4) & ~sP27(X2,X3,X4) & ~sP28(X2,X3,X4) & ~sP23(X1,X0,X2,X3,X4) & ~sP30(X0,X1,X2,X3,X4) & ! [X5,X6] : (s('type.ast.pat',X4) != s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X5))) | s('type.ast.t',X3) != s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X6))) | ~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))) | sP31(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK104,sK105])],[f664])).
fof(f664,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP31(X0,X1,X2,X3,X4) | sP29(X0,X3,X2,X4) | sP24(X2,X3,X4) | sP25(X2,X3,X4) | sP26(X2,X3,X4) | sP27(X2,X3,X4) | sP28(X2,X3,X4) | sP23(X1,X0,X2,X3,X4) | sP30(X0,X1,X2,X3,X4) | ? [X7,X8] : (s('type.ast.pat',X4) = s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X7))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X8))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X7),s('type.ast.t',X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))) & ((~sP29(X0,X3,X2,X4) & ~sP24(X2,X3,X4) & ~sP25(X2,X3,X4) & ~sP26(X2,X3,X4) & ~sP27(X2,X3,X4) & ~sP28(X2,X3,X4) & ~sP23(X1,X0,X2,X3,X4) & ~sP30(X0,X1,X2,X3,X4) & ! [X5,X6] : (s('type.ast.pat',X4) != s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X5))) | s('type.ast.t',X3) != s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X6))) | ~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))) | sP31(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f663])).
fof(f663,plain,(
  ! [X0,X1,X4,X3,X2] : ((~sP31(X0,X1,X4,X3,X2) | sP29(X0,X3,X4,X2) | sP24(X4,X3,X2) | sP25(X4,X3,X2) | sP26(X4,X3,X2) | sP27(X4,X3,X2) | sP28(X4,X3,X2) | sP23(X1,X0,X4,X3,X2) | sP30(X0,X1,X4,X3,X2) | ? [X5,X6] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X5))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X6))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))))) & ((~sP29(X0,X3,X4,X2) & ~sP24(X4,X3,X2) & ~sP25(X4,X3,X2) & ~sP26(X4,X3,X2) & ~sP27(X4,X3,X2) & ~sP28(X4,X3,X2) & ~sP23(X1,X0,X4,X3,X2) & ~sP30(X0,X1,X4,X3,X2) & ! [X5,X6] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X5))) | s('type.ast.t',X3) != s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X6))) | ~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))))) | sP31(X0,X1,X4,X3,X2)))),
  inference(flattening,[],[f662])).
fof(f662,plain,(
  ! [X0,X1,X4,X3,X2] : ((~sP31(X0,X1,X4,X3,X2) | (sP29(X0,X3,X4,X2) | sP24(X4,X3,X2) | sP25(X4,X3,X2) | sP26(X4,X3,X2) | sP27(X4,X3,X2) | sP28(X4,X3,X2) | sP23(X1,X0,X4,X3,X2) | sP30(X0,X1,X4,X3,X2) | ? [X5,X6] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X5))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X6))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))))) & ((~sP29(X0,X3,X4,X2) & ~sP24(X4,X3,X2) & ~sP25(X4,X3,X2) & ~sP26(X4,X3,X2) & ~sP27(X4,X3,X2) & ~sP28(X4,X3,X2) & ~sP23(X1,X0,X4,X3,X2) & ~sP30(X0,X1,X4,X3,X2) & ! [X5,X6] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X5))) | s('type.ast.t',X3) != s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X6))) | ~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))))) | sP31(X0,X1,X4,X3,X2)))),
  inference(nnf_transformation,[],[f488])).
fof(f488,plain,(
  ! [X0,X1,X4,X3,X2] : (sP31(X0,X1,X4,X3,X2) <=> (sP29(X0,X3,X4,X2) | sP24(X4,X3,X2) | sP25(X4,X3,X2) | sP26(X4,X3,X2) | sP27(X4,X3,X2) | sP28(X4,X3,X2) | sP23(X1,X0,X4,X3,X2) | sP30(X0,X1,X4,X3,X2) | ? [X5,X6] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X5))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X6))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP31])])).
fof(f20480,plain,(
  ( ! [X14,X15,X16] : (~sP31(sK52,X14,sK56,X15,X16)) )),
  inference(resolution,[],[f1549,f1138])).
fof(f1138,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X2),s('type.ast.t',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) | ~sP31(X0,X1,X4,X3,X2)) )),
  inference(cnf_transformation,[],[f690])).
fof(f690,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X2),s('type.ast.t',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) | sP31(X0,X1,X4,X3,X2)) & (~sP31(X0,X1,X4,X3,X2) | p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X2),s('type.ast.t',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))))),
  inference(nnf_transformation,[],[f489])).
fof(f489,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X2),s('type.ast.t',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) <=> sP31(X0,X1,X4,X3,X2))),
  inference(definition_folding,[],[f344,f488,f487,f486,f485,f484,f483,f482,f481,f480])).
fof(f480,plain,(
  ! [X1,X0,X4,X3,X2] : (sP23(X1,X0,X4,X3,X2) <=> ? [X9,X10,X11,X12,X13,X14] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.SOME_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9))),s('type.list.list'('type.ast.pat'),X10))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X14),s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X13))))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X14)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),X14))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.list.list'('type.string.char')),X12))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X10),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.ast.t','type.ast.t'),'const.typeSystem.type_subst_1'(s('type.finite_map.fmap'('type.list.list'('type.string.char'),'type.ast.t'),'const.alist.alist_to_fmap_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.list.list'('type.ast.t')),'const.pair.,_2'(s('type.list.list'('type.list.list'('type.string.char')),X12),s('type.list.list'('type.ast.t'),X14))))))))),s('type.list.list'('type.ast.t'),X11))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) & s('type.option.option'('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))),'const.option.SOME_1'(s('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')),'const.pair.,_2'(s('type.list.list'('type.list.list'('type.string.char')),X12),s('type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.ast.t'),X11),s('type.semanticPrimitives.tid_or_exn',X13))))))) = s('type.option.option'('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))),'const.semanticPrimitives.lookup_alist_mod_env_2'(s('type.ast.id'('type.list.list'('type.string.char')),X9),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP23])])).
fof(f481,plain,(
  ! [X4,X3,X2] : (sP24(X4,X3,X2) <=> ? [X19] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X19))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP24])])).
fof(f482,plain,(
  ! [X4,X3,X2] : (sP25(X4,X3,X2) <=> ? [X18] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X18))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP25])])).
fof(f483,plain,(
  ! [X4,X3,X2] : (sP26(X4,X3,X2) <=> ? [X17] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X17))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP26])])).
fof(f484,plain,(
  ! [X4,X3,X2] : (sP27(X4,X3,X2) <=> ? [X16] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X16))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP27])])).
fof(f485,plain,(
  ! [X4,X3,X2] : (sP28(X4,X3,X2) <=> ? [X15] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X15))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP28])])).
fof(f486,plain,(
  ! [X0,X3,X4,X2] : (sP29(X0,X3,X4,X2) <=> ? [X20] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pvar_1'(s('type.list.list'('type.string.char'),X20))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X20),s('type.ast.t',X3))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'))) & p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP29])])).
fof(f344,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X2),s('type.ast.t',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) <=> (? [X20] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pvar_1'(s('type.list.list'('type.string.char'),X20))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X20),s('type.ast.t',X3))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'))) & p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X3))))) | ? [X19] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X19))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | ? [X18] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X18))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | ? [X17] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X17))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | ? [X16] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X16))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | ? [X15] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X15))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0'))) | ? [X9,X10,X11,X12,X13,X14] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.SOME_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9))),s('type.list.list'('type.ast.pat'),X10))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X14),s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X13))))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X14)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),X14))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.list.list'('type.string.char')),X12))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X10),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.ast.t','type.ast.t'),'const.typeSystem.type_subst_1'(s('type.finite_map.fmap'('type.list.list'('type.string.char'),'type.ast.t'),'const.alist.alist_to_fmap_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.list.list'('type.ast.t')),'const.pair.,_2'(s('type.list.list'('type.list.list'('type.string.char')),X12),s('type.list.list'('type.ast.t'),X14))))))))),s('type.list.list'('type.ast.t'),X11))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) & s('type.option.option'('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))),'const.option.SOME_1'(s('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')),'const.pair.,_2'(s('type.list.list'('type.list.list'('type.string.char')),X12),s('type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.ast.t'),X11),s('type.semanticPrimitives.tid_or_exn',X13))))))) = s('type.option.option'('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))),'const.semanticPrimitives.lookup_alist_mod_env_2'(s('type.ast.id'('type.list.list'('type.string.char')),X9),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1)))) | ? [X7,X8] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X7))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X8),s('type.ast.tctor','const.ast.TC_tup_0'))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X7),s('type.list.list'('type.ast.t'),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) | ? [X5,X6] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X5))) & s('type.ast.t',X3) = s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X6))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X5),s('type.ast.t',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X23,X24,X25,X26,X27] : (p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X23),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X24),s('type.ast.pat',X25),s('type.ast.t',X26),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27)))) <=> (? [X28] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Pvar_1'(s('type.list.list'('type.string.char'),X28))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X28),s('type.ast.t',X26))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'))) & p(s(bool,happ(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X23),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.ast.t',X26))))) | ? [X28] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X28))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X26) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27) & p(s(bool,'T_0'))) | ? [X29] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X29))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X26) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27) & p(s(bool,'T_0'))) | ? [X30] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X30))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X26) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27) & p(s(bool,'T_0'))) | ? [X31] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X31))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X26) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27) & p(s(bool,'T_0'))) | ? [X31] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X31))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X26) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27) & p(s(bool,'T_0'))) | ? [X32,X8,X9,X33,X34,X35] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.SOME_1'(s('type.ast.id'('type.list.list'('type.string.char')),X32))),s('type.list.list'('type.ast.pat'),X8))) & s('type.ast.t',X26) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X35),s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X34))))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X23),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X35)))) & s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),X35))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.list.list'('type.string.char')),X33))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X23),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X24),s('type.list.list'('type.ast.pat'),X8),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.ast.t','type.ast.t'),'const.typeSystem.type_subst_1'(s('type.finite_map.fmap'('type.list.list'('type.string.char'),'type.ast.t'),'const.alist.alist_to_fmap_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.list.list'('type.ast.t')),'const.pair.,_2'(s('type.list.list'('type.list.list'('type.string.char')),X33),s('type.list.list'('type.ast.t'),X35))))))))),s('type.list.list'('type.ast.t'),X9))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27)))) & s('type.option.option'('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))),'const.option.SOME_1'(s('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')),'const.pair.,_2'(s('type.list.list'('type.list.list'('type.string.char')),X33),s('type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'),'const.pair.,_2'(s('type.list.list'('type.ast.t'),X9),s('type.semanticPrimitives.tid_or_exn',X34))))))) = s('type.option.option'('type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))),'const.semanticPrimitives.lookup_alist_mod_env_2'(s('type.ast.id'('type.list.list'('type.string.char')),X32),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X24)))) | ? [X8,X9] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Pcon_2'(s('type.option.option'('type.ast.id'('type.list.list'('type.string.char'))),'const.option.NONE_0'),s('type.list.list'('type.ast.pat'),X8))) & s('type.ast.t',X26) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X9),s('type.ast.tctor','const.ast.TC_tup_0'))) & p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X23),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X24),s('type.list.list'('type.ast.pat'),X8),s('type.list.list'('type.ast.t'),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27))))) | ? [X12,X13] : (s('type.ast.pat',X25) = s('type.ast.pat','const.ast.Pref_1'(s('type.ast.pat',X12))) & s('type.ast.t',X26) = s('type.ast.t','const.ast.Tref_1'(s('type.ast.t',X13))) & p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X23),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X24),s('type.ast.pat',X12),s('type.ast.t',X13),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X27)))))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.typeSystem.type_p_cases|split|0')).
fof(f1549,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X0),s('type.ast.pat',X1),s('type.ast.t',X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f1407,f768])).
fof(f768,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))))),
  inference(cnf_transformation,[],[f512])).
fof(f1407,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X0),s('type.ast.pat',X1),s('type.ast.t',X2),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(resolution,[],[f769,f770])).
fof(f770,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4),s('type.typeSystem.tenv_val',X5)))))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X5)))) | ~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X2),s('type.ast.t',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X2),s('type.ast.t',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) | ! [X5] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X5)))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4),s('type.typeSystem.tenv_val',X5))))))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.ast.pat',X2),s('type.ast.t',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) => ! [X5] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X5)))) => p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4),s('type.typeSystem.tenv_val',X5))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7,X12,X13,X10] : (p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',X6),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X7),s('type.ast.pat',X12),s('type.ast.t',X13),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10)))) => ! [X11] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X11)))) => p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10),s('type.typeSystem.tenv_val',X11))))))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.typeSysProps.type_p_bvl|split|0')).
fof(f769,plain,(
  ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))))),
  inference(cnf_transformation,[],[f512])).
fof(f20712,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X25),s('type.ast.pat',X26),s('type.ast.t',X27),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'))))) ) | $spl160),
  inference(cnf_transformation,[],[f20712_D])).
fof(f20712_D,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X25),s('type.ast.pat',X26),s('type.ast.t',X27),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f20667,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X25),s('type.ast.pat',X26),s('type.ast.t',X27),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1549,f1081])).
fof(f20706,plain,(
  ( ! [X30,X28,X31,X32] : (~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X30),s('type.ast.pat',X31),s('type.ast.t',X32),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) | $spl158),
  inference(cnf_transformation,[],[f20706_D])).
fof(f20706_D,plain,(
  ( ! [X30,X28,X31,X32] : (~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X30),s('type.ast.pat',X31),s('type.ast.t',X32),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f1401,plain,(
  ( ! [X21] : (p(s(bool,X21))) ) | $spl12),
  inference(cnf_transformation,[],[f1401_D])).
fof(f1401_D,plain,(
  ( ! [X21] : (p(s(bool,X21))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f20575,plain,(
  ( ! [X422,X424,X426,X423,X425] : (p(s(bool,X422)) | ~p(s(bool,X423)) | s(bool,X423) != s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X424),s('type.ast.pat',X425),s('type.ast.t',X426),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f1549,f1274])).
fof(f1274,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f799])).
fof(f799,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f516])).
fof(f516,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f515])).
fof(f515,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f404])).
fof(f404,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X47,X66,X46,X67] : ((s(bool,X47) = s(bool,X66) & (p(s(bool,X66)) => s(bool,X46) = s(bool,X67))) => ((p(s(bool,X47)) => p(s(bool,X46))) <=> (p(s(bool,X66)) => p(s(bool,X67)))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.IMP_CONG')).
fof(f1399,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X16))) ) | $spl8),
  inference(cnf_transformation,[],[f1399_D])).
fof(f1399_D,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X16))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f20490,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X68),s('type.ast.pat',X69),s('type.ast.t',X70),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))) | s(bool,X66) != s(bool,X71)) )),
  inference(resolution,[],[f1549,f802])).
fof(f802,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f516])).
fof(f1400,plain,(
  ( ! [X22] : (~p(s(bool,X22))) ) | $spl10),
  inference(cnf_transformation,[],[f1400_D])).
fof(f1400_D,plain,(
  ( ! [X22] : (~p(s(bool,X22))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f20489,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (p(s(bool,X60)) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X63),s('type.ast.pat',X64),s('type.ast.t',X65),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f1549,f801])).
fof(f801,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f516])).
fof(f1396,plain,(
  ( ! [X11,X9] : (s(bool,X9) != s(bool,X11) | p(s(bool,X9))) ) | $spl2),
  inference(cnf_transformation,[],[f1396_D])).
fof(f1396_D,plain,(
  ( ! [X11,X9] : (s(bool,X9) != s(bool,X11) | p(s(bool,X9))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f20487,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X54),s('type.ast.pat',X55),s('type.ast.t',X56),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f1549,f800])).
fof(f800,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f516])).
fof(f20485,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X41),s('type.ast.pat',X42),s('type.ast.t',X43),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))) | s(bool,X39) != s(bool,X44)) )),
  inference(resolution,[],[f1549,f798])).
fof(f798,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f516])).
fof(f20483,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.typeSystem.type_p_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X30),s('type.ast.pat',X31),s('type.ast.t',X32),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f1549,f796])).
fof(f796,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f516])).
fof(f20476,plain,(
  ( ! [X21,X20] : (sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21))) ) | $spl156),
  inference(cnf_transformation,[],[f20476_D])).
fof(f20476_D,plain,(
  ( ! [X21,X20] : (sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f20475,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f20474,f1314])).
fof(f1314,plain,(
  ( ! [X76,X75] : (sP2(X75,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X76)) )),
  inference(resolution,[],[f768,f829])).
fof(f829,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f530])).
fof(f530,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f529])).
fof(f529,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f528])).
fof(f528,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f20474,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP2(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f20446,f775])).
fof(f775,plain,(
  ( ! [X0,X1] : (s('type.typeSystem.tenv_val',X1) = s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',X1)))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : s('type.typeSystem.tenv_val',X1) = s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',X1)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6,X77] : s('type.typeSystem.tenv_val',X77) = s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',X77)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.typeSystem.bind_var_list_def|split|0')).
fof(f20446,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP2(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1437,f1081])).
fof(f1437,plain,(
  ( ! [X70,X71] : (~sP2(X70,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X71) | sP3(X70,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X71) | p(s(bool,X71))) )),
  inference(resolution,[],[f769,f826])).
fof(f826,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f526])).
fof(f526,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f525])).
fof(f525,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f453])).
fof(f453,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f20400,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21))) ) | $spl154),
  inference(cnf_transformation,[],[f20400_D])).
fof(f20400_D,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f20399,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f20398,f1313])).
fof(f1313,plain,(
  ( ! [X74,X73] : (sP2(X73,X74,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f828])).
fof(f828,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f530])).
fof(f20398,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP2(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f20370,f775])).
fof(f20370,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP2(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1436,f1081])).
fof(f1436,plain,(
  ( ! [X68,X69] : (~sP2(X68,X69,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | sP3(X68,X69,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X69))) )),
  inference(resolution,[],[f769,f826])).
fof(f20170,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X21))) ) | $spl152),
  inference(cnf_transformation,[],[f20170_D])).
fof(f20170_D,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X21))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f20169,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f20168,f1300])).
fof(f1300,plain,(
  ( ! [X47,X48] : (sP0(X47,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X48)) )),
  inference(resolution,[],[f768,f812])).
fof(f812,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f521])).
fof(f521,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f520])).
fof(f520,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f449])).
fof(f449,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f20168,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP0(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f20147,f775])).
fof(f20147,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP0(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1421,f1081])).
fof(f1421,plain,(
  ( ! [X39,X38] : (~sP0(X38,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X39) | sP1(X38,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X39) | ~p(s(bool,X39))) )),
  inference(resolution,[],[f769,f809])).
fof(f809,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f519])).
fof(f519,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f518])).
fof(f518,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f19881,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X4),s('type.num.num',X5),s('type.ast.t',X6),s('type.typeSystem.tenv_val',sK57)))))) | s(bool,X8) != s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X6))) | ~p(s(bool,X8))) ) | $spl150),
  inference(cnf_transformation,[],[f19881_D])).
fof(f19881_D,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X4),s('type.num.num',X5),s('type.ast.t',X6),s('type.typeSystem.tenv_val',sK57)))))) | s(bool,X8) != s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X6))) | ~p(s(bool,X8))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f19786,plain,(
  ( ! [X436,X432,X434,X433,X435] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X432),s('type.num.num',X433),s('type.ast.t',X434),s('type.typeSystem.tenv_val',sK57)))))) | p(s(bool,X435)) | ~p(s(bool,X436)) | s(bool,X436) != s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X433),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X434)))) )),
  inference(resolution,[],[f1282,f1274])).
fof(f1282,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X2)))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.ast.t',X2),s('type.typeSystem.tenv_val',sK57))))))) )),
  inference(resolution,[],[f768,f778])).
fof(f778,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.ast.t',X2),s('type.typeSystem.tenv_val',X3)))))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X3)))) | ~p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X3))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X2))))) )),
  inference(cnf_transformation,[],[f514])).
fof(f514,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.ast.t',X2),s('type.typeSystem.tenv_val',X3)))))) | (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X3))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X2)))) & p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X3)))))) & (~p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X3))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X2)))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X3)))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.ast.t',X2),s('type.typeSystem.tenv_val',X3))))))))),
  inference(flattening,[],[f513])).
fof(f513,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.ast.t',X2),s('type.typeSystem.tenv_val',X3)))))) | (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X3))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X2)))) & p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X3)))))) & ((~p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X3))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X2)))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X3))))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.ast.t',X2),s('type.typeSystem.tenv_val',X3))))))))),
  inference(nnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.ast.t',X2),s('type.typeSystem.tenv_val',X3)))))) <=> (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X3))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X2)))) & p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X3))))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X47,X6,X13,X10] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X47),s('type.num.num',X6),s('type.ast.t',X13),s('type.typeSystem.tenv_val',X10)))))) <=> (p(s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X10))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X13)))) & p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X10))))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.typeSoundInvariants.tenv_val_ok_def|split|2')).
fof(f19701,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X46),s('type.num.num',X47),s('type.ast.t',X48),s('type.typeSystem.tenv_val',sK57)))))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X47),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X48))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f1282,f802])).
fof(f19700,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X40),s('type.num.num',X41),s('type.ast.t',X42),s('type.typeSystem.tenv_val',sK57)))))) | p(s(bool,X43)) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X41),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X42)))) )),
  inference(resolution,[],[f1282,f801])).
fof(f19698,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X28),s('type.num.num',X29),s('type.ast.t',X30),s('type.typeSystem.tenv_val',sK57)))))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X30)))) )),
  inference(resolution,[],[f1282,f800])).
fof(f19696,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X16),s('type.num.num',X17),s('type.ast.t',X18),s('type.typeSystem.tenv_val',sK57)))))) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X18))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f1282,f798])).
fof(f19694,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X4),s('type.num.num',X5),s('type.ast.t',X6),s('type.typeSystem.tenv_val',sK57)))))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.typeSystem.check_freevars_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'),s('type.ast.t',X6)))) )),
  inference(resolution,[],[f1282,f796])).
fof(f15819,plain,(
  ( ! [X37,X38] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X37))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X38)))) ) | $spl132),
  inference(cnf_transformation,[],[f15819_D])).
fof(f15819_D,plain,(
  ( ! [X37,X38] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X37))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X38)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f19327,plain,(
  ( ! [X52,X53,X51] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X53))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X52))) | ~p(s(bool,X51))) )),
  inference(duplicate_literal_removal,[],[f19285])).
fof(f19285,plain,(
  ( ! [X52,X53,X51] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X53))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X52))) | ~p(s(bool,X51)) | ~p(s(bool,X51))) )),
  inference(superposition,[],[f19081,f987])).
fof(f987,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f617])).
fof(f617,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f19081,plain,(
  ( ! [X14,X12,X13] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X14))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X12),s('type.list.list'(bool),X13))) | ~p(s(bool,X12))) )),
  inference(forward_demodulation,[],[f1544,f1486])).
fof(f1486,plain,(
  s(bool,'F_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))),
  inference(resolution,[],[f769,f935])).
fof(f935,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f585])).
fof(f585,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f584])).
fof(f584,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X13))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.IMP_F_EQ_F')).
fof(f1544,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X12),s('type.list.list'(bool),X13))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))),s('type.list.list'(bool),X14)))) )),
  inference(superposition,[],[f769,f1247])).
fof(f1247,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,X3) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4)))) )),
  inference(cnf_transformation,[],[f759])).
fof(f759,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4))) | (s(X0,X1) = s(X0,X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))) & (s(X0,X1) != s(X0,X3) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),X4) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f758])).
fof(f758,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4))) | (s(X0,X1) = s(X0,X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))) & ((s(X0,X1) != s(X0,X3) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),X4)) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4)))))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4))) <=> (s(X0,X1) = s(X0,X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4)))),
  inference(flattening,[],[f357])).
fof(f357,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4))) <=> (s(X0,X1) = s(X0,X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4)))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X37] : ! [X23,X24,X80,X81] : (s('type.list.list'(X37),'const.list.CONS_2'(s(X37,X23),s('type.list.list'(X37),X24))) = s('type.list.list'(X37),'const.list.CONS_2'(s(X37,X80),s('type.list.list'(X37),X81))) <=> (s(X37,X23) = s(X37,X80) & s('type.list.list'(X37),X24) = s('type.list.list'(X37),X81)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.list.list_|49||49|')).
fof(f19328,plain,(
  ( ! [X50,X48,X49] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X50))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X49))) | ~p(s(bool,X48))) )),
  inference(duplicate_literal_removal,[],[f19284])).
fof(f19284,plain,(
  ( ! [X50,X48,X49] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X50))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X49))) | ~p(s(bool,X48)) | ~p(s(bool,X48))) )),
  inference(superposition,[],[f19081,f985])).
fof(f985,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f616])).
fof(f616,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f18749,plain,(
  ( ! [X26,X27,X25] : (s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',sK57),s('type.list.list'('type.typeSystem.tenv_val'),X27))) != s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',X25),s('type.list.list'('type.typeSystem.tenv_val'),X26)))) ) | $spl148),
  inference(cnf_transformation,[],[f18749_D])).
fof(f18749_D,plain,(
  ( ! [X26,X27,X25] : (s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',sK57),s('type.list.list'('type.typeSystem.tenv_val'),X27))) != s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',X25),s('type.list.list'('type.typeSystem.tenv_val'),X26)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f18748,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',sK57),s('type.list.list'('type.typeSystem.tenv_val'),X27))) != s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',X25),s('type.list.list'('type.typeSystem.tenv_val'),X26))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f18747,f1375])).
fof(f1375,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',sK57),s('type.list.list'('type.typeSystem.tenv_val'),X2))) != s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',X0),s('type.list.list'('type.typeSystem.tenv_val'),X1))) | p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X0))))) )),
  inference(superposition,[],[f768,f1247])).
fof(f18747,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',sK57),s('type.list.list'('type.typeSystem.tenv_val'),X27))) != s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',X25),s('type.list.list'('type.typeSystem.tenv_val'),X26))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X25)))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f18708,f775])).
fof(f18708,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',X25),s('type.list.list'('type.typeSystem.tenv_val'),X26))) != s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))),s('type.list.list'('type.typeSystem.tenv_val'),X27))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X25)))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1528,f1081])).
fof(f1528,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val',X0),s('type.list.list'('type.typeSystem.tenv_val'),X1))) != s('type.list.list'('type.typeSystem.tenv_val'),'const.list.CONS_2'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))),s('type.list.list'('type.typeSystem.tenv_val'),X2))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X0))))) )),
  inference(superposition,[],[f769,f1247])).
fof(f17735,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK61(X26,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X25))))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27))))) ) | $spl146),
  inference(cnf_transformation,[],[f17735_D])).
fof(f17735_D,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK61(X26,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X25))))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f17734,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK61(X26,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X25))))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f17630,f775])).
fof(f17630,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK61(X26,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X25))))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1485,f1081])).
fof(f1485,plain,(
  ( ! [X154,X155,X153] : (~p(s(bool,happ(s(fun(X153,bool),X154),s(X153,sK61(X154,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X153))))) | p(s(bool,happ(s(fun(X153,bool),X154),s(X153,X155))))) )),
  inference(resolution,[],[f769,f920])).
fof(f920,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK61(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK60(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK61(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK60,sK61])],[f572])).
fof(f572,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f571])).
fof(f571,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f570])).
fof(f570,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f299])).
fof(f299,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X37] : ! [X59,X60] : (! [X47] : (p(s(bool,happ(s(fun(X37,bool),X60),s(X37,X47)))) | p(s(bool,X59))) <=> (! [X47] : p(s(bool,happ(s(fun(X37,bool),X60),s(X37,X47)))) | p(s(bool,X59))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.LEFT_FORALL_OR_THM')).
fof(f17723,plain,(
  ( ! [X177,X180,X178,X176] : (p(s(bool,happ(s(fun(X176,bool),X177),s(X176,X178)))) | s(bool,X180) != s(bool,happ(s(fun(X176,bool),X177),s(X176,sK61(X177,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X176)))) | ~p(s(bool,X180))) ) | $spl144),
  inference(cnf_transformation,[],[f17723_D])).
fof(f17723_D,plain,(
  ( ! [X177,X180,X178,X176] : (p(s(bool,happ(s(fun(X176,bool),X177),s(X176,X178)))) | s(bool,X180) != s(bool,happ(s(fun(X176,bool),X177),s(X176,sK61(X177,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X176)))) | ~p(s(bool,X180))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f17615,plain,(
  ( ! [X606,X605,X607,X608,X604] : (p(s(bool,happ(s(fun(X604,bool),X605),s(X604,X606)))) | p(s(bool,X607)) | ~p(s(bool,X608)) | s(bool,X608) != s(bool,happ(s(fun(X604,bool),X605),s(X604,sK61(X605,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X604))))) )),
  inference(resolution,[],[f1485,f1274])).
fof(f17530,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (p(s(bool,happ(s(fun(X218,bool),X219),s(X218,X220)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,happ(s(fun(X218,bool),X219),s(X218,sK61(X219,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X218)))) | s(bool,X221) != s(bool,X223)) )),
  inference(resolution,[],[f1485,f802])).
fof(f17529,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (p(s(bool,happ(s(fun(X212,bool),X213),s(X212,X214)))) | p(s(bool,X215)) | ~p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X216) != s(bool,happ(s(fun(X212,bool),X213),s(X212,sK61(X213,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X212))))) )),
  inference(resolution,[],[f1485,f801])).
fof(f17527,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (p(s(bool,happ(s(fun(X200,bool),X201),s(X200,X202)))) | p(s(bool,X203)) | ~p(s(bool,X204)) | s(bool,X203) != s(bool,X205) | s(bool,X204) != s(bool,happ(s(fun(X200,bool),X201),s(X200,sK61(X201,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X200))))) )),
  inference(resolution,[],[f1485,f800])).
fof(f17525,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (p(s(bool,happ(s(fun(X188,bool),X189),s(X188,X190)))) | ~p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X192) != s(bool,happ(s(fun(X188,bool),X189),s(X188,sK61(X189,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X188)))) | s(bool,X191) != s(bool,X193)) )),
  inference(resolution,[],[f1485,f798])).
fof(f17523,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (p(s(bool,happ(s(fun(X176,bool),X177),s(X176,X178)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X179) != s(bool,X181) | s(bool,X180) != s(bool,happ(s(fun(X176,bool),X177),s(X176,sK61(X177,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X176))))) )),
  inference(resolution,[],[f1485,f796])).
fof(f17426,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK60(X26,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X25))))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27))))) ) | $spl142),
  inference(cnf_transformation,[],[f17426_D])).
fof(f17426_D,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK60(X26,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X25))))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f17425,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK60(X26,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X25))))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f17321,f775])).
fof(f17321,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK60(X26,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X25))))) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1484,f1081])).
fof(f1484,plain,(
  ( ! [X152,X151,X150] : (~p(s(bool,happ(s(fun(X150,bool),X151),s(X150,sK60(X151,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X150))))) | p(s(bool,happ(s(fun(X150,bool),X151),s(X150,X152))))) )),
  inference(resolution,[],[f769,f918])).
fof(f918,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK60(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f573])).
fof(f17414,plain,(
  ( ! [X177,X180,X178,X176] : (p(s(bool,happ(s(fun(X176,bool),X177),s(X176,X178)))) | s(bool,X180) != s(bool,happ(s(fun(X176,bool),X177),s(X176,sK60(X177,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X176)))) | ~p(s(bool,X180))) ) | $spl140),
  inference(cnf_transformation,[],[f17414_D])).
fof(f17414_D,plain,(
  ( ! [X177,X180,X178,X176] : (p(s(bool,happ(s(fun(X176,bool),X177),s(X176,X178)))) | s(bool,X180) != s(bool,happ(s(fun(X176,bool),X177),s(X176,sK60(X177,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X176)))) | ~p(s(bool,X180))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f17306,plain,(
  ( ! [X606,X605,X607,X608,X604] : (p(s(bool,happ(s(fun(X604,bool),X605),s(X604,X606)))) | p(s(bool,X607)) | ~p(s(bool,X608)) | s(bool,X608) != s(bool,happ(s(fun(X604,bool),X605),s(X604,sK60(X605,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X604))))) )),
  inference(resolution,[],[f1484,f1274])).
fof(f17221,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (p(s(bool,happ(s(fun(X218,bool),X219),s(X218,X220)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,happ(s(fun(X218,bool),X219),s(X218,sK60(X219,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X218)))) | s(bool,X221) != s(bool,X223)) )),
  inference(resolution,[],[f1484,f802])).
fof(f17220,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (p(s(bool,happ(s(fun(X212,bool),X213),s(X212,X214)))) | p(s(bool,X215)) | ~p(s(bool,X216)) | ~p(s(bool,X217)) | s(bool,X216) != s(bool,happ(s(fun(X212,bool),X213),s(X212,sK60(X213,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X212))))) )),
  inference(resolution,[],[f1484,f801])).
fof(f17218,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (p(s(bool,happ(s(fun(X200,bool),X201),s(X200,X202)))) | p(s(bool,X203)) | ~p(s(bool,X204)) | s(bool,X203) != s(bool,X205) | s(bool,X204) != s(bool,happ(s(fun(X200,bool),X201),s(X200,sK60(X201,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X200))))) )),
  inference(resolution,[],[f1484,f800])).
fof(f17216,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (p(s(bool,happ(s(fun(X188,bool),X189),s(X188,X190)))) | ~p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X192) != s(bool,happ(s(fun(X188,bool),X189),s(X188,sK60(X189,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X188)))) | s(bool,X191) != s(bool,X193)) )),
  inference(resolution,[],[f1484,f798])).
fof(f17214,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (p(s(bool,happ(s(fun(X176,bool),X177),s(X176,X178)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | s(bool,X179) != s(bool,X181) | s(bool,X180) != s(bool,happ(s(fun(X176,bool),X177),s(X176,sK60(X177,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X176))))) )),
  inference(resolution,[],[f1484,f796])).
fof(f16598,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))))) ) | $spl138),
  inference(cnf_transformation,[],[f16598_D])).
fof(f16598_D,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f16539,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))))) )),
  inference(resolution,[],[f1936,f1274])).
fof(f1936,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))))),
  inference(resolution,[],[f767,f1140])).
fof(f1140,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))))) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) )),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) | (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X3)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))))))),
  inference(ennf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) => (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X3)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X7,X8,X9,X39] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X6),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X7),s('type.list.list'('type.ast.pat'),X8),s('type.list.list'('type.ast.t'),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X39)))) => (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X6),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X9)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X6),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X39))))))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.typeSysProps.type_p_freevars|split|1')).
fof(f16460,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) | s(bool,X59) != s(bool,X60)) )),
  inference(resolution,[],[f1936,f802])).
fof(f16457,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))))) )),
  inference(resolution,[],[f1936,f801])).
fof(f16456,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))))) )),
  inference(resolution,[],[f1936,f800])).
fof(f16455,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) | s(bool,X44) != s(bool,X45)) )),
  inference(resolution,[],[f1936,f798])).
fof(f16453,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),'const.list.MAP_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t'),'type.ast.t'),'const.pair.SND_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))))) )),
  inference(resolution,[],[f1936,f796])).
fof(f16328,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X25),s('type.num.num',X26),s('type.ast.t',X27),s('type.typeSystem.tenv_val',sK57))))))) ) | $spl136),
  inference(cnf_transformation,[],[f16328_D])).
fof(f16328_D,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X25),s('type.num.num',X26),s('type.ast.t',X27),s('type.typeSystem.tenv_val',sK57))))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f16327,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X25),s('type.num.num',X26),s('type.ast.t',X27),s('type.typeSystem.tenv_val',sK57)))))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f16280,f775])).
fof(f16280,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X25),s('type.num.num',X26),s('type.ast.t',X27),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))))) | sP22(X28,X29,sK56,X30,X31) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X28),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X29),s('type.list.list'('type.ast.pat'),X31),s('type.list.list'('type.ast.t'),X30),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1408,f1081])).
fof(f1408,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X4),s('type.ast.t',X5),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))))))) )),
  inference(resolution,[],[f769,f777])).
fof(f777,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X3)))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.ast.t',X2),s('type.typeSystem.tenv_val',X3))))))) )),
  inference(cnf_transformation,[],[f514])).
fof(f16321,plain,(
  ( ! [X12,X13,X11,X9] : (~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X11),s('type.num.num',X12),s('type.ast.t',X13),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))))) ) | $spl134),
  inference(cnf_transformation,[],[f16321_D])).
fof(f16321_D,plain,(
  ( ! [X12,X13,X11,X9] : (~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X11),s('type.num.num',X12),s('type.ast.t',X13),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f16198,plain,(
  ( ! [X405,X407,X403,X404,X406] : (p(s(bool,X403)) | ~p(s(bool,X404)) | s(bool,X404) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X405),s('type.num.num',X406),s('type.ast.t',X407),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))))) )),
  inference(resolution,[],[f1408,f1274])).
fof(f16113,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X48) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X49),s('type.num.num',X50),s('type.ast.t',X51),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))))) | s(bool,X47) != s(bool,X52)) )),
  inference(resolution,[],[f1408,f802])).
fof(f16112,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X44),s('type.num.num',X45),s('type.ast.t',X46),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))))) )),
  inference(resolution,[],[f1408,f801])).
fof(f16110,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X35),s('type.num.num',X36),s('type.ast.t',X37),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))))) )),
  inference(resolution,[],[f1408,f800])).
fof(f16108,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X22),s('type.num.num',X23),s('type.ast.t',X24),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))))) | s(bool,X20) != s(bool,X25)) )),
  inference(resolution,[],[f1408,f798])).
fof(f16106,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.Bind_name_4'(s('type.list.list'('type.string.char'),X11),s('type.num.num',X12),s('type.ast.t',X13),s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))))) )),
  inference(resolution,[],[f1408,f796])).
fof(f15812,plain,(
  ( ! [X47,X45,X46] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X46))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X47))) | p(s(bool,X45))) )),
  inference(duplicate_literal_removal,[],[f15770])).
fof(f15770,plain,(
  ( ! [X47,X45,X46] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X46))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X47))) | p(s(bool,X45)) | p(s(bool,X45))) )),
  inference(superposition,[],[f15581,f983])).
fof(f983,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f615])).
fof(f615,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f15581,plain,(
  ( ! [X14,X12,X13] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X14))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X12),s('type.list.list'(bool),X13))) | p(s(bool,X12))) )),
  inference(forward_demodulation,[],[f1391,f1369])).
fof(f1369,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))),
  inference(resolution,[],[f768,f985])).
fof(f1391,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X12),s('type.list.list'(bool),X13))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))),s('type.list.list'(bool),X14)))) )),
  inference(superposition,[],[f768,f1247])).
fof(f15813,plain,(
  ( ! [X43,X44,X42] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X43))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X44))) | p(s(bool,X42))) )),
  inference(duplicate_literal_removal,[],[f15769])).
fof(f15769,plain,(
  ( ! [X43,X44,X42] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X43))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X44))) | p(s(bool,X42)) | p(s(bool,X42))) )),
  inference(superposition,[],[f15581,f981])).
fof(f981,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f614])).
fof(f614,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f15814,plain,(
  ( ! [X37,X38,X36] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X37))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X38))) | p(s(bool,X36))) )),
  inference(duplicate_literal_removal,[],[f15767])).
fof(f15767,plain,(
  ( ! [X37,X38,X36] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X37))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X38))) | p(s(bool,X36)) | p(s(bool,X36))) )),
  inference(superposition,[],[f15581,f935])).
fof(f15290,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f15289,f4927])).
fof(f4927,plain,(
  ( ! [X0,X1] : (sP7(X1,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X0) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f2296,f849])).
fof(f849,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f540])).
fof(f540,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f539])).
fof(f539,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f2296,plain,(
  ( ! [X2,X3] : (~sP5('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X3,X2) | ~p(s(bool,X3)) | sP7(X2,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X3)) )),
  inference(subsumption_resolution,[],[f2291,f855])).
fof(f855,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f542])).
fof(f542,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f2291,plain,(
  ( ! [X2,X3] : (sP7(X2,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X3) | ~p(s(bool,X3)) | ~sP5('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X3,X2) | ~sP4(X3,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X2)) )),
  inference(resolution,[],[f1327,f841])).
fof(f841,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP4(X2,X1,X0) & sP5(X1,X2,X0) & sP6(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X2,X1,X0) | ~sP5(X1,X2,X0) | ~sP6(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f534])).
fof(f534,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP7(X2,X1,X0)))),
  inference(flattening,[],[f533])).
fof(f533,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP7(X2,X1,X0)))),
  inference(nnf_transformation,[],[f458])).
fof(f458,plain,(
  ! [X2,X1,X0] : (sP7(X2,X1,X0) <=> (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f1327,plain,(
  ( ! [X101,X102] : (sP6(X101,X102,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f845])).
fof(f845,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f538])).
fof(f538,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f537])).
fof(f537,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f536])).
fof(f536,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f457])).
fof(f457,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f15289,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP7(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f15241,f775])).
fof(f15241,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP7(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1533,f1081])).
fof(f1533,plain,(
  ( ! [X6,X7] : (~sP7(X7,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f769,f858])).
fof(f858,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f544])).
fof(f544,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP7(X2,X1,X0)) & (~sP7(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP7(X2,X1,X0))),
  inference(definition_folding,[],[f279,f458,f457,f456,f455])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X79,X74,X12] : ((p(s(bool,X12)) <=> s(bool,X74) = s(bool,X79)) <=> ((p(s(bool,X12)) | p(s(bool,X74)) | p(s(bool,X79))) & (p(s(bool,X12)) | ~p(s(bool,X79)) | ~p(s(bool,X74))) & (p(s(bool,X74)) | ~p(s(bool,X79)) | ~p(s(bool,X12))) & (p(s(bool,X79)) | ~p(s(bool,X74)) | ~p(s(bool,X12)))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.sat.sat_2|49|_prove_ltb')).
fof(f15071,plain,(
  ( ! [X21,X20] : (sP7(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21))) ) | $spl130),
  inference(cnf_transformation,[],[f15071_D])).
fof(f15071_D,plain,(
  ( ! [X21,X20] : (sP7(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f15070,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP7(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f15069,f4927])).
fof(f15069,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP7(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f15029,f775])).
fof(f15029,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP7(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1531,f1081])).
fof(f1531,plain,(
  ( ! [X2,X3] : (sP7(X3,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X2) | ~p(s(bool,X2)) | p(s(bool,X3))) )),
  inference(superposition,[],[f769,f856])).
fof(f856,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f544])).
fof(f14495,plain,(
  ( ! [X21,X20] : (sP7(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X21))) ) | $spl128),
  inference(cnf_transformation,[],[f14495_D])).
fof(f14495_D,plain,(
  ( ! [X21,X20] : (sP7(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X21))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f14494,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP7(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f14493,f11740])).
fof(f11740,plain,(
  ( ! [X4,X5] : (sP7(X4,X5,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X5)) | ~p(s(bool,X4))) )),
  inference(resolution,[],[f11402,f849])).
fof(f11402,plain,(
  ( ! [X2,X3] : (~sP5(X3,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X2) | sP7(X2,X3,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f11372,f845])).
fof(f11372,plain,(
  ( ! [X95,X96] : (~sP6(X95,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X96) | sP7(X95,X96,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~sP5(X96,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X95)) )),
  inference(subsumption_resolution,[],[f1324,f1333])).
fof(f1333,plain,(
  ( ! [X114,X113] : (sP4('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X113,X114)) )),
  inference(resolution,[],[f768,f855])).
fof(f1324,plain,(
  ( ! [X95,X96] : (sP7(X95,X96,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~sP6(X95,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X96) | ~sP5(X96,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X95) | ~sP4('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X96,X95)) )),
  inference(resolution,[],[f768,f841])).
fof(f14493,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP7(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f14452,f775])).
fof(f14452,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP7(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1530,f1081])).
fof(f1530,plain,(
  ( ! [X0,X1] : (sP7(X1,X0,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(superposition,[],[f769,f856])).
fof(f14419,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f14418,f1332])).
fof(f1332,plain,(
  ( ! [X111,X112] : (sP4(X111,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X112)) )),
  inference(resolution,[],[f768,f854])).
fof(f854,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f543])).
fof(f14418,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP4(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f14385,f775])).
fof(f14385,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP4(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1459,f1081])).
fof(f1459,plain,(
  ( ! [X114,X115] : (~sP4(X114,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X115) | p(s(bool,X115)) | p(s(bool,X114))) )),
  inference(resolution,[],[f769,f852])).
fof(f852,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f543])).
fof(f14348,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f14347,f1331])).
fof(f1331,plain,(
  ( ! [X109,X110] : (sP4(X109,X110,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f853])).
fof(f853,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f543])).
fof(f14347,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP4(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f14315,f775])).
fof(f14315,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP4(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1458,f1081])).
fof(f1458,plain,(
  ( ! [X112,X113] : (~sP4(X112,X113,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X113)) | p(s(bool,X112))) )),
  inference(resolution,[],[f769,f852])).
fof(f14277,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f14276,f1333])).
fof(f14276,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP4('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f14243,f775])).
fof(f14243,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP4('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1457,f1081])).
fof(f1457,plain,(
  ( ! [X111,X110] : (~sP4('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X110,X111) | p(s(bool,X111)) | p(s(bool,X110))) )),
  inference(resolution,[],[f769,f852])).
fof(f14025,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f14024,f1330])).
fof(f1330,plain,(
  ( ! [X107,X108] : (sP5(X107,X108,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f849])).
fof(f14024,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP5(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f14002,f775])).
fof(f14002,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP5(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1454,f1081])).
fof(f1454,plain,(
  ( ! [X105,X104] : (~sP5(X104,X105,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X105)) | ~p(s(bool,X104))) )),
  inference(resolution,[],[f769,f848])).
fof(f848,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f541])).
fof(f13833,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f13832,f1327])).
fof(f13832,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP6(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f13804,f775])).
fof(f13804,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP6(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1451,f1081])).
fof(f1451,plain,(
  ( ! [X99,X98] : (~sP6(X98,X99,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X99)) | ~p(s(bool,X98))) )),
  inference(resolution,[],[f769,f844])).
fof(f844,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f538])).
fof(f13659,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f13658,f11740])).
fof(f13658,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP7(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f13603,f775])).
fof(f13603,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP7(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1448,f1081])).
fof(f1448,plain,(
  ( ! [X92,X93] : (~sP7(X92,X93,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X93)) | ~p(s(bool,X92))) )),
  inference(resolution,[],[f769,f840])).
fof(f840,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f535])).
fof(f13524,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f13523,f1315])).
fof(f1315,plain,(
  ( ! [X78,X77] : (sP3(X77,X78,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X77))) )),
  inference(resolution,[],[f768,f832])).
fof(f832,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f532])).
fof(f532,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f531])).
fof(f531,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f277,f453,f452])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X79,X74,X12] : ((p(s(bool,X12)) <=> (p(s(bool,X74)) | p(s(bool,X79)))) <=> ((p(s(bool,X12)) | ~p(s(bool,X74))) & (p(s(bool,X12)) | ~p(s(bool,X79))) & (p(s(bool,X74)) | p(s(bool,X79)) | ~p(s(bool,X12)))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.sat.sat_25_prove_ltb')).
fof(f13523,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f13460,f775])).
fof(f13460,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1445,f1081])).
fof(f1445,plain,(
  ( ! [X87,X86] : (~sP3(X87,X86,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X87)) | p(s(bool,X86))) )),
  inference(resolution,[],[f769,f834])).
fof(f834,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP3(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f532])).
fof(f13348,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f13347,f1317])).
fof(f1317,plain,(
  ( ! [X81,X82] : (sP3(X81,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X82) | ~p(s(bool,X81))) )),
  inference(resolution,[],[f768,f833])).
fof(f833,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f532])).
fof(f13347,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f13284,f775])).
fof(f13284,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1444,f1081])).
fof(f1444,plain,(
  ( ! [X85,X84] : (~sP3(X85,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X84) | ~p(s(bool,X85)) | p(s(bool,X84))) )),
  inference(resolution,[],[f769,f834])).
fof(f13154,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f13153,f1308])).
fof(f1308,plain,(
  ( ! [X64,X63] : (~sP3(X63,X64,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X63))) )),
  inference(resolution,[],[f768,f820])).
fof(f820,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f527])).
fof(f13153,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f13111,f775])).
fof(f13111,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP3(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1443,f1081])).
fof(f1443,plain,(
  ( ! [X83,X82] : (sP3(X82,X83,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X83)) | p(s(bool,X82))) )),
  inference(resolution,[],[f769,f831])).
fof(f831,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f532])).
fof(f13043,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f13042,f1319])).
fof(f1319,plain,(
  ( ! [X85,X86] : (~sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X85,X86) | p(s(bool,X86)) | p(s(bool,X85))) )),
  inference(resolution,[],[f768,f834])).
fof(f13042,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f13011,f775])).
fof(f13011,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | p(s(bool,X20)) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1442,f1081])).
fof(f1442,plain,(
  ( ! [X80,X81] : (sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X80,X81) | p(s(bool,X80)) | p(s(bool,X81))) )),
  inference(resolution,[],[f769,f831])).
fof(f12427,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f12426,f1309])).
fof(f1309,plain,(
  ( ! [X66,X65] : (~sP3(X65,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X66) | p(s(bool,X65))) )),
  inference(resolution,[],[f768,f821])).
fof(f821,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f527])).
fof(f12426,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f12384,f775])).
fof(f12384,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP3(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1441,f1081])).
fof(f1441,plain,(
  ( ! [X78,X79] : (sP3(X78,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X79) | p(s(bool,X79)) | p(s(bool,X78))) )),
  inference(resolution,[],[f769,f831])).
fof(f12360,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f12359,f1314])).
fof(f12359,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP2(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f12333,f775])).
fof(f12333,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP2(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1439,f1081])).
fof(f1439,plain,(
  ( ! [X74,X75] : (~sP2(X74,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X75) | p(s(bool,X75)) | ~p(s(bool,X74))) )),
  inference(resolution,[],[f769,f827])).
fof(f827,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f530])).
fof(f12164,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f12163,f1313])).
fof(f12163,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP2(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f12137,f775])).
fof(f12137,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP2(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1438,f1081])).
fof(f1438,plain,(
  ( ! [X72,X73] : (~sP2(X72,X73,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X73)) | ~p(s(bool,X72))) )),
  inference(resolution,[],[f769,f827])).
fof(f12075,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f12074,f1303])).
fof(f1303,plain,(
  ( ! [X54,X53] : (sP1(X53,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X54) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f768,f816])).
fof(f816,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f524])).
fof(f524,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f523])).
fof(f523,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f451])).
fof(f451,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f406,f450,f449])).
fof(f406,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X79,X74,X12] : ((p(s(bool,X12)) <=> (p(s(bool,X74)) => p(s(bool,X79)))) <=> ((p(s(bool,X12)) | p(s(bool,X74))) & (p(s(bool,X12)) | ~p(s(bool,X79))) & (~p(s(bool,X74)) | p(s(bool,X79)) | ~p(s(bool,X12)))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.sat.sat_27_prove_ltb')).
fof(f12074,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f12031,f775])).
fof(f12031,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1428,f1081])).
fof(f1428,plain,(
  ( ! [X52,X53] : (~sP1(X53,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X52) | ~p(s(bool,X53)) | ~p(s(bool,X52))) )),
  inference(resolution,[],[f769,f817])).
fof(f817,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f524])).
fof(f11858,plain,(
  ( ! [X21,X20] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) | $spl126),
  inference(cnf_transformation,[],[f11858_D])).
fof(f11858_D,plain,(
  ( ! [X21,X20] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f11857,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f11856,f1304])).
fof(f1304,plain,(
  ( ! [X56,X55] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X55,X56) | ~p(s(bool,X55))) )),
  inference(resolution,[],[f768,f816])).
fof(f11856,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f11855,f1302])).
fof(f1302,plain,(
  ( ! [X52,X51] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X51,X52) | p(s(bool,X52))) )),
  inference(resolution,[],[f768,f815])).
fof(f815,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f524])).
fof(f11855,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | p(s(bool,X20)) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f11819,f775])).
fof(f11819,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | p(s(bool,X20)) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1426,f1081])).
fof(f1426,plain,(
  ( ! [X48,X49] : (sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X48,X49) | p(s(bool,X48)) | ~p(s(bool,X49))) )),
  inference(resolution,[],[f769,f814])).
fof(f814,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f524])).
fof(f11699,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21))) ) | $spl124),
  inference(cnf_transformation,[],[f11699_D])).
fof(f11699_D,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f11698,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f11697,f1303])).
fof(f11697,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f11664,f775])).
fof(f11664,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X21)) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1425,f1081])).
fof(f1425,plain,(
  ( ! [X47,X46] : (sP1(X46,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X47) | ~p(s(bool,X47)) | p(s(bool,X46))) )),
  inference(resolution,[],[f769,f814])).
fof(f11588,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f11587,f1300])).
fof(f11587,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP0(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | ~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f11565,f775])).
fof(f11565,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP0(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | ~p(s(bool,X21)) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1422,f1081])).
fof(f1422,plain,(
  ( ! [X41,X40] : (~sP0(X40,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X41) | ~p(s(bool,X41)) | ~p(s(bool,X40))) )),
  inference(resolution,[],[f769,f810])).
fof(f810,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f522])).
fof(f11138,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21)) ) | $spl122),
  inference(cnf_transformation,[],[f11138_D])).
fof(f11138_D,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f11137,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(bool,X20) != s(bool,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f11136,f1334])).
fof(f1334,plain,(
  ( ! [X116,X115] : (sP7('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X115,X116) | s(bool,X115) != s(bool,X116)) )),
  inference(resolution,[],[f768,f857])).
fof(f857,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f544])).
fof(f11136,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP7('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | s(bool,X20) != s(bool,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f11095,f775])).
fof(f11095,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP7('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | s(bool,X20) != s(bool,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1461,f1081])).
fof(f1461,plain,(
  ( ! [X118,X119] : (~sP7('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X119,X118) | s(bool,X118) != s(bool,X119)) )),
  inference(resolution,[],[f769,f859])).
fof(f859,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f544])).
fof(f10878,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,X21)) ) | $spl120),
  inference(cnf_transformation,[],[f10878_D])).
fof(f10878_D,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,X21)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f10877,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(bool,X20) = s(bool,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f10876,f1335])).
fof(f1335,plain,(
  ( ! [X118,X117] : (~sP7('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X118,X117) | s(bool,X117) = s(bool,X118)) )),
  inference(resolution,[],[f768,f858])).
fof(f10876,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP7('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | s(bool,X20) = s(bool,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f10824,f775])).
fof(f10824,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP7('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | s(bool,X20) = s(bool,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1460,f1081])).
fof(f1460,plain,(
  ( ! [X116,X117] : (sP7('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X116,X117) | s(bool,X116) = s(bool,X117)) )),
  inference(resolution,[],[f769,f856])).
fof(f10179,plain,(
  ( ! [X24,X23,X25,X22,X20] : (p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f10178,f1342])).
fof(f1342,plain,(
  ( ! [X132,X131] : (sP8(X131,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X132)) )),
  inference(resolution,[],[f768,f872])).
fof(f872,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f550])).
fof(f550,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f281,f460])).
fof(f460,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f281,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f175])).
fof(f175,axiom,(
  ! [X37,X36,X85] : ((p(s(bool,X37)) | (p(s(bool,X36)) & p(s(bool,X85)))) <=> ((p(s(bool,X37)) | p(s(bool,X36))) & (p(s(bool,X37)) | p(s(bool,X85)))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.LEFT_OR_OVER_AND')).
fof(f10178,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP8(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f10151,f775])).
fof(f10151,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP8(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1465,f1081])).
fof(f1465,plain,(
  ( ! [X127,X126] : (~sP8(X126,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X127) | p(s(bool,X126))) )),
  inference(resolution,[],[f769,f867])).
fof(f867,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f549])).
fof(f549,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f548])).
fof(f548,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f547])).
fof(f547,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f460])).
fof(f10037,plain,(
  ( ! [X21,X20] : (~sP8('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | p(s(bool,X20))) ) | $spl118),
  inference(cnf_transformation,[],[f10037_D])).
fof(f10037_D,plain,(
  ( ! [X21,X20] : (~sP8('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | p(s(bool,X20))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f10036,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP8('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f10005,f775])).
fof(f10005,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP8('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1464,f1081])).
fof(f1464,plain,(
  ( ! [X125,X124] : (~sP8('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X124,X125) | p(s(bool,X124))) )),
  inference(resolution,[],[f769,f867])).
fof(f9926,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f9925,f1342])).
fof(f9925,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP8(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f9898,f775])).
fof(f9898,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP8(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1463,f1081])).
fof(f1463,plain,(
  ( ! [X123,X122] : (~sP8(X123,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X122) | p(s(bool,X122))) )),
  inference(resolution,[],[f769,f866])).
fof(f866,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f549])).
fof(f9795,plain,(
  ( ! [X21,X20] : (~sP8(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X21))) ) | $spl116),
  inference(cnf_transformation,[],[f9795_D])).
fof(f9795_D,plain,(
  ( ! [X21,X20] : (~sP8(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X21))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f9794,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP8(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f9764,f775])).
fof(f9764,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP8(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1462,f1081])).
fof(f1462,plain,(
  ( ! [X121,X120] : (~sP8(X121,X120,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X120))) )),
  inference(resolution,[],[f769,f866])).
fof(f9074,plain,(
  ( ! [X24,X23,X25,X22,X20] : (~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f9073,f1318])).
fof(f1318,plain,(
  ( ! [X83,X84] : (sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X83,X84) | ~p(s(bool,X83))) )),
  inference(resolution,[],[f768,f833])).
fof(f9073,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f9035,f775])).
fof(f9035,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1433,f1081])).
fof(f1433,plain,(
  ( ! [X62,X63] : (~sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X62,X63) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f769,f821])).
fof(f9012,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f9011,f1316])).
fof(f1316,plain,(
  ( ! [X80,X79] : (sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X79,X80) | ~p(s(bool,X80))) )),
  inference(resolution,[],[f768,f832])).
fof(f9011,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f8973,f775])).
fof(f8973,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | ~p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1432,f1081])).
fof(f1432,plain,(
  ( ! [X61,X60] : (~sP3('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X61,X60) | ~p(s(bool,X60))) )),
  inference(resolution,[],[f769,f820])).
fof(f8672,plain,(
  ( ! [X21,X20] : (sP1(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X20))) ) | $spl114),
  inference(cnf_transformation,[],[f8672_D])).
fof(f8672_D,plain,(
  ( ! [X21,X20] : (sP1(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X20))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f8671,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f8645,f775])).
fof(f8645,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP1(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1427,f1081])).
fof(f1427,plain,(
  ( ! [X50,X51] : (sP1(X50,X51,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X50))) )),
  inference(resolution,[],[f769,f815])).
fof(f8558,plain,(
  ( ! [X24,X23,X25,X22,X20] : (~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f8557,f1304])).
fof(f8557,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f8525,f775])).
fof(f8525,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | ~p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1419,f1081])).
fof(f1419,plain,(
  ( ! [X35,X34] : (~sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X34,X35) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f769,f804])).
fof(f804,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f519])).
fof(f8445,plain,(
  ( ! [X21,X20] : (~sP1(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X20))) ) | $spl112),
  inference(cnf_transformation,[],[f8445_D])).
fof(f8445_D,plain,(
  ( ! [X21,X20] : (~sP1(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X20))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f8444,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP1(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f8411,f775])).
fof(f8411,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP1(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X20)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1418,f1081])).
fof(f1418,plain,(
  ( ! [X33,X32] : (~sP1(X32,X33,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X32))) )),
  inference(resolution,[],[f769,f803])).
fof(f803,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f519])).
fof(f8330,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f8329,f1302])).
fof(f8329,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f8294,f775])).
fof(f8294,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | p(s(bool,X21)) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1417,f1081])).
fof(f1417,plain,(
  ( ! [X30,X31] : (~sP1('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X31,X30) | p(s(bool,X30))) )),
  inference(resolution,[],[f769,f803])).
fof(f7664,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f7663,f1353])).
fof(f1353,plain,(
  ( ! [X145] : (sP10(X145,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f894])).
fof(f7663,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~sP10(X15,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f7635,f775])).
fof(f7635,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~sP10(X15,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1554,f1081])).
fof(f1554,plain,(
  ( ! [X145] : (~sP10(X145,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X145))) )),
  inference(subsumption_resolution,[],[f1479,f950])).
fof(f950,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f595])).
fof(f595,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f594])).
fof(f594,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) | p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.OR_CLAUSES|split|2')).
fof(f1479,plain,(
  ( ! [X145] : (p(s(bool,'F_0')) | p(s(bool,X145)) | ~sP10(X145,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) )),
  inference(resolution,[],[f769,f893])).
fof(f893,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f559])).
fof(f7617,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f7616,f1354])).
fof(f1354,plain,(
  ( ! [X146] : (sP10('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X146)) )),
  inference(resolution,[],[f768,f895])).
fof(f7616,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~sP10('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X15) | p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f7589,f775])).
fof(f7589,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~sP10('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X15) | p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1553,f1081])).
fof(f1553,plain,(
  ( ! [X144] : (~sP10('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X144) | p(s(bool,X144))) )),
  inference(subsumption_resolution,[],[f1478,f950])).
fof(f1478,plain,(
  ( ! [X144] : (p(s(bool,'F_0')) | p(s(bool,X144)) | ~sP10('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X144)) )),
  inference(resolution,[],[f769,f893])).
fof(f5910,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f5909,f1345])).
fof(f1345,plain,(
  ( ! [X137] : (~sP9('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X137) | ~p(s(bool,X137))) )),
  inference(resolution,[],[f768,f878])).
fof(f878,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1)))),
  inference(rectify,[],[f553])).
fof(f553,plain,(
  ! [X1,X0] : ((~sP9(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP9(X1,X0)))),
  inference(flattening,[],[f552])).
fof(f552,plain,(
  ! [X1,X0] : ((~sP9(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP9(X1,X0)))),
  inference(nnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X1,X0] : (sP9(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f5909,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP9('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X15) | ~p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f5888,f775])).
fof(f5888,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP9('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X15) | ~p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1473,f1081])).
fof(f1473,plain,(
  ( ! [X139] : (sP9('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X139) | ~p(s(bool,X139))) )),
  inference(resolution,[],[f769,f882])).
fof(f882,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f554])).
fof(f5729,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f5728,f1346])).
fof(f1346,plain,(
  ( ! [X138] : (~sP9(X138,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X138))) )),
  inference(resolution,[],[f768,f878])).
fof(f5728,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP9(X15,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | ~p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f5707,f775])).
fof(f5707,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP9(X15,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1472,f1081])).
fof(f1472,plain,(
  ( ! [X138] : (sP9(X138,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X138))) )),
  inference(resolution,[],[f769,f879])).
fof(f879,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f554])).
fof(f5569,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f5568,f1348])).
fof(f1348,plain,(
  ( ! [X140] : (sP9(X140,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X140))) )),
  inference(resolution,[],[f768,f882])).
fof(f5568,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~sP9(X15,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f5545,f775])).
fof(f5545,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~sP9(X15,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1471,f1081])).
fof(f1471,plain,(
  ( ! [X137] : (~sP9(X137,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | p(s(bool,X137))) )),
  inference(resolution,[],[f769,f877])).
fof(f877,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f554])).
fof(f5346,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(subsumption_resolution,[],[f5345,f1347])).
fof(f1347,plain,(
  ( ! [X139] : (sP9('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X139) | p(s(bool,X139))) )),
  inference(resolution,[],[f768,f879])).
fof(f5345,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~sP9('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X15) | p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f5322,f775])).
fof(f5322,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~sP9('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X15) | p(s(bool,X15)) | sP22(X16,X17,sK56,X18,X19) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X16),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X17),s('type.list.list'('type.ast.pat'),X19),s('type.list.list'('type.ast.t'),X18),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1470,f1081])).
fof(f1470,plain,(
  ( ! [X136] : (~sP9('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X136) | p(s(bool,X136))) )),
  inference(resolution,[],[f769,f877])).
fof(f5229,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55)))) ) | $spl110),
  inference(cnf_transformation,[],[f5229_D])).
fof(f5229_D,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f5192,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55)))) )),
  inference(resolution,[],[f1935,f1274])).
fof(f1935,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55))))),
  inference(resolution,[],[f767,f1139])).
fof(f1139,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',X0),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),X3)))) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) )),
  inference(cnf_transformation,[],[f440])).
fof(f5113,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55))) | s(bool,X59) != s(bool,X60)) )),
  inference(resolution,[],[f1935,f802])).
fof(f5110,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55)))) )),
  inference(resolution,[],[f1935,f801])).
fof(f5109,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55)))) )),
  inference(resolution,[],[f1935,f800])).
fof(f5108,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55))) | s(bool,X44) != s(bool,X45)) )),
  inference(resolution,[],[f1935,f798])).
fof(f5106,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun('type.ast.t',bool),'const.typeSystem.check_freevars_2'(s('type.num.num',sK52),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))),s('type.list.list'('type.ast.t'),sK55)))) )),
  inference(resolution,[],[f1935,f796])).
fof(f5034,plain,(
  ( ! [X21,X20] : (sP5('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) | $spl108),
  inference(cnf_transformation,[],[f5034_D])).
fof(f5034_D,plain,(
  ( ! [X21,X20] : (sP5('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f5033,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP5('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f5011,f775])).
fof(f5011,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP5('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1456,f1081])).
fof(f1456,plain,(
  ( ! [X109,X108] : (sP5('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X108,X109)) )),
  inference(resolution,[],[f769,f851])).
fof(f851,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f541])).
fof(f4926,plain,(
  ( ! [X21,X20] : (sP5(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21)) ) | $spl106),
  inference(cnf_transformation,[],[f4926_D])).
fof(f4926_D,plain,(
  ( ! [X21,X20] : (sP5(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4925,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP5(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f4903,f775])).
fof(f4903,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP5(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1455,f1081])).
fof(f1455,plain,(
  ( ! [X107,X106] : (sP5(X106,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X107)) )),
  inference(resolution,[],[f769,f850])).
fof(f850,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f541])).
fof(f4850,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | $spl103),
  inference(cnf_transformation,[],[f4850_D])).
fof(f4850_D,plain,(
  s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f4858,plain,(
  ( ! [X8] : (sP7(X8,'F_0','const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) ) | $spl104),
  inference(cnf_transformation,[],[f4858_D])).
fof(f4858_D,plain,(
  ( ! [X8] : (sP7(X8,'F_0','const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f4855,plain,(
  ( ! [X8] : (sP7(X8,'F_0','const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) )),
  inference(resolution,[],[f4496,f1452])).
fof(f1452,plain,(
  ( ! [X101,X100] : (sP6(X100,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X101)) )),
  inference(resolution,[],[f769,f846])).
fof(f846,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f538])).
fof(f4496,plain,(
  ( ! [X4,X5] : (~sP6(X5,X4,'F_0') | sP7(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f4495,f995])).
fof(f995,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.typeSysProps.type_p_bvl|split|1',aHL_FALSITY)).
fof(f4495,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP7(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP6(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f4492,f851])).
fof(f4492,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP7(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP6(X5,X4,'F_0') | ~sP5('F_0',X4,X5)) )),
  inference(resolution,[],[f4482,f842])).
fof(f842,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f535])).
fof(f4482,plain,(
  ( ! [X4,X5] : (sP4(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f4471,f853])).
fof(f4471,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP4(X4,'F_0',X5)) )),
  inference(resolution,[],[f2400,f837])).
fof(f837,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f535])).
fof(f2400,plain,(
  ( ! [X6,X7] : (sP7(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f2393,f856])).
fof(f2393,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f2385,f768])).
fof(f2385,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))))),
  inference(superposition,[],[f1367,f985])).
fof(f1367,plain,(
  s(bool,'F_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))),
  inference(resolution,[],[f768,f980])).
fof(f980,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f614])).
fof(f4846,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f4777,f1452])).
fof(f4777,plain,(
  ( ! [X14,X13] : (~sP6(X13,X14,'F_0') | s(bool,'T_0') != s(bool,X14) | ~p(s(bool,X13))) )),
  inference(subsumption_resolution,[],[f4776,f853])).
fof(f4776,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | ~sP6(X13,X14,'F_0') | ~sP4(X14,'F_0',X13)) )),
  inference(subsumption_resolution,[],[f4775,f933])).
fof(f933,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.IMP_CLAUSES|split|4')).
fof(f4775,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | p(s(bool,'F_0')) | ~sP6(X13,X14,'F_0') | ~sP4(X14,'F_0',X13)) )),
  inference(subsumption_resolution,[],[f4762,f851])).
fof(f4762,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | p(s(bool,'F_0')) | ~sP6(X13,X14,'F_0') | ~sP5('F_0',X14,X13) | ~sP4(X14,'F_0',X13)) )),
  inference(resolution,[],[f2402,f842])).
fof(f2402,plain,(
  ( ! [X10,X11] : (~sP7(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) )),
  inference(superposition,[],[f2393,f858])).
fof(f4841,plain,(
  ( ! [X21,X20] : (sP6('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) | $spl100),
  inference(cnf_transformation,[],[f4841_D])).
fof(f4841_D,plain,(
  ( ! [X21,X20] : (sP6('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4840,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP6('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f4818,f775])).
fof(f4818,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP6('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1453,f1081])).
fof(f1453,plain,(
  ( ! [X103,X102] : (sP6('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X102,X103)) )),
  inference(resolution,[],[f769,f847])).
fof(f847,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f538])).
fof(f4735,plain,(
  ( ! [X21,X20] : (sP6(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21)) ) | $spl98),
  inference(cnf_transformation,[],[f4735_D])).
fof(f4735_D,plain,(
  ( ! [X21,X20] : (sP6(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f4734,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP6(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f4713,f775])).
fof(f4713,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP6(X20,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1452,f1081])).
fof(f4630,plain,(
  ( ! [X21,X20] : (sP2('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) | $spl96),
  inference(cnf_transformation,[],[f4630_D])).
fof(f4630_D,plain,(
  ( ! [X21,X20] : (sP2('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f4629,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP2('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f4608,f775])).
fof(f4608,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP2('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1440,f1081])).
fof(f1440,plain,(
  ( ! [X76,X77] : (sP2('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X76,X77)) )),
  inference(resolution,[],[f769,f830])).
fof(f830,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f530])).
fof(f4544,plain,(
  ( ! [X21,X20] : (sP0('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) | $spl94),
  inference(cnf_transformation,[],[f4544_D])).
fof(f4544_D,plain,(
  ( ! [X21,X20] : (sP0('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4543,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP0('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f4523,f775])).
fof(f4523,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP0('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))),X20,X21) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1424,f1081])).
fof(f1424,plain,(
  ( ! [X45,X44] : (sP0('const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))),X44,X45)) )),
  inference(resolution,[],[f769,f813])).
fof(f813,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f522])).
fof(f4402,plain,(
  ( ! [X21,X20] : (sP0(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) ) | $spl92),
  inference(cnf_transformation,[],[f4402_D])).
fof(f4402_D,plain,(
  ( ! [X21,X20] : (sP0(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4401,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP0(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(forward_demodulation,[],[f4381,f775])).
fof(f4381,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (sP0(X20,X21,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57))))) | sP22(X22,X23,sK56,X24,X25) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X22),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X23),s('type.list.list'('type.ast.pat'),X25),s('type.list.list'('type.ast.t'),X24),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f1423,f1081])).
fof(f1423,plain,(
  ( ! [X43,X42] : (sP0(X42,X43,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) )),
  inference(resolution,[],[f769,f811])).
fof(f811,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f522])).
fof(f4353,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl90),
  inference(cnf_transformation,[],[f4353_D])).
fof(f4353_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4352,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f4341])).
fof(f4341,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f2399,f859])).
fof(f2399,plain,(
  ( ! [X4,X5] : (sP7(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) )),
  inference(superposition,[],[f2393,f856])).
fof(f4148,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),sK55))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.pat'),'const.list.NIL_0'))) | $spl88),
  inference(cnf_transformation,[],[f4148_D])).
fof(f4148_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),sK55))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.pat'),'const.list.NIL_0'))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3979,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),sK55))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.pat'),'const.list.NIL_0'))) | sP22(X0,X1,X2,X3,sK54) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) )),
  inference(superposition,[],[f1937,f1079])).
fof(f1937,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),sK55))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.pat'),sK54)))),
  inference(resolution,[],[f767,f1141])).
fof(f1141,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.pat'),X2))) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4))))) )),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.pat'),X2))))),
  inference(ennf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X2),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.pat'),X2))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7,X8,X9,X10] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X6),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X7),s('type.list.list'('type.ast.pat'),X8),s('type.list.list'('type.ast.t'),X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X10)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.t'),X9))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.ast.pat'),X8))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.typeSysProps.type_ps_length')).
fof(f3551,plain,(
  $false | $spl84),
  inference(trivial_inequality_removal,[],[f3541])).
fof(f3541,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl84),
  inference(superposition,[],[f2721,f1369])).
fof(f3550,plain,(
  $false | $spl84),
  inference(trivial_inequality_removal,[],[f3542])).
fof(f3542,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl84),
  inference(superposition,[],[f2721,f1256])).
fof(f1256,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f365])).
fof(f365,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X37] : ! [X60] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X60),s('type.list.list'(X37),'const.list.NIL_0')))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.list.EVERY_DEF|split|0')).
fof(f3549,plain,(
  $false | $spl84),
  inference(equality_resolution,[],[f2721])).
fof(f3498,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f3296,f3179])).
fof(f3179,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(equality_resolution,[],[f1399])).
fof(f3296,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f3179,f998])).
fof(f998,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f619])).
fof(f619,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f340])).
fof(f340,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f46])).
fof(f46,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f3497,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f3295,f3179])).
fof(f3295,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f3179,f996])).
fof(f996,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f618])).
fof(f618,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f339])).
fof(f339,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f49])).
fof(f49,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f3496,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f3290,f3179])).
fof(f3290,plain,(
  ( ! [X240] : (p(s(bool,X240))) ) | $spl8),
  inference(resolution,[],[f3179,f947])).
fof(f947,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f593])).
fof(f593,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f592])).
fof(f592,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f430])).
fof(f430,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.IMP_CLAUSES|split|0')).
fof(f3495,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f3288,f3179])).
fof(f3288,plain,(
  ( ! [X238] : (p(s(bool,X238))) ) | $spl8),
  inference(resolution,[],[f3179,f944])).
fof(f944,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f591])).
fof(f591,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f590])).
fof(f590,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.IMP_CLAUSES|split|3')).
fof(f3494,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f3287,f3179])).
fof(f3287,plain,(
  ( ! [X237] : (p(s(bool,X237))) ) | $spl8),
  inference(resolution,[],[f3179,f941])).
fof(f941,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f588])).
fof(f588,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.IMP_CLAUSES|split|1')).
fof(f3493,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f3285,f3179])).
fof(f3285,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(resolution,[],[f3179,f938])).
fof(f938,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f587])).
fof(f587,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f586])).
fof(f586,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.IMP_CLAUSES|split|2')).
fof(f3297,plain,(
  $false | $spl8),
  inference(resolution,[],[f3179,f1000])).
fof(f1000,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  p(s(bool,'T_0'))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.TRUTH')).
fof(f3298,plain,(
  $false | $spl8),
  inference(resolution,[],[f3179,f1001])).
fof(f1001,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.typeSysProps.type_p_bvl|split|1',aHL_TRUTH)).
fof(f3326,plain,(
  $false | $spl8),
  inference(resolution,[],[f3179,f1278])).
fof(f1278,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f928])).
fof(f928,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f580])).
fof(f580,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X37] : ! [X47] : (s(X37,X47) = s(X37,X47) <=> p(s(bool,'T_0')))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.bool.REFL_CLAUSE')).
fof(f3343,plain,(
  $false | $spl8),
  inference(resolution,[],[f3179,f767])).
fof(f3419,plain,(
  $false | $spl8),
  inference(resolution,[],[f3179,f994])).
fof(f994,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f158])).
fof(f158,axiom,(
  ! [X28] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X28))))),
  file('thm.typeSysProps.type_p_bvl|split|1','thm.arithmetic.ZERO_LESS_EQ')).
fof(f2721,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33)) ) | $spl84),
  inference(cnf_transformation,[],[f2721_D])).
fof(f2721_D,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3181,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3136])).
fof(f3136,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | ~p(s(bool,X38)) | ~p(s(bool,X38))) ) | $spl8),
  inference(superposition,[],[f1399,f987])).
fof(f3182,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3135])).
fof(f3135,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | ~p(s(bool,X36))) ) | $spl8),
  inference(superposition,[],[f1399,f985])).
fof(f3115,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl86),
  inference(cnf_transformation,[],[f3115_D])).
fof(f3115_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3110,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f3064])).
fof(f3064,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34)) | p(s(bool,X34))) ) | $spl2),
  inference(superposition,[],[f1396,f983])).
fof(f3111,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f3063])).
fof(f3063,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32)) | p(s(bool,X32))) ) | $spl2),
  inference(superposition,[],[f1396,f981])).
fof(f3112,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f3061])).
fof(f3061,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl2),
  inference(superposition,[],[f1396,f935])).
fof(f2878,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f2850])).
fof(f2850,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(superposition,[],[f1397,f1369])).
fof(f2877,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f1397])).
fof(f1398,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | $spl6),
  inference(cnf_transformation,[],[f1398_D])).
fof(f1398_D,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X12)) | ~p(s(bool,X13))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2742,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,X110) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109)) )),
  inference(forward_demodulation,[],[f2741,f1369])).
fof(f2741,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109) | s(bool,X110) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(subsumption_resolution,[],[f2471,f1274])).
fof(f2471,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,'T_0')) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109) | s(bool,X110) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(superposition,[],[f802,f1369])).
fof(f2738,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X96) | p(s(bool,X97)) | ~p(s(bool,X98)) | s(bool,X97) != s(bool,X98)) )),
  inference(subsumption_resolution,[],[f2467,f800])).
fof(f2467,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X96) | p(s(bool,X97)) | ~p(s(bool,X96)) | ~p(s(bool,X98)) | s(bool,X97) != s(bool,X98)) )),
  inference(superposition,[],[f802,f1369])).
fof(f2737,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,X93) | p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,X95)) )),
  inference(subsumption_resolution,[],[f2466,f1274])).
fof(f2466,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,X93) | p(s(bool,X94)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,X95)) )),
  inference(superposition,[],[f802,f1369])).
fof(f2726,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,X50) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49)) )),
  inference(forward_demodulation,[],[f2725,f1369])).
fof(f2725,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(subsumption_resolution,[],[f2451,f1274])).
fof(f2451,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(superposition,[],[f798,f1369])).
fof(f2722,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2447,f1274])).
fof(f2447,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38)) )),
  inference(superposition,[],[f798,f1369])).
fof(f2720,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X33) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2446,f800])).
fof(f2446,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X33) | p(s(bool,X34)) | ~p(s(bool,X33)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X35)) )),
  inference(superposition,[],[f798,f1369])).
fof(f2124,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X2)))))) ) | $spl80),
  inference(cnf_transformation,[],[f2124_D])).
fof(f2124_D,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X2)))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2125,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X0)) ) | $spl82),
  inference(cnf_transformation,[],[f2125_D])).
fof(f2125_D,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X0)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1887,plain,(
  s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | $spl23),
  inference(cnf_transformation,[],[f1887_D])).
fof(f1887_D,plain,(
  s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f2235,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X2)))))) )),
  inference(subsumption_resolution,[],[f2231,f1278])).
fof(f2231,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X2)))))) )),
  inference(resolution,[],[f1588,f1230])).
fof(f1230,plain,(
  ( ! [X2,X0,X3,X1] : (sP42(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3)))))) )),
  inference(cnf_transformation,[],[f750])).
fof(f750,plain,(
  ! [X0,X1,X2] : ((~sP42(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',sK188(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP42(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK188])],[f749])).
fof(f749,plain,(
  ! [X0,X1,X2] : ((~sP42(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X4))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP42(X0,X1,X2)))),
  inference(rectify,[],[f748])).
fof(f748,plain,(
  ! [X22,X21,X20] : ((~sP42(X22,X21,X20) | ? [X37] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X37))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X37] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X37))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP42(X22,X21,X20)))),
  inference(nnf_transformation,[],[f501])).
fof(f501,plain,(
  ! [X22,X21,X20] : (sP42(X22,X21,X20) <=> ? [X37] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X37))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP42])])).
fof(f1588,plain,(
  ( ! [X33,X32] : (~sP42(sK56,X32,X33)) )),
  inference(subsumption_resolution,[],[f1587,f768])).
fof(f1587,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP42(sK56,X32,X33)) )),
  inference(forward_demodulation,[],[f1521,f775])).
fof(f1521,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP42(sK56,X32,X33)) )),
  inference(superposition,[],[f769,f1228])).
fof(f1228,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP42(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f750])).
fof(f1916,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X2)))))) ) | $spl32),
  inference(cnf_transformation,[],[f1916_D])).
fof(f1916_D,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X2)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1917,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X0)) ) | $spl34),
  inference(cnf_transformation,[],[f1917_D])).
fof(f1917_D,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X0)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2221,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X2)))))) )),
  inference(subsumption_resolution,[],[f2217,f1278])).
fof(f2217,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X2)))))) )),
  inference(resolution,[],[f1586,f1225])).
fof(f1225,plain,(
  ( ! [X2,X0,X3,X1] : (sP43(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3)))))) )),
  inference(cnf_transformation,[],[f747])).
fof(f747,plain,(
  ! [X0,X1,X2] : ((~sP43(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',sK187(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP43(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK187])],[f746])).
fof(f746,plain,(
  ! [X0,X1,X2] : ((~sP43(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X4))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP43(X0,X1,X2)))),
  inference(rectify,[],[f745])).
fof(f745,plain,(
  ! [X22,X21,X20] : ((~sP43(X22,X21,X20) | ? [X36] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X36))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X36] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X36))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP43(X22,X21,X20)))),
  inference(nnf_transformation,[],[f502])).
fof(f502,plain,(
  ! [X22,X21,X20] : (sP43(X22,X21,X20) <=> ? [X36] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X36))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP43])])).
fof(f1586,plain,(
  ( ! [X30,X31] : (~sP43(sK56,X30,X31)) )),
  inference(subsumption_resolution,[],[f1585,f768])).
fof(f1585,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP43(sK56,X30,X31)) )),
  inference(forward_demodulation,[],[f1520,f775])).
fof(f1520,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP43(sK56,X30,X31)) )),
  inference(superposition,[],[f769,f1223])).
fof(f1223,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP43(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f747])).
fof(f1903,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X2)))))) ) | $spl28),
  inference(cnf_transformation,[],[f1903_D])).
fof(f1903_D,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X2)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1904,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X0)) ) | $spl30),
  inference(cnf_transformation,[],[f1904_D])).
fof(f1904_D,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X0)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2210,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X2)))))) )),
  inference(subsumption_resolution,[],[f2206,f1278])).
fof(f2206,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X2)))))) )),
  inference(resolution,[],[f1584,f1220])).
fof(f1220,plain,(
  ( ! [X2,X0,X3,X1] : (sP44(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3)))))) )),
  inference(cnf_transformation,[],[f744])).
fof(f744,plain,(
  ! [X0,X1,X2] : ((~sP44(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),sK186(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP44(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK186])],[f743])).
fof(f743,plain,(
  ! [X0,X1,X2] : ((~sP44(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X4))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP44(X0,X1,X2)))),
  inference(rectify,[],[f742])).
fof(f742,plain,(
  ! [X22,X21,X20] : ((~sP44(X22,X21,X20) | ? [X35] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X35))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X35] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X35))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP44(X22,X21,X20)))),
  inference(nnf_transformation,[],[f503])).
fof(f503,plain,(
  ! [X22,X21,X20] : (sP44(X22,X21,X20) <=> ? [X35] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X35))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP44])])).
fof(f1584,plain,(
  ( ! [X28,X29] : (~sP44(sK56,X28,X29)) )),
  inference(subsumption_resolution,[],[f1583,f768])).
fof(f1583,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP44(sK56,X28,X29)) )),
  inference(forward_demodulation,[],[f1519,f775])).
fof(f1519,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP44(sK56,X28,X29)) )),
  inference(superposition,[],[f769,f1218])).
fof(f1218,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP44(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f744])).
fof(f1893,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X2)))))) ) | $spl24),
  inference(cnf_transformation,[],[f1893_D])).
fof(f1893_D,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X2)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1894,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X0)) ) | $spl26),
  inference(cnf_transformation,[],[f1894_D])).
fof(f1894_D,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X0)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2202,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X2)))))) )),
  inference(subsumption_resolution,[],[f2198,f1278])).
fof(f2198,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X2)))))) )),
  inference(resolution,[],[f1582,f1215])).
fof(f1215,plain,(
  ( ! [X2,X0,X3,X1] : (sP45(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3)))))) )),
  inference(cnf_transformation,[],[f741])).
fof(f741,plain,(
  ! [X0,X1,X2] : ((~sP45(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),sK185(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP45(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK185])],[f740])).
fof(f740,plain,(
  ! [X0,X1,X2] : ((~sP45(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X4))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP45(X0,X1,X2)))),
  inference(rectify,[],[f739])).
fof(f739,plain,(
  ! [X22,X21,X20] : ((~sP45(X22,X21,X20) | ? [X34] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X34))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X34] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X34))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP45(X22,X21,X20)))),
  inference(nnf_transformation,[],[f504])).
fof(f504,plain,(
  ! [X22,X21,X20] : (sP45(X22,X21,X20) <=> ? [X34] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X34))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP45])])).
fof(f1582,plain,(
  ( ! [X26,X27] : (~sP45(sK56,X26,X27)) )),
  inference(subsumption_resolution,[],[f1581,f768])).
fof(f1581,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP45(sK56,X26,X27)) )),
  inference(forward_demodulation,[],[f1518,f775])).
fof(f1518,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP45(sK56,X26,X27)) )),
  inference(superposition,[],[f769,f1213])).
fof(f1213,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP45(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f741])).
fof(f1884,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X2)))))) ) | $spl18),
  inference(cnf_transformation,[],[f1884_D])).
fof(f1884_D,plain,(
  ( ! [X2,X1] : (s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X2)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1885,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X0)) ) | $spl20),
  inference(cnf_transformation,[],[f1885_D])).
fof(f1885_D,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2197,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X2)))))) )),
  inference(subsumption_resolution,[],[f2193,f1278])).
fof(f2193,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X2)))))) )),
  inference(resolution,[],[f1580,f1210])).
fof(f1210,plain,(
  ( ! [X2,X0,X3,X1] : (sP46(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3)))))) )),
  inference(cnf_transformation,[],[f738])).
fof(f738,plain,(
  ! [X0,X1,X2] : ((~sP46(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK184(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP46(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK184])],[f737])).
fof(f737,plain,(
  ! [X0,X1,X2] : ((~sP46(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X4))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP46(X0,X1,X2)))),
  inference(rectify,[],[f736])).
fof(f736,plain,(
  ! [X22,X21,X20] : ((~sP46(X22,X21,X20) | ? [X33] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X33))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X33] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X33))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP46(X22,X21,X20)))),
  inference(nnf_transformation,[],[f505])).
fof(f505,plain,(
  ! [X22,X21,X20] : (sP46(X22,X21,X20) <=> ? [X33] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X33))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP46])])).
fof(f1580,plain,(
  ( ! [X24,X25] : (~sP46(sK56,X24,X25)) )),
  inference(subsumption_resolution,[],[f1579,f768])).
fof(f1579,plain,(
  ( ! [X24,X25] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP46(sK56,X24,X25)) )),
  inference(forward_demodulation,[],[f1517,f775])).
fof(f1517,plain,(
  ( ! [X24,X25] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP46(sK56,X24,X25)) )),
  inference(superposition,[],[f769,f1208])).
fof(f1208,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP46(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f738])).
fof(f2180,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X2)))))) )),
  inference(subsumption_resolution,[],[f2176,f1278])).
fof(f2176,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X2)))))) )),
  inference(resolution,[],[f1578,f1129])).
fof(f1129,plain,(
  ( ! [X2,X0,X3,X1] : (sP24(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3)))))) )),
  inference(cnf_transformation,[],[f686])).
fof(f686,plain,(
  ! [X0,X1,X2] : ((~sP24(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',sK113(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP24(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK113])],[f685])).
fof(f685,plain,(
  ! [X0,X1,X2] : ((~sP24(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X4))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP24(X0,X1,X2)))),
  inference(rectify,[],[f684])).
fof(f684,plain,(
  ! [X4,X3,X2] : ((~sP24(X4,X3,X2) | ? [X19] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X19))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0')))) & (! [X19] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X19))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) | ~p(s(bool,'T_0'))) | sP24(X4,X3,X2)))),
  inference(nnf_transformation,[],[f481])).
fof(f1578,plain,(
  ( ! [X23,X22] : (~sP24(sK56,X22,X23)) )),
  inference(subsumption_resolution,[],[f1577,f768])).
fof(f1577,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP24(sK56,X22,X23)) )),
  inference(forward_demodulation,[],[f1516,f775])).
fof(f1516,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP24(sK56,X22,X23)) )),
  inference(superposition,[],[f769,f1127])).
fof(f1127,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP24(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f686])).
fof(f2166,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X2)))))) )),
  inference(subsumption_resolution,[],[f2162,f1278])).
fof(f2162,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X2)))))) )),
  inference(resolution,[],[f1576,f1124])).
fof(f1124,plain,(
  ( ! [X2,X0,X3,X1] : (sP25(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3)))))) )),
  inference(cnf_transformation,[],[f683])).
fof(f683,plain,(
  ! [X0,X1,X2] : ((~sP25(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',sK112(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP25(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK112])],[f682])).
fof(f682,plain,(
  ! [X0,X1,X2] : ((~sP25(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X4))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP25(X0,X1,X2)))),
  inference(rectify,[],[f681])).
fof(f681,plain,(
  ! [X4,X3,X2] : ((~sP25(X4,X3,X2) | ? [X18] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X18))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0')))) & (! [X18] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X18))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) | ~p(s(bool,'T_0'))) | sP25(X4,X3,X2)))),
  inference(nnf_transformation,[],[f482])).
fof(f1576,plain,(
  ( ! [X21,X20] : (~sP25(sK56,X20,X21)) )),
  inference(subsumption_resolution,[],[f1575,f768])).
fof(f1575,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP25(sK56,X20,X21)) )),
  inference(forward_demodulation,[],[f1515,f775])).
fof(f1515,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP25(sK56,X20,X21)) )),
  inference(superposition,[],[f769,f1122])).
fof(f1122,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP25(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f683])).
fof(f2155,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X2)))))) )),
  inference(subsumption_resolution,[],[f2151,f1278])).
fof(f2151,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X2)))))) )),
  inference(resolution,[],[f1574,f1119])).
fof(f1119,plain,(
  ( ! [X2,X0,X3,X1] : (sP26(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3)))))) )),
  inference(cnf_transformation,[],[f680])).
fof(f680,plain,(
  ! [X0,X1,X2] : ((~sP26(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),sK111(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP26(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK111])],[f679])).
fof(f679,plain,(
  ! [X0,X1,X2] : ((~sP26(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X4))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP26(X0,X1,X2)))),
  inference(rectify,[],[f678])).
fof(f678,plain,(
  ! [X4,X3,X2] : ((~sP26(X4,X3,X2) | ? [X17] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X17))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0')))) & (! [X17] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X17))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) | ~p(s(bool,'T_0'))) | sP26(X4,X3,X2)))),
  inference(nnf_transformation,[],[f483])).
fof(f1574,plain,(
  ( ! [X19,X18] : (~sP26(sK56,X18,X19)) )),
  inference(subsumption_resolution,[],[f1573,f768])).
fof(f1573,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP26(sK56,X18,X19)) )),
  inference(forward_demodulation,[],[f1514,f775])).
fof(f1514,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP26(sK56,X18,X19)) )),
  inference(superposition,[],[f769,f1117])).
fof(f1117,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP26(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f680])).
fof(f2147,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X2)))))) )),
  inference(subsumption_resolution,[],[f2143,f1278])).
fof(f2143,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X2)))))) )),
  inference(resolution,[],[f1572,f1114])).
fof(f1114,plain,(
  ( ! [X2,X0,X3,X1] : (sP27(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3)))))) )),
  inference(cnf_transformation,[],[f677])).
fof(f677,plain,(
  ! [X0,X1,X2] : ((~sP27(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),sK110(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP27(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK110])],[f676])).
fof(f676,plain,(
  ! [X0,X1,X2] : ((~sP27(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X4))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP27(X0,X1,X2)))),
  inference(rectify,[],[f675])).
fof(f675,plain,(
  ! [X4,X3,X2] : ((~sP27(X4,X3,X2) | ? [X16] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X16))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0')))) & (! [X16] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X16))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) | ~p(s(bool,'T_0'))) | sP27(X4,X3,X2)))),
  inference(nnf_transformation,[],[f484])).
fof(f1572,plain,(
  ( ! [X17,X16] : (~sP27(sK56,X16,X17)) )),
  inference(subsumption_resolution,[],[f1571,f768])).
fof(f1571,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP27(sK56,X16,X17)) )),
  inference(forward_demodulation,[],[f1513,f775])).
fof(f1513,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP27(sK56,X16,X17)) )),
  inference(superposition,[],[f769,f1112])).
fof(f1112,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP27(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f677])).
fof(f2142,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X2)))))) )),
  inference(subsumption_resolution,[],[f2138,f1278])).
fof(f2138,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X2)))))) )),
  inference(resolution,[],[f1570,f1109])).
fof(f1109,plain,(
  ( ! [X2,X0,X3,X1] : (sP28(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3)))))) )),
  inference(cnf_transformation,[],[f674])).
fof(f674,plain,(
  ! [X0,X1,X2] : ((~sP28(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK109(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP28(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK109])],[f673])).
fof(f673,plain,(
  ! [X0,X1,X2] : ((~sP28(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X4))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP28(X0,X1,X2)))),
  inference(rectify,[],[f672])).
fof(f672,plain,(
  ! [X4,X3,X2] : ((~sP28(X4,X3,X2) | ? [X15] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X15))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) & p(s(bool,'T_0')))) & (! [X15] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X15))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X4) | ~p(s(bool,'T_0'))) | sP28(X4,X3,X2)))),
  inference(nnf_transformation,[],[f485])).
fof(f1570,plain,(
  ( ! [X14,X15] : (~sP28(sK56,X14,X15)) )),
  inference(subsumption_resolution,[],[f1569,f768])).
fof(f1569,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP28(sK56,X14,X15)) )),
  inference(forward_demodulation,[],[f1512,f775])).
fof(f1512,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP28(sK56,X14,X15)) )),
  inference(superposition,[],[f769,f1107])).
fof(f1107,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP28(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f674])).
fof(f2123,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X2)))))) )),
  inference(subsumption_resolution,[],[f2119,f1278])).
fof(f2119,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X2)))))) )),
  inference(resolution,[],[f1566,f1060])).
fof(f1060,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3)))))) )),
  inference(cnf_transformation,[],[f650])).
fof(f650,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',sK84(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP12(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK84])],[f649])).
fof(f649,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X4))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X3))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f648])).
fof(f648,plain,(
  ! [X22,X21,X20] : ((~sP12(X22,X21,X20) | ? [X37] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X37))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X37] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X37))))) | s('type.ast.t','const.ast.Tint_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP12(X22,X21,X20)))),
  inference(nnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X22,X21,X20] : (sP12(X22,X21,X20) <=> ? [X37] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.IntLit_1'(s('type.integer.int',X37))))) & s('type.ast.t','const.ast.Tint_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f1566,plain,(
  ( ! [X8,X9] : (~sP12(sK56,X8,X9)) )),
  inference(subsumption_resolution,[],[f1565,f768])).
fof(f1565,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP12(sK56,X8,X9)) )),
  inference(forward_demodulation,[],[f1510,f775])).
fof(f1510,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP12(sK56,X8,X9)) )),
  inference(superposition,[],[f769,f1058])).
fof(f1058,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f650])).
fof(f2115,plain,(
  ( ! [X33,X32] : (~sP42(sK56,X32,X33)) ) | $spl78),
  inference(cnf_transformation,[],[f2115_D])).
fof(f2115_D,plain,(
  ( ! [X33,X32] : (~sP42(sK56,X32,X33)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2104,plain,(
  p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | $spl56),
  inference(cnf_transformation,[],[f2104_D])).
fof(f2104_D,plain,(
  p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2063,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP42(sK56,X32,X33)) )),
  inference(superposition,[],[f767,f1228])).
fof(f2114,plain,(
  ( ! [X30,X31] : (~sP43(sK56,X30,X31)) ) | $spl76),
  inference(cnf_transformation,[],[f2114_D])).
fof(f2114_D,plain,(
  ( ! [X30,X31] : (~sP43(sK56,X30,X31)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2062,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP43(sK56,X30,X31)) )),
  inference(superposition,[],[f767,f1223])).
fof(f2113,plain,(
  ( ! [X28,X29] : (~sP44(sK56,X28,X29)) ) | $spl74),
  inference(cnf_transformation,[],[f2113_D])).
fof(f2113_D,plain,(
  ( ! [X28,X29] : (~sP44(sK56,X28,X29)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2061,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP44(sK56,X28,X29)) )),
  inference(superposition,[],[f767,f1218])).
fof(f2112,plain,(
  ( ! [X26,X27] : (~sP45(sK56,X26,X27)) ) | $spl72),
  inference(cnf_transformation,[],[f2112_D])).
fof(f2112_D,plain,(
  ( ! [X26,X27] : (~sP45(sK56,X26,X27)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2060,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP45(sK56,X26,X27)) )),
  inference(superposition,[],[f767,f1213])).
fof(f2111,plain,(
  ( ! [X24,X25] : (~sP46(sK56,X24,X25)) ) | $spl70),
  inference(cnf_transformation,[],[f2111_D])).
fof(f2111_D,plain,(
  ( ! [X24,X25] : (~sP46(sK56,X24,X25)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2059,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP46(sK56,X24,X25)) )),
  inference(superposition,[],[f767,f1208])).
fof(f2110,plain,(
  ( ! [X23,X22] : (~sP24(sK56,X22,X23)) ) | $spl68),
  inference(cnf_transformation,[],[f2110_D])).
fof(f2110_D,plain,(
  ( ! [X23,X22] : (~sP24(sK56,X22,X23)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2058,plain,(
  ( ! [X23,X22] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP24(sK56,X22,X23)) )),
  inference(superposition,[],[f767,f1127])).
fof(f2109,plain,(
  ( ! [X21,X20] : (~sP25(sK56,X20,X21)) ) | $spl66),
  inference(cnf_transformation,[],[f2109_D])).
fof(f2109_D,plain,(
  ( ! [X21,X20] : (~sP25(sK56,X20,X21)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2057,plain,(
  ( ! [X21,X20] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP25(sK56,X20,X21)) )),
  inference(superposition,[],[f767,f1122])).
fof(f2108,plain,(
  ( ! [X19,X18] : (~sP26(sK56,X18,X19)) ) | $spl64),
  inference(cnf_transformation,[],[f2108_D])).
fof(f2108_D,plain,(
  ( ! [X19,X18] : (~sP26(sK56,X18,X19)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2056,plain,(
  ( ! [X19,X18] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP26(sK56,X18,X19)) )),
  inference(superposition,[],[f767,f1117])).
fof(f2107,plain,(
  ( ! [X17,X16] : (~sP27(sK56,X16,X17)) ) | $spl62),
  inference(cnf_transformation,[],[f2107_D])).
fof(f2107_D,plain,(
  ( ! [X17,X16] : (~sP27(sK56,X16,X17)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2055,plain,(
  ( ! [X17,X16] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP27(sK56,X16,X17)) )),
  inference(superposition,[],[f767,f1112])).
fof(f2106,plain,(
  ( ! [X14,X15] : (~sP28(sK56,X14,X15)) ) | $spl60),
  inference(cnf_transformation,[],[f2106_D])).
fof(f2106_D,plain,(
  ( ! [X14,X15] : (~sP28(sK56,X14,X15)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2054,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP28(sK56,X14,X15)) )),
  inference(superposition,[],[f767,f1107])).
fof(f2053,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | sP22(X10,X11,sK56,X12,X13) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X10),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X11),s('type.list.list'('type.ast.pat'),X13),s('type.list.list'('type.ast.t'),X12),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))))) )),
  inference(superposition,[],[f767,f1081])).
fof(f2102,plain,(
  ( ! [X8,X9] : (~sP12(sK56,X8,X9)) ) | $spl54),
  inference(cnf_transformation,[],[f2102_D])).
fof(f2102_D,plain,(
  ( ! [X8,X9] : (~sP12(sK56,X8,X9)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2052,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0')))) | ~sP12(sK56,X8,X9)) )),
  inference(superposition,[],[f767,f1058])).
fof(f2101,plain,(
  p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) | $spl52),
  inference(cnf_transformation,[],[f2101_D])).
fof(f2101_D,plain,(
  p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2043,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) | sP22(X0,X1,X2,sK55,X3) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),X3),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) )),
  inference(superposition,[],[f767,f1080])).
fof(f2098,plain,(
  p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) | $spl48),
  inference(cnf_transformation,[],[f2098_D])).
fof(f2098_D,plain,(
  p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2038,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),'const.list.NIL_0'),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) | sP22(X0,X1,X2,X3,sK54) | ~p(s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',X0),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),X1),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) )),
  inference(superposition,[],[f767,f1079])).
fof(f2092,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) | $spl42),
  inference(cnf_transformation,[],[f2092_D])).
fof(f2092_D,plain,(
  ( ! [X21] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2028,plain,(
  ( ! [X163,X162] : (p(s(bool,X162)) | p(s(bool,X163)) | s(bool,X163) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f767,f1274])).
fof(f1949,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f767,f802])).
fof(f2093,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) | $spl44),
  inference(cnf_transformation,[],[f2093_D])).
fof(f2093_D,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1948,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X39) | s(bool,X40) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f767,f802])).
fof(f1946,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f767,f801])).
fof(f1945,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f767,f800])).
fof(f1944,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f767,f798])).
fof(f1943,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X24) | s(bool,X25) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f767,f798])).
fof(f1942,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.typeSystem.type_ps_5'(s('type.num.num',sK52),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn')))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.list.list'('type.string.char')),'type.pair.prod'('type.list.list'('type.ast.t'),'type.semanticPrimitives.tid_or_exn'))))),sK53),s('type.list.list'('type.ast.pat'),sK54),s('type.list.list'('type.ast.t'),sK55),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)))) )),
  inference(resolution,[],[f767,f796])).
fof(f1886,plain,(
  s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | $spl22),
  inference(cnf_transformation,[],[f1886_D])).
fof(f1886_D,plain,(
  s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2089,plain,(
  sP22(sK52,sK53,sK56,sK55,sK54) | $spl38),
  inference(cnf_transformation,[],[f2089_D])).
fof(f2089_D,plain,(
  sP22(sK52,sK53,sK56,sK55,sK54) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1932,plain,(
  sP22(sK52,sK53,sK56,sK55,sK54) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56)),
  inference(resolution,[],[f767,f1081])).
fof(f2091,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK55) | $spl40),
  inference(cnf_transformation,[],[f2091_D])).
fof(f2091_D,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK55) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1931,plain,(
  sP22(sK52,sK53,sK56,sK55,sK54) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK55)),
  inference(resolution,[],[f767,f1080])).
fof(f2087,plain,(
  s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),sK54) | $spl36),
  inference(cnf_transformation,[],[f2087_D])).
fof(f2087_D,plain,(
  s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),sK54) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1930,plain,(
  sP22(sK52,sK53,sK56,sK55,sK54) | s('type.list.list'('type.ast.pat'),'const.list.NIL_0') = s('type.list.list'('type.ast.pat'),sK54)),
  inference(resolution,[],[f767,f1079])).
fof(f1915,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X2)))))) )),
  inference(subsumption_resolution,[],[f1911,f1278])).
fof(f1911,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X2)))))) )),
  inference(resolution,[],[f1564,f1055])).
fof(f1055,plain,(
  ( ! [X2,X0,X3,X1] : (sP13(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3)))))) )),
  inference(cnf_transformation,[],[f647])).
fof(f647,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',sK83(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP13(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK83])],[f646])).
fof(f646,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X4))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X3))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP13(X0,X1,X2)))),
  inference(rectify,[],[f645])).
fof(f645,plain,(
  ! [X22,X21,X20] : ((~sP13(X22,X21,X20) | ? [X36] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X36))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X36] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X36))))) | s('type.ast.t','const.ast.Tchar_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP13(X22,X21,X20)))),
  inference(nnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X22,X21,X20] : (sP13(X22,X21,X20) <=> ? [X36] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Char_1'(s('type.string.char',X36))))) & s('type.ast.t','const.ast.Tchar_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f1564,plain,(
  ( ! [X6,X7] : (~sP13(sK56,X6,X7)) )),
  inference(subsumption_resolution,[],[f1563,f768])).
fof(f1563,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP13(sK56,X6,X7)) )),
  inference(forward_demodulation,[],[f1509,f775])).
fof(f1509,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP13(sK56,X6,X7)) )),
  inference(superposition,[],[f769,f1053])).
fof(f1053,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP13(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f647])).
fof(f1902,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X2)))))) )),
  inference(subsumption_resolution,[],[f1898,f1278])).
fof(f1898,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X2)))))) )),
  inference(resolution,[],[f1562,f1050])).
fof(f1050,plain,(
  ( ! [X2,X0,X3,X1] : (sP14(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3)))))) )),
  inference(cnf_transformation,[],[f644])).
fof(f644,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),sK82(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP14(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK82])],[f643])).
fof(f643,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X4))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X3))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP14(X0,X1,X2)))),
  inference(rectify,[],[f642])).
fof(f642,plain,(
  ! [X22,X21,X20] : ((~sP14(X22,X21,X20) | ? [X35] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X35))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X35] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X35))))) | s('type.ast.t','const.ast.Tstring_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP14(X22,X21,X20)))),
  inference(nnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X22,X21,X20] : (sP14(X22,X21,X20) <=> ? [X35] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.StrLit_1'(s('type.list.list'('type.string.char'),X35))))) & s('type.ast.t','const.ast.Tstring_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f1562,plain,(
  ( ! [X4,X5] : (~sP14(sK56,X4,X5)) )),
  inference(subsumption_resolution,[],[f1561,f768])).
fof(f1561,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP14(sK56,X4,X5)) )),
  inference(forward_demodulation,[],[f1508,f775])).
fof(f1508,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP14(sK56,X4,X5)) )),
  inference(superposition,[],[f769,f1048])).
fof(f1048,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP14(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f644])).
fof(f1892,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X2)))))) )),
  inference(subsumption_resolution,[],[f1888,f1278])).
fof(f1888,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X2)))))) )),
  inference(resolution,[],[f1560,f1045])).
fof(f1045,plain,(
  ( ! [X2,X0,X3,X1] : (sP15(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3)))))) )),
  inference(cnf_transformation,[],[f641])).
fof(f641,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),sK81(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP15(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK81])],[f640])).
fof(f640,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X4))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X3))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP15(X0,X1,X2)))),
  inference(rectify,[],[f639])).
fof(f639,plain,(
  ! [X22,X21,X20] : ((~sP15(X22,X21,X20) | ? [X34] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X34))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X34] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X34))))) | s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP15(X22,X21,X20)))),
  inference(nnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X22,X21,X20] : (sP15(X22,X21,X20) <=> ? [X34] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word8_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))),X34))))) & s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f1560,plain,(
  ( ! [X2,X3] : (~sP15(sK56,X2,X3)) )),
  inference(subsumption_resolution,[],[f1559,f768])).
fof(f1559,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP15(sK56,X2,X3)) )),
  inference(forward_demodulation,[],[f1507,f775])).
fof(f1507,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP15(sK56,X2,X3)) )),
  inference(superposition,[],[f769,f1043])).
fof(f1043,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP15(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f641])).
fof(f1883,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X2)))))) )),
  inference(subsumption_resolution,[],[f1879,f1278])).
fof(f1879,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X0) | s('type.ast.pat',X1) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X2)))))) )),
  inference(resolution,[],[f1558,f1040])).
fof(f1040,plain,(
  ( ! [X2,X0,X3,X1] : (sP16(X0,X1,X2) | ~p(s(bool,'T_0')) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3)))))) )),
  inference(cnf_transformation,[],[f638])).
fof(f638,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK80(X2,X1,X0)))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP16(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK80])],[f637])).
fof(f637,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | ? [X4] : (s('type.ast.pat',X2) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X4))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X1) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) & p(s(bool,'T_0')))) & (! [X3] : (s('type.ast.pat',X2) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X3))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X1) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~p(s(bool,'T_0'))) | sP16(X0,X1,X2)))),
  inference(rectify,[],[f636])).
fof(f636,plain,(
  ! [X22,X21,X20] : ((~sP16(X22,X21,X20) | ? [X33] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X33))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0')))) & (! [X33] : (s('type.ast.pat',X20) != s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X33))))) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t',X21) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) | ~p(s(bool,'T_0'))) | sP16(X22,X21,X20)))),
  inference(nnf_transformation,[],[f471])).
fof(f471,plain,(
  ! [X22,X21,X20] : (sP16(X22,X21,X20) <=> ? [X33] : (s('type.ast.pat',X20) = s('type.ast.pat','const.ast.Plit_1'(s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X33))))) & s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t',X21) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X22) & p(s(bool,'T_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP16])])).
fof(f1558,plain,(
  ( ! [X0,X1] : (~sP16(sK56,X0,X1)) )),
  inference(subsumption_resolution,[],[f1557,f768])).
fof(f1557,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) | ~sP16(sK56,X0,X1)) )),
  inference(forward_demodulation,[],[f1506,f775])).
fof(f1506,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0'),s('type.typeSystem.tenv_val',sK57)))))) | ~sP16(sK56,X0,X1)) )),
  inference(superposition,[],[f769,f1038])).
fof(f1038,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X0) | ~sP16(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f638])).
fof(f1865,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1716,f1000])).
fof(f1716,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f1401,f999])).
fof(f999,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f619])).
fof(f1864,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1715,f1000])).
fof(f1715,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f1401,f997])).
fof(f997,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f618])).
fof(f1863,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1706,f1401])).
fof(f1706,plain,(
  ( ! [X318] : (~p(s(bool,X318))) ) | $spl12),
  inference(resolution,[],[f1401,f933])).
fof(f1714,plain,(
  $false | $spl12),
  inference(resolution,[],[f1401,f995])).
fof(f1739,plain,(
  $false | $spl12),
  inference(resolution,[],[f1401,f769])).
fof(f1550,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) ) | $spl14),
  inference(cnf_transformation,[],[f1550_D])).
fof(f1550_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1501,plain,(
  ( ! [X158,X159] : (p(s(bool,X158)) | ~p(s(bool,X159)) | s(bool,X159) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) )),
  inference(resolution,[],[f769,f1274])).
fof(f1416,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f769,f802])).
fof(f1415,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) )),
  inference(resolution,[],[f769,f801])).
fof(f1413,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) )),
  inference(resolution,[],[f769,f800])).
fof(f1551,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) ) | $spl16),
  inference(cnf_transformation,[],[f1551_D])).
fof(f1551_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1412,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f769,f800])).
fof(f1411,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f769,f798])).
fof(f1410,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57))))) | s(bool,X9) != s(bool,X10)) )),
  inference(resolution,[],[f769,f796])).
fof(f1409,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val','const.typeSystem.bind_var_list_3'(s('type.num.num',sK52),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK56),s('type.typeSystem.tenv_val',sK57)))))) )),
  inference(resolution,[],[f769,f796])).
fof(f1395,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) ) | $spl0),
  inference(cnf_transformation,[],[f1395_D])).
fof(f1395_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1371,plain,(
  ( ! [X152,X151] : (p(s(bool,X151)) | p(s(bool,X152)) | s(bool,X152) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f1274])).
fof(f1292,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f768,f802])).
fof(f1397,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) ) | $spl4),
  inference(cnf_transformation,[],[f1397_D])).
fof(f1397_D,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1291,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28) | s(bool,X29) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f802])).
fof(f1289,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f801])).
fof(f1288,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f800])).
fof(f1287,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f768,f798])).
fof(f1286,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,X14) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f798])).
fof(f1285,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK57)))) )),
  inference(resolution,[],[f768,f796])).
% SZS output end Proof for thm.typeSysProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 908
% Passive clauses: 3843
% Generated clauses: 20054
% Final active clauses: 898
% Final passive clauses: 3051
% Input formulas: 261
% Initial clauses: 507
% 
% Selected by SInE selection: 94
% SInE iterations: 1
% 
% Duplicate literals: 852
% Trivial inequalities: 122
% Fw subsumption resolutions: 2348
% Fw demodulations: 3786
% Bw demodulations: 3
% 
% Simple tautologies: 1475
% Equational tautologies: 2
% Forward subsumptions: 7561
% 
% Binary resolution: 6533
% Forward superposition: 4093
% Backward superposition: 2295
% Equality resolution: 22
% 
% Split clauses: 178
% Split components: 397
% SAT solver clauses: 214
% SAT solver unit clauses: 23
% SAT solver binary clauses: 144
% 
% Sat splits: 178
% Sat splitting refutations: 23
% 
% Memory used [KB]: 11129
% Time elapsed: 0.424 s
% ------------------------------
% ------------------------------
% Success in time 13.53 s
