Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.weakening
% (141855)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 731
% Passive clauses: 6290
% Generated clauses: 16462
% Final active clauses: 684
% Final passive clauses: 5153
% Input formulas: 37
% Initial clauses: 149
% 
% Duplicate literals: 553
% Trivial inequalities: 26
% Fw subsumption resolutions: 1713
% Fw demodulations: 326
% Bw demodulations: 46
% 
% Simple tautologies: 191
% Equational tautologies: 60
% Forward subsumptions: 6640
% 
% Binary resolution: 7965
% Factoring: 1354
% Forward superposition: 2474
% Backward superposition: 1552
% Self superposition: 26
% Equality factoring: 28
% Equality resolution: 6
% 
% Split clauses: 625
% Split components: 1268
% SAT solver clauses: 1033
% SAT solver unit clauses: 14
% SAT solver binary clauses: 115
% 
% Sat splits: 625
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1030
% TWLsolver calls for satisfiability: 463
% 
% Memory used [KB]: 7803
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.weakening
% (141878)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 379
% Passive clauses: 2026
% Generated clauses: 10487
% Final active clauses: 353
% Final passive clauses: 1458
% Input formulas: 37
% Initial clauses: 149
% 
% Duplicate literals: 161
% Trivial inequalities: 50
% Fw subsumption resolutions: 732
% Fw demodulations: 192
% Bw demodulations: 211
% Condensations: 659
% Global subsumptions: 13
% 
% Simple tautologies: 191
% Equational tautologies: 16
% Forward subsumptions: 4055
% 
% Binary resolution: 2870
% Factoring: 91
% Forward superposition: 2044
% Backward superposition: 1151
% Self superposition: 48
% Equality factoring: 28
% Equality resolution: 6
% 
% Split clauses: 178
% Split components: 356
% SAT solver clauses: 2558
% SAT solver unit clauses: 162
% SAT solver binary clauses: 70
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 178
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2506
% TWLsolver calls for satisfiability: 1508
% 
% Memory used [KB]: 10490
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.weakening
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.weakening
% SZS output start Proof for thm.weakening
fof(f4118,plain,(
  $false),
  inference(subsumption_resolution,[],[f4117,f363])).
fof(f363,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f213,f210])).
fof(f210,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.weakening.type_e_weakening|split|2',aHL_TRUTH)).
fof(f213,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.weakening.type_e_weakening|split|2','thm.bool.NOT_CLAUSES|split|2')).
fof(f4117,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f4116,f465])).
fof(f465,plain,(
  s(bool,'F_0') = s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12)))),
  inference(unit_resulting_resolution,[],[f209,f228])).
fof(f228,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.weakening.type_e_weakening|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f209,plain,(
  ~p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12))))),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10)))) & ~p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f94])).
fof(f94,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) & ~p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0))))) & ~p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0))))) => p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ((p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X8)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X9),s('type.typeSystem.type_environment',X6))))) => p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ((p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X8)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X9),s('type.typeSystem.type_environment',X6))))) => p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X8)))))),
  file('thm.weakening.type_e_weakening|split|2','thm.weakening.type_e_weakening|split|2')).
fof(f4116,plain,(
  p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12))))),
  inference(subsumption_resolution,[],[f4110,f210])).
fof(f4110,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12))))),
  inference(superposition,[],[f1355,f395])).
fof(f395,plain,(
  s(bool,'T_0') = s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10)))),
  inference(unit_resulting_resolution,[],[f208,f224])).
fof(f224,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.weakening.type_e_weakening|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f208,plain,(
  p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10))))),
  inference(cnf_transformation,[],[f132])).
fof(f1355,plain,(
  ( ! [X0] : (~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X0),s('type.typeSystem.type_environment',sK10)))) | p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12))))) )),
  inference(subsumption_resolution,[],[f1315,f210])).
fof(f1315,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X0),s('type.typeSystem.type_environment',sK10)))) | p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12))))) )),
  inference(superposition,[],[f298,f441])).
fof(f441,plain,(
  s(bool,'T_0') = s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12)))),
  inference(unit_resulting_resolution,[],[f207,f224])).
fof(f207,plain,(
  p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),sK11),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),sK12))))),
  inference(cnf_transformation,[],[f132])).
fof(f298,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) | p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))) | ! [X3] : (~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) | p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2)))) => ! [X3] : (p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) => p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X2))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X6),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X8)))) => ! [X9] : (p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X9),s('type.typeSystem.type_environment',X6)))) => p(s(bool,'const.typeSystem.type_funs_3'(s('type.typeSystem.type_environment',X9),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.ast.exp'))),X7),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.ast.t')),X8))))))),
  file('thm.weakening.type_e_weakening|split|2','thm.weakening.weakening_29_prove_ltb|split|2')).
% SZS output end Proof for thm.weakening
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 398
% Passive clauses: 650
% Generated clauses: 3861
% Final active clauses: 377
% Final passive clauses: 202
% Input formulas: 37
% Initial clauses: 149
% 
% Duplicate literals: 149
% Trivial inequalities: 16
% Fw subsumption resolutions: 228
% Fw demodulations: 324
% Bw demodulations: 70
% Condensations: 112
% Global subsumptions: 3
% 
% Simple tautologies: 57
% Equational tautologies: 53
% Forward subsumptions: 2202
% 
% Unit resulting resolution: 1514
% Forward superposition: 599
% Backward superposition: 554
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 1
% 
% Split clauses: 59
% Split components: 123
% SAT solver clauses: 775
% SAT solver unit clauses: 217
% SAT solver binary clauses: 150
% 
% Sat splits: 59
% Sat splitting refutations: 2
% 
% Memory used [KB]: 2814
% Time elapsed: 0.115 s
% ------------------------------
% ------------------------------
% Success in time 1.452 s
