Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.decSem
% (136169)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 160
% Passive clauses: 3142
% Generated clauses: 21145
% Final active clauses: 129
% Final passive clauses: 2761
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 712
% Trivial inequalities: 49
% Fw subsumption resolutions: 1407
% Fw demodulations: 3033
% Bw demodulations: 24
% 
% Simple tautologies: 48
% Equational tautologies: 45
% Forward subsumptions: 12296
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 457
% Factoring: 2
% Forward superposition: 7820
% Backward superposition: 7437
% Self superposition: 231
% Equality factoring: 78
% Equality resolution: 1
% 
% Split clauses: 185
% Split components: 370
% SAT solver clauses: 594
% SAT solver unit clauses: 40
% SAT solver binary clauses: 120
% 
% Sat splits: 185
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 594
% TWLsolver calls for satisfiability: 140
% 
% Memory used [KB]: 21236
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.decSem
% (136205)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 129
% Passive clauses: 1116
% Generated clauses: 17121
% Final active clauses: 92
% Final passive clauses: 704
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 623
% Trivial inequalities: 105
% Fw subsumption resolutions: 959
% Fw demodulations: 2016
% Bw demodulations: 308
% Condensations: 1558
% 
% Simple tautologies: 118
% Equational tautologies: 44
% Forward subsumptions: 8461
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 365
% Factoring: 1
% Forward superposition: 5188
% Backward superposition: 4465
% Self superposition: 184
% Equality factoring: 98
% Equality resolution: 8
% 
% Split clauses: 92
% Split components: 184
% SAT solver clauses: 1755
% SAT solver unit clauses: 131
% SAT solver binary clauses: 42
% 
% Sat splits: 92
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 1755
% TWLsolver calls for satisfiability: 272
% 
% Memory used [KB]: 15991
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.decSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.decSem
% SZS output start Proof for thm.decSem
fof(f21271,plain,(
  $false),
  inference(subsumption_resolution,[],[f21270,f216])).
fof(f216,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.decSem.state_fupdfupds_comp|split|5',aHL_TRUTH)).
fof(f21270,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f21268,f4436])).
fof(f4436,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(condensation,[],[f4435])).
fof(f4435,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f4363,f300])).
fof(f300,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f218,f219])).
fof(f219,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.decSem.state_fupdfupds_comp|split|5',aHL_BOOL_CASES)).
fof(f218,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.decSem.state_fupdfupds_comp|split|5',aHL_FALSITY)).
fof(f4363,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f2576,f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f219,f219])).
fof(f2576,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f2500,f258])).
fof(f258,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X8,X16] : ! [X18,X22] : s(X16,X22) = s(X16,'const.pair.SND_1'(s('type.pair.prod'(X8,X16),'const.pair.,_2'(s(X8,X18),s(X16,X22)))))),
  file('thm.decSem.state_fupdfupds_comp|split|5','thm.pair.SND')).
fof(f2500,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3))))))),s(X0,X3))))) )),
  inference(unit_resulting_resolution,[],[f258,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X43,X44,X45] : (p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X44))),s(X43,X45)))) <=> s(X43,X44) = s(X43,X45))),
  file('thm.decSem.state_fupdfupds_comp|split|5',arity_0_HL_EQ)).
fof(f21268,plain,(
  ~p(s(bool,happ(s(fun(fun(sK2,'type.decSem.state'(sK0)),bool),happ(s(fun(fun(sK2,'type.decSem.state'(sK0)),fun(fun(sK2,'type.decSem.state'(sK0)),bool)),'=_0'),s(fun(sK2,'type.decSem.state'(sK0)),'const.combin.o_2'(s(fun('type.decSem.state'(sK3),'type.decSem.state'(sK0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.decSem.state'(sK3)),sK4))))),s(fun(sK2,'type.decSem.state'(sK0)),'const.combin.o_2'(s(fun('type.decSem.state'(sK3),'type.decSem.state'(sK0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.decSem.state'(sK3)),sK4))))))),
  inference(backward_demodulation,[],[f21209,f12785])).
fof(f12785,plain,(
  ~p(s(bool,happ(s(fun(fun(sK2,'type.decSem.state'(sK0)),bool),happ(s(fun(fun(sK2,'type.decSem.state'(sK0)),fun(fun(sK2,'type.decSem.state'(sK0)),bool)),'=_0'),s(fun(sK2,'type.decSem.state'(sK0)),'const.combin.o_2'(s(fun('type.decSem.state'(sK3),'type.decSem.state'(sK0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.decSem.state'(sK3)),sK4))))),s(fun(sK2,'type.decSem.state'(sK0)),'const.combin.o_2'(s(fun('type.decSem.state'(sK1),'type.decSem.state'(sK0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.decSem.state'(sK1)),'const.combin.o_2'(s(fun('type.decSem.state'(sK3),'type.decSem.state'(sK1)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.decSem.state'(sK3)),sK4))))))))),
  inference(unit_resulting_resolution,[],[f215,f269])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f215,plain,(
  s(fun(sK2,'type.decSem.state'(sK0)),'const.combin.o_2'(s(fun('type.decSem.state'(sK3),'type.decSem.state'(sK0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.decSem.state'(sK3)),sK4))) != s(fun(sK2,'type.decSem.state'(sK0)),'const.combin.o_2'(s(fun('type.decSem.state'(sK1),'type.decSem.state'(sK0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.decSem.state'(sK1)),'const.combin.o_2'(s(fun('type.decSem.state'(sK3),'type.decSem.state'(sK1)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.decSem.state'(sK3)),sK4)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK2,'type.decSem.state'(sK0)),'const.combin.o_2'(s(fun('type.decSem.state'(sK3),'type.decSem.state'(sK0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6),s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))))),s(fun(sK2,'type.decSem.state'(sK3)),sK4))) != s(fun(sK2,'type.decSem.state'(sK0)),'const.combin.o_2'(s(fun('type.decSem.state'(sK1),'type.decSem.state'(sK0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK6))),s(fun(sK2,'type.decSem.state'(sK1)),'const.combin.o_2'(s(fun('type.decSem.state'(sK3),'type.decSem.state'(sK1)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(sK3),'type.ffi.ffi_state'(sK1)),sK5))),s(fun(sK2,'type.decSem.state'(sK3)),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f193])).
fof(f193,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X3),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.decSem.state'(X3)),X4))) != s(fun(X2,'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.decSem.state'(X1)),'const.combin.o_2'(s(fun('type.decSem.state'(X3),'type.decSem.state'(X1)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.decSem.state'(X3)),X4)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X3),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.decSem.state'(X3)),X4))) = s(fun(X2,'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.decSem.state'(X1)),'const.combin.o_2'(s(fun('type.decSem.state'(X3),'type.decSem.state'(X1)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.decSem.state'(X3)),X4)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X3),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6),s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))))),s(fun(X2,'type.decSem.state'(X3)),X4))) = s(fun(X2,'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X6))),s(fun(X2,'type.decSem.state'(X1)),'const.combin.o_2'(s(fun('type.decSem.state'(X3),'type.decSem.state'(X1)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X1)),X5))),s(fun(X2,'type.decSem.state'(X3)),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X8,'type.decSem.state'(X6)),'const.combin.o_2'(s(fun('type.decSem.state'(X9),'type.decSem.state'(X6)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))),s(fun(X8,'type.decSem.state'(X9)),X10))) = s(fun(X8,'type.decSem.state'(X6)),'const.combin.o_2'(s(fun('type.decSem.state'(X7),'type.decSem.state'(X6)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun(X8,'type.decSem.state'(X7)),'const.combin.o_2'(s(fun('type.decSem.state'(X9),'type.decSem.state'(X7)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))),s(fun(X8,'type.decSem.state'(X9)),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X8,'type.decSem.state'(X6)),'const.combin.o_2'(s(fun('type.decSem.state'(X9),'type.decSem.state'(X6)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))),s(fun(X8,'type.decSem.state'(X9)),X10))) = s(fun(X8,'type.decSem.state'(X6)),'const.combin.o_2'(s(fun('type.decSem.state'(X7),'type.decSem.state'(X6)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun(X8,'type.decSem.state'(X7)),'const.combin.o_2'(s(fun('type.decSem.state'(X9),'type.decSem.state'(X7)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))),s(fun(X8,'type.decSem.state'(X9)),X10)))))),
  file('thm.decSem.state_fupdfupds_comp|split|5','thm.decSem.state_fupdfupds_comp|split|5')).
fof(f21209,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X5,'type.decSem.state'(X1)),'const.combin.o_2'(s(fun('type.decSem.state'(X0),'type.decSem.state'(X1)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X3),s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X2)),X4))))),s(fun(X5,'type.decSem.state'(X0)),X6))) = s(fun(X5,'type.decSem.state'(X1)),'const.combin.o_2'(s(fun('type.decSem.state'(X2),'type.decSem.state'(X1)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X3))),s(fun(X5,'type.decSem.state'(X2)),'const.combin.o_2'(s(fun('type.decSem.state'(X0),'type.decSem.state'(X2)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X2)),X4))),s(fun(X5,'type.decSem.state'(X0)),X6)))))) )),
  inference(superposition,[],[f277,f272])).
fof(f272,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X2),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.decSem.state'(X1),'type.decSem.state'(X2)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X2),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.decSem.state'(X1),'type.decSem.state'(X2)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4),s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3))))) = s(fun('type.decSem.state'(X1),'type.decSem.state'(X0)),'const.combin.o_2'(s(fun('type.decSem.state'(X2),'type.decSem.state'(X0)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),X4))),s(fun('type.decSem.state'(X1),'type.decSem.state'(X2)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X2)),X3)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X9,X7] : ! [X11,X12] : s(fun('type.decSem.state'(X9),'type.decSem.state'(X6)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X6)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12),s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11))))) = s(fun('type.decSem.state'(X9),'type.decSem.state'(X6)),'const.combin.o_2'(s(fun('type.decSem.state'(X7),'type.decSem.state'(X6)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X12))),s(fun('type.decSem.state'(X9),'type.decSem.state'(X7)),'const.decSem.state_ffi_fupd_1'(s(fun('type.ffi.ffi_state'(X9),'type.ffi.ffi_state'(X7)),X11)))))),
  file('thm.decSem.state_fupdfupds_comp|split|5','thm.decSem.state_fupdfupds_comp|split|4')).
fof(f277,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X16,X8,X21,X17] : ! [X12,X11,X10] : s(fun(X21,X16),'const.combin.o_2'(s(fun(X8,X16),X12),s(fun(X21,X8),'const.combin.o_2'(s(fun(X17,X8),X11),s(fun(X21,X17),X10))))) = s(fun(X21,X16),'const.combin.o_2'(s(fun(X17,X16),'const.combin.o_2'(s(fun(X8,X16),X12),s(fun(X17,X8),X11))),s(fun(X21,X17),X10)))),
  file('thm.decSem.state_fupdfupds_comp|split|5','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.decSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 170
% Passive clauses: 1768
% Generated clauses: 22694
% Final active clauses: 103
% Final passive clauses: 1036
% Input formulas: 68
% Initial clauses: 79
% 
% Duplicate literals: 638
% Trivial inequalities: 49
% Fw subsumption resolutions: 653
% Fw demodulations: 4007
% Bw demodulations: 537
% Condensations: 977
% Global subsumptions: 334
% 
% Simple tautologies: 125
% Equational tautologies: 135
% Forward subsumptions: 13573
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 26
% 
% Unit resulting resolution: 734
% Factoring: 1
% Forward superposition: 6567
% Backward superposition: 6039
% Self superposition: 200
% Equality factoring: 76
% Equality resolution: 2
% 
% Split clauses: 28
% Split components: 56
% SAT solver clauses: 2510
% SAT solver unit clauses: 766
% SAT solver binary clauses: 906
% 
% Sat splits: 28
% Sat splitting refutations: 350
% 
% Memory used [KB]: 26993
% Time elapsed: 0.936 s
% ------------------------------
% ------------------------------
% Success in time 2.27 s
