Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvl_to_bviProof
% (139664)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 166
% Passive clauses: 1791
% Generated clauses: 16697
% Final active clauses: 153
% Final passive clauses: 1471
% Input formulas: 22
% Initial clauses: 56
% 
% Duplicate literals: 835
% Trivial inequalities: 361
% Fw subsumption resolutions: 2977
% Fw demodulations: 2066
% Bw demodulations: 9
% 
% Simple tautologies: 411
% Equational tautologies: 43
% Forward subsumptions: 6474
% 
% Binary resolution: 468
% Forward superposition: 4731
% Backward superposition: 5080
% Self superposition: 161
% Equality factoring: 132
% Equality resolution: 42
% 
% Split clauses: 1224
% Split components: 2574
% SAT solver clauses: 1891
% SAT solver unit clauses: 24
% SAT solver binary clauses: 221
% 
% Sat splits: 1224
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1890
% TWLsolver calls for satisfiability: 203
% 
% Memory used [KB]: 8571
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvl_to_bviProof
% (139681)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 155
% Passive clauses: 744
% Generated clauses: 16552
% Final active clauses: 137
% Final passive clauses: 430
% Input formulas: 22
% Initial clauses: 56
% 
% Duplicate literals: 1049
% Trivial inequalities: 353
% Fw subsumption resolutions: 1689
% Fw demodulations: 1671
% Bw demodulations: 133
% Condensations: 1481
% Global subsumptions: 4
% 
% Simple tautologies: 377
% Equational tautologies: 41
% Forward subsumptions: 6575
% Bw demodulations to eq. taut.: 44
% 
% Binary resolution: 306
% Forward superposition: 4732
% Backward superposition: 5077
% Self superposition: 183
% Equality factoring: 144
% Equality resolution: 42
% 
% Split clauses: 164
% Split components: 340
% SAT solver clauses: 1277
% SAT solver unit clauses: 51
% SAT solver binary clauses: 35
% SAT solver learnt clauses: 8
% 
% Sat splits: 164
% 
% TWLsolver clauses: 1261
% TWLsolver calls for satisfiability: 387
% 
% Memory used [KB]: 11129
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvl_to_bviProof
% (139775)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 610
% Passive clauses: 3723
% Generated clauses: 192809
% Final active clauses: 170
% Final passive clauses: 1187
% Input formulas: 22
% Initial clauses: 56
% 
% Duplicate literals: 11083
% Trivial inequalities: 2205
% Fw subsumption resolutions: 21652
% Fw demodulations: 26201
% Bw demodulations: 1779
% Condensations: 20613
% Global subsumptions: 119
% 
% Simple tautologies: 2057
% Equational tautologies: 1589
% Forward subsumptions: 103329
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 165
% 
% Unit resulting resolution: 5851
% Forward superposition: 52318
% Backward superposition: 44098
% Self superposition: 863
% Equality factoring: 482
% Equality resolution: 86
% 
% Split clauses: 353
% Split components: 718
% SAT solver clauses: 4717
% SAT solver unit clauses: 34
% SAT solver binary clauses: 370
% 
% Sat splits: 353
% Sat splitting refutations: 167
% 
% Memory used [KB]: 69593
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.bvl_to_bviProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvl_to_bviProof
% SZS output start Proof for thm.bvl_to_bviProof
fof(f14559,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f144,f148,f151,f150,f146,f156,f157,f158,f159,f160,f162,f180,f163,f170,f178,f189,f190,f191,f192,f193,f194,f196,f203,f217,f218,f219,f220,f221,f223,f212,f224,f231,f239,f258,f309,f310,f311,f312,f313,f332,f314,f316,f323,f331,f414,f418,f422,f427,f434,f438,f473,f474,f507,f508,f500,f499,f513,f514,f515,f516,f517,f519,f541,f520,f527,f540,f560,f561,f586,f587,f618,f613,f612,f619,f719,f722,f726,f731,f738,f742,f181,f762,f764,f765,f767,f769,f776,f807,f821,f823,f824,f825,f828,f835,f867,f1104,f1106,f1107,f1109,f1111,f1118,f1151,f1260,f1262,f1263,f1265,f1267,f1274,f1307,f1320,f1322,f1323,f1324,f1327,f1334,f1368,f1558,f1560,f1561,f1563,f1565,f1572,f1606,f1749,f1750,f1786,f1788,f1789,f1791,f1793,f1800,f1813,f1815,f1816,f1818,f1820,f1827,f1837,f1876,f1878,f1879,f1881,f1883,f1890,f1924,f1943,f1945,f1946,f1947,f1950,f1957,f1996,f2009,f1940,f636,f149,f2013,f2015,f2016,f2017,f2020,f2027,f2039,f2194,f2195,f2207,f2209,f2210,f2212,f2214,f2221,f2233,f2238,f2240,f2241,f2242,f2245,f2252,f2270,f2301,f2295,f2326,f2327,f2649,f2650,f2661,f2663,f2664,f2665,f2668,f2675,f2687,f2726,f2728,f2729,f2731,f2733,f2740,f2767,f2769,f2770,f2771,f2774,f2781,f2940,f2941,f3187,f3254,f3256,f3257,f3258,f3261,f3268,f3314,f3274,f3276,f3277,f3279,f3281,f3288,f3315,f3479,f3481,f3482,f3483,f3486,f3493,f3541,f3499,f3501,f3502,f3503,f3506,f3513,f3542,f3551,f3553,f3554,f3556,f3558,f3565,f2751,f3591,f3593,f3594,f3596,f3598,f3605,f3627,f3629,f3630,f3631,f3634,f3641,f3753,f3648,f3650,f3651,f3653,f3655,f3662,f3754,f3669,f3671,f3672,f3674,f3676,f3683,f3755,f3785,f3787,f3788,f3789,f3792,f3799,f3814,f3837,f3839,f3840,f3841,f3844,f3851,f3916,f3857,f3859,f3860,f3862,f3864,f3871,f3917,f3877,f3879,f3880,f3882,f3884,f3891,f3918,f3926,f3928,f3929,f3930,f3933,f3940,f3952,f3956,f3958,f3959,f3961,f3963,f3970,f3981,f3984,f3986,f3987,f3988,f3991,f3998,f4011,f4016,f4018,f4019,f4020,f4023,f4030,f4104,f4039,f4041,f4042,f4044,f4046,f4053,f4107,f4059,f4061,f4062,f4064,f4066,f4073,f4108,f4170,f4172,f4173,f4175,f4177,f4184,f4223,f4314,f4316,f4317,f4319,f4321,f4328,f1810,f4341,f4343,f4344,f4345,f4348,f4355,f2793,f4372,f4374,f4375,f4376,f4379,f4386,f4399,f4404,f4406,f4407,f4408,f4411,f4418,f4431,f4596,f4598,f4599,f4600,f4603,f4610,f4623,f4628,f4630,f4631,f4633,f4635,f4642,f4660,f5076,f5077,f5121,f5256,f5258,f5261,f5263,f5265,f5269,f3623,f5375,f5377,f5378,f5379,f5382,f5389,f5431,f5433,f5434,f5435,f5438,f5445,f5461,f5466,f5468,f5469,f5470,f5473,f5480,f5530,f5488,f5490,f5491,f5492,f5495,f5502,f5531,f5543,f5545,f5546,f5548,f5550,f5557,f5577,f5581,f5583,f5584,f5586,f5588,f5595,f5606,f5643,f5645,f5646,f5648,f5650,f5657,f5668,f5691,f5693,f5694,f5696,f5698,f5705,f5719,f5757,f5759,f5760,f5761,f5764,f5771,f5818,f5779,f5781,f5782,f5783,f5786,f5793,f5819,f5827,f5829,f5830,f5831,f5834,f5841,f5854,f5858,f5860,f5861,f5863,f5865,f5872,f5883,f5927,f5929,f5930,f5931,f5934,f5941,f5991,f5949,f5951,f5952,f5953,f5956,f5963,f5992,f6054,f6055,f6056,f6057,f6058,f6059,f6060,f6061,f6317,f6318,f6331,f6332,f6442,f6444,f6445,f6447,f6449,f6456,f6467,f6480,f6482,f6483,f6485,f6487,f6494,f6508,f6516,f6518,f6519,f6520,f6523,f6530,f6545,f6550,f6552,f6553,f6555,f6557,f6564,f6585,f6649,f6651,f6652,f6654,f6656,f6663,f6677,f6685,f6687,f6688,f6689,f6692,f6699,f6738,f6740,f6741,f6743,f6745,f6752,f6764,f6778,f6780,f6781,f6783,f6785,f6792,f6806,f6978,f6979,f6989,f6990,f7011,f7013,f7014,f7016,f7018,f7025,f7039,f7047,f7049,f7050,f7051,f7054,f7061,f7091,f7093,f7094,f7096,f7098,f7105,f7116,f7120,f7122,f7123,f7124,f7127,f7134,f7172,f7174,f7175,f7177,f7179,f7186,f7200,f7225,f7227,f7228,f7229,f7232,f7239,f7253,f7274,f7276,f7277,f7278,f7281,f7288,f7303,f7312,f7314,f7315,f7317,f7319,f7326,f7365,f7383,f7385,f7386,f7388,f7390,f7397,f7461,f7511,f7513,f7514,f7516,f7518,f7525,f7538,f7590,f7592,f7806,f7808,f7809,f7811,f7813,f7820,f7837,f7840,f7842,f7843,f7845,f7847,f7854,f7874,f7926,f8007,f8009,f8010,f8012,f8014,f8021,f8040,f8203,f8205,f8208,f8210,f8212,f8214,f5409,f8218,f8227,f8228,f8236,f8238,f8239,f8240,f8243,f8250,f8281,f8283,f8284,f8286,f8288,f8295,f8306,f8330,f8332,f8333,f8334,f8337,f8344,f8362,f8379,f8381,f8382,f8384,f8386,f8393,f8404,f8429,f8431,f8432,f8433,f8436,f8443,f8461,f8466,f8468,f8469,f8471,f8473,f8480,f8492,f8515,f8517,f8518,f8520,f8522,f8529,f8546,f8574,f8576,f8577,f8579,f8581,f8588,f8608,f8620,f8622,f8623,f8624,f8627,f8634,f8660,f8662,f8663,f8664,f8667,f8674,f8712,f8714,f8715,f8717,f8719,f8726,f8748,f8940,f8941,f8957,f8958,f8978,f8980,f8981,f8983,f8985,f8992,f9012,f9021,f9023,f9024,f9026,f9028,f9035,f9055,f9057,f9058,f9060,f9062,f9069,f9081,f9145,f9147,f9148,f9150,f9152,f9159,f9052,f9175,f9177,f9178,f9179,f9182,f9189,f9249,f9204,f9206,f9207,f9208,f9211,f9218,f9255,f9266,f9268,f9269,f9271,f9273,f9280,f9293,f9338,f9340,f9341,f9342,f9345,f9352,f9366,f9381,f9383,f9384,f9386,f9388,f9395,f9408,f9418,f9420,f9421,f9423,f9425,f9432,f9475,f9532,f9534,f9535,f9537,f9539,f9546,f9557,f9699,f9701,f9704,f9706,f9708,f9710,f6712,f9714,f9723,f9724,f9869,f9871,f9874,f9876,f9878,f9882,f8266,f9886,f9887,f9895,f9896,f10042,f10044,f10047,f10049,f10051,f10055,f7148,f10059,f10060,f10068,f10069,f10090,f10092,f10093,f10095,f10097,f10104,f10115,f10142,f10144,f10145,f10147,f10149,f10156,f10173,f10316,f10318,f10321,f10323,f10325,f10327,f7075,f10331,f10340,f10341,f10350,f10352,f10353,f10355,f10357,f10364,f10384,f10415,f10417,f10418,f10420,f10422,f10429,f10476,f10586,f10588,f10589,f10591,f10593,f10600,f10647,f10683,f10685,f10686,f10688,f10690,f10697,f10710,f10714,f10716,f10717,f10719,f10721,f10728,f10741,f10743,f10745,f10746,f10747,f10750,f10757,f10854,f10882,f10884,f10885,f10887,f10889,f10896,f10911,f10948,f10950,f10951,f10953,f10955,f10962,f10977,f11081,f11083,f11084,f11085,f11088,f11095,f11224,f11103,f11105,f11106,f11107,f11110,f11117,f11225,f11386,f11388,f11389,f11391,f11393,f11400,f11413,f11416,f11418,f11419,f11421,f11423,f11430,f11478,f11490,f11492,f11493,f11495,f11497,f11504,f11586,f11679,f11681,f11682,f11684,f11686,f11693,f11713,f11905,f11907,f11908,f11910,f11912,f11919,f11934,f11954,f11956,f11957,f11959,f11961,f11968,f11982,f12052,f12054,f12055,f12056,f12059,f12066,f12161,f12087,f12089,f12090,f12091,f12094,f12101,f12162,f12293,f12295,f12296,f12298,f12300,f12307,f12320,f12381,f12383,f12384,f12386,f12388,f12395,f12413,f12459,f12461,f12462,f12464,f12466,f12473,f12486,f12653,f12655,f12656,f12658,f12660,f12667,f12691,f12718,f12720,f12721,f12723,f12725,f12732,f12745,f12761,f12763,f12764,f12765,f12768,f12775,f12802,f12814,f12816,f12817,f12819,f12821,f12828,f12902,f7927,f12907,f12909,f12910,f12912,f12914,f12921,f12848,f12941,f12943,f12944,f12945,f12948,f12955,f12976,f13199,f13200,f13203,f13205,f13208,f13210,f13212,f13216,f8648,f13220,f13229,f13230,f13287,f13289,f13294,f13295,f13296,f13297,f13300,f13301,f13302,f13303,f13310,f13311,f13312,f13524,f415,f13525,f13528,f13530,f13533,f13535,f13537,f13541,f8688,f13545,f13554,f13555,f1751,f13612,f13614,f13615,f13617,f13619,f13626,f13645,f13663,f13665,f13666,f13668,f13670,f13677,f13690,f13709,f13711,f13712,f13713,f13716,f13723,f13738,f13750,f13752,f13753,f13755,f13757,f13764,f13782,f13785,f13607,f13791,f11065,f4086,f1874,f2010,f632,f2011,f145,f13471,f13546,f13155,f13221,f12644,f13795,f12642,f12640,f13798,f12638,f13799,f12633,f13802,f12627,f12630,f12628,f13801,f13796,f12596,f13792,f12594,f12592,f153,f12563,f8219,f12561,f9715,f12559,f10332,f12556,f13804,f12555,f13807,f12553,f13808,f12551,f13809,f12550,f13793,f5122,f12577,f12575,f12573,f13806,f13803,f13811,f13813,f13814,f13816,f13818,f13825,f211,f14001,f14003,f14004,f14006,f14008,f14015,f14029,f14082,f14083,f14085,f14087,f14088,f14303,f14305,f14306,f501,f14308,f14433,f14435,f14436,f14437,f14440,f14447,f14463,f14506,f14507,f14509,f14511,f14512,f14513,f14522,f14524,f14525,f179,f14527,f183,f14529,f182,f14536,f184,f14550,f14558])).
fof(f14558,plain,(
  $false | $spl66),
  inference(subsumption_resolution,[],[f14557,f141])).
fof(f141,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb',aHL_TRUTH)).
fof(f14557,plain,(
  ~p(s(bool,'T_0')) | $spl66),
  inference(forward_demodulation,[],[f14556,f523])).
fof(f523,plain,(
  s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))),
  inference(resolution,[],[f88,f113])).
fof(f113,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f88,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))))),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',sK3),s('type.bvlSem.state'(sK0),sK1))))) & (~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) | ~p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3)))) | (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f53])).
fof(f53,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X6)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X6),s('type.bvlSem.state'(X0),X2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1))))) & (~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) | ~p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3)))) | (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.list.list'('type.bvlSem.v'),X5)))))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X6)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X6),s('type.bvlSem.state'(X0),X2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1)))))) & (~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) | ~p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3)))) | (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.list.list'('type.bvlSem.v'),X5)))))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X6)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X6),s('type.bvlSem.state'(X0),X2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1)))))) => (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3)))) & (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.list.list'('type.bvlSem.v'),X5)))))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X6)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X6),s('type.bvlSem.state'(X0),X2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1)))))) => (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3)))) & (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.list.list'('type.bvlSem.v'),X5)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12] : ((p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X6),X8)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X8))))),s('type.list.list'('type.bvlSem.v'),X12)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X10),s('type.list.list'('type.bvlSem.v'),X12),s('type.bvlSem.state'(X6),X8))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'const.pair.,_2'(s('type.bvlSem.v',X9),s('type.bvlSem.state'(X6),X7)))))) => (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X6),X7)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X7))))),s('type.bvlSem.v',X9)))) & (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X8))))),s('type.list.list'('type.bvlSem.v'),X11)))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X7))))),s('type.list.list'('type.bvlSem.v'),X11)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12] : ((p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X6),X8)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X8))))),s('type.list.list'('type.bvlSem.v'),X12)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X10),s('type.list.list'('type.bvlSem.v'),X12),s('type.bvlSem.state'(X6),X8))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'const.pair.,_2'(s('type.bvlSem.v',X9),s('type.bvlSem.state'(X6),X7)))))) => (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X6),X7)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X7))))),s('type.bvlSem.v',X9)))) & (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X8))))),s('type.list.list'('type.bvlSem.v'),X11)))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X7))))),s('type.list.list'('type.bvlSem.v'),X11)))))))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb')).
fof(f14556,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) | $spl66),
  inference(subsumption_resolution,[],[f14555,f138])).
fof(f138,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ((p(s(bool,'T_0')) & p(s(bool,X7))) <=> p(s(bool,X7)))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.AND_CLAUSES|split|0')).
fof(f14555,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) | $spl66),
  inference(forward_demodulation,[],[f14554,f166])).
fof(f166,plain,(
  s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))),
  inference(resolution,[],[f87,f113])).
fof(f87,plain,(
  p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2))))),
  inference(cnf_transformation,[],[f64])).
fof(f14554,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) | $spl66),
  inference(equality_resolution,[],[f2011])).
fof(f14550,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) ) | $spl404),
  inference(cnf_transformation,[],[f14550_D])).
fof(f14550_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f184,plain,(
  ( ! [X20] : (p(s(bool,X20))) ) | $spl20),
  inference(cnf_transformation,[],[f184_D])).
fof(f184_D,plain,(
  ( ! [X20] : (p(s(bool,X20))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f14536,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) ) | $spl7),
  inference(resolution,[],[f14028,f142])).
fof(f142,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f102])).
fof(f102,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13,X18,X14,X19] : ((s(bool,X13) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X14) = s(bool,X19))) => ((p(s(bool,X13)) => p(s(bool,X14))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.IMP_CONG')).
fof(f14028,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) | $spl7),
  inference(forward_demodulation,[],[f14000,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',X2))) = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))),s('type.bvlSem.v',X2)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',X2))) = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))),s('type.bvlSem.v',X2)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',X2))) = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))),s('type.bvlSem.v',X2)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X26,X27] : s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X26),s('type.bvlSem.v',X27))) = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X26))),s('type.bvlSem.v',X27)))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','ARITY_1_2_const.bvl_to_bviProof.bv_ok')).
fof(f14000,plain,(
  ~p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) | $spl7),
  inference(resolution,[],[f152,f96])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X23,X24] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X23),s('type.list.list'(X6),X24)))) <=> ! [X25] : (p(s(bool,'const.bool.IN_2'(s(X6,X25),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X24)))))) => p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X25))))))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.list.EVERY_MEM')).
fof(f152,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) | $spl7),
  inference(cnf_transformation,[],[f152_D])).
fof(f152_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f182,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,X16)) ) | $spl16),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,X16)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f14529,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')))) | s(bool,X23) != s(bool,X25)) ) | $spl7),
  inference(resolution,[],[f14028,f105])).
fof(f105,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f183,plain,(
  ( ! [X21] : (~p(s(bool,X21))) ) | $spl18),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  ( ! [X21] : (~p(s(bool,X21))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f14527,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) ) | $spl7),
  inference(resolution,[],[f14028,f104])).
fof(f104,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f179,plain,(
  ( ! [X10,X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,X10)) ) | $spl10),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  ( ! [X10,X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,X10)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f14525,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) ) | $spl7),
  inference(resolution,[],[f14028,f103])).
fof(f103,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f14524,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')))) | s(bool,X8) != s(bool,X10)) ) | $spl7),
  inference(resolution,[],[f14028,f101])).
fof(f101,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f14522,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) ) | $spl7),
  inference(resolution,[],[f14028,f99])).
fof(f99,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f14513,plain,(
  $false | ($spl1 | $spl7 | $spl74)),
  inference(trivial_inequality_removal,[],[f14488])).
fof(f14488,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl1 | $spl7 | $spl74)),
  inference(superposition,[],[f14468,f13819])).
fof(f13819,plain,(
  s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | $spl1),
  inference(resolution,[],[f146,f109])).
fof(f109,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f14468,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8)) ) | ($spl7 | $spl74)),
  inference(subsumption_resolution,[],[f14467,f110])).
fof(f110,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f14467,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8))) ) | ($spl7 | $spl74)),
  inference(forward_demodulation,[],[f2270,f14009])).
fof(f14009,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | $spl7),
  inference(resolution,[],[f152,f109])).
fof(f14512,plain,(
  $false | ($spl7 | $spl74)),
  inference(trivial_inequality_removal,[],[f14490])).
fof(f14490,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl7 | $spl74)),
  inference(superposition,[],[f14468,f14009])).
fof(f14511,plain,(
  $false | ($spl3 | $spl7 | $spl74)),
  inference(trivial_inequality_removal,[],[f14491])).
fof(f14491,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl7 | $spl74)),
  inference(superposition,[],[f14468,f2215])).
fof(f2215,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5))) | $spl3),
  inference(resolution,[],[f148,f109])).
fof(f14509,plain,(
  $false | ($spl3 | $spl7 | $spl74)),
  inference(trivial_inequality_removal,[],[f14498])).
fof(f14498,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl7 | $spl74)),
  inference(superposition,[],[f14468,f3815])).
fof(f3815,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | $spl3),
  inference(subsumption_resolution,[],[f3807,f140])).
fof(f140,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb',aHL_FALSITY)).
fof(f3807,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | $spl3),
  inference(superposition,[],[f3282,f2215])).
fof(f3282,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'const.list.EVERY_2'(s(fun(X48,bool),X49),s('type.list.list'(X48),X50)))) | s(bool,'F_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,sK7(X50,X49,X48))))) )),
  inference(resolution,[],[f96,f109])).
fof(f14507,plain,(
  $false | ($spl3 | $spl7 | $spl74)),
  inference(trivial_inequality_removal,[],[f14501])).
fof(f14501,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl7 | $spl74)),
  inference(superposition,[],[f14468,f2734])).
fof(f2734,plain,(
  s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | $spl3),
  inference(resolution,[],[f2232,f109])).
fof(f2232,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | $spl3),
  inference(forward_demodulation,[],[f2206,f97])).
fof(f2206,plain,(
  ~p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | $spl3),
  inference(resolution,[],[f148,f96])).
fof(f14506,plain,(
  $false | ($spl7 | $spl74)),
  inference(equality_resolution,[],[f14468])).
fof(f14463,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl402),
  inference(cnf_transformation,[],[f14463_D])).
fof(f14463_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f14447,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl7),
  inference(resolution,[],[f13999,f142])).
fof(f13999,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | $spl7),
  inference(resolution,[],[f152,f95])).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f14440,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X27) != s(bool,X28)) ) | $spl7),
  inference(resolution,[],[f13999,f105])).
fof(f14437,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl7),
  inference(resolution,[],[f13999,f104])).
fof(f14436,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl7),
  inference(resolution,[],[f13999,f103])).
fof(f14435,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X12) != s(bool,X13)) ) | $spl7),
  inference(resolution,[],[f13999,f101])).
fof(f14433,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl7),
  inference(resolution,[],[f13999,f99])).
fof(f14308,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')))) | $spl400),
  inference(cnf_transformation,[],[f14308_D])).
fof(f14308_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')))) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f501,plain,(
  ( ! [X57] : (s(bool,'F_0') != s(bool,X57)) ) | $spl34),
  inference(cnf_transformation,[],[f501_D])).
fof(f501_D,plain,(
  ( ! [X57] : (s(bool,'F_0') != s(bool,X57)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f14306,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18) | s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v'))))) ) | ($spl7 | $spl14)),
  inference(subsumption_resolution,[],[f14161,f108])).
fof(f108,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f71])).
fof(f14161,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18) | s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')))) | p(s(bool,X18))) ) | ($spl7 | $spl14)),
  inference(superposition,[],[f3800,f14009])).
fof(f3800,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X78) != s(bool,'const.list.EVERY_2'(s(fun(X75,bool),X76),s('type.list.list'(X75),X77))) | s(bool,'F_0') = s(bool,happ(s(fun(X75,bool),X76),s(X75,sK7(X77,X76,X75)))) | p(s(bool,X78))) ) | $spl14),
  inference(resolution,[],[f3282,f181])).
fof(f14305,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | $spl398),
  inference(cnf_transformation,[],[f14305_D])).
fof(f14305_D,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f14303,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X5))) ) | ($spl7 | $spl90)),
  inference(subsumption_resolution,[],[f14150,f109])).
fof(f14150,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X5))) ) | ($spl7 | $spl90)),
  inference(superposition,[],[f3541,f14009])).
fof(f14088,plain,(
  $false | ($spl1 | $spl26)),
  inference(trivial_inequality_removal,[],[f14065])).
fof(f14065,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl1 | $spl26)),
  inference(superposition,[],[f14045,f13819])).
fof(f14045,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) | ($spl1 | $spl26)),
  inference(subsumption_resolution,[],[f14044,f110])).
fof(f14044,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X9))) ) | ($spl1 | $spl26)),
  inference(forward_demodulation,[],[f239,f13819])).
fof(f14087,plain,(
  $false | ($spl1 | $spl3 | $spl26)),
  inference(trivial_inequality_removal,[],[f14067])).
fof(f14067,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl1 | $spl3 | $spl26)),
  inference(superposition,[],[f14045,f2215])).
fof(f14085,plain,(
  $false | ($spl1 | $spl3 | $spl26)),
  inference(trivial_inequality_removal,[],[f14074])).
fof(f14074,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl1 | $spl3 | $spl26)),
  inference(superposition,[],[f14045,f3815])).
fof(f14083,plain,(
  $false | ($spl1 | $spl3 | $spl26)),
  inference(trivial_inequality_removal,[],[f14077])).
fof(f14077,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl1 | $spl3 | $spl26)),
  inference(superposition,[],[f14045,f2734])).
fof(f14082,plain,(
  $false | ($spl1 | $spl26)),
  inference(equality_resolution,[],[f14045])).
fof(f14029,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl396),
  inference(cnf_transformation,[],[f14029_D])).
fof(f14029_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f14015,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f152,f142])).
fof(f14008,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X22) != s(bool,X24)) ) | $spl7),
  inference(resolution,[],[f152,f105])).
fof(f14006,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f152,f104])).
fof(f14004,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f152,f103])).
fof(f14003,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X7) != s(bool,X9)) ) | $spl7),
  inference(resolution,[],[f152,f101])).
fof(f14001,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f152,f99])).
fof(f211,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl22),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f13825,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl1),
  inference(resolution,[],[f146,f142])).
fof(f13818,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | s(bool,X25) != s(bool,X27)) ) | $spl1),
  inference(resolution,[],[f146,f105])).
fof(f13816,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl1),
  inference(resolution,[],[f146,f104])).
fof(f13814,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl1),
  inference(resolution,[],[f146,f103])).
fof(f13813,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | s(bool,X10) != s(bool,X12)) ) | $spl1),
  inference(resolution,[],[f146,f101])).
fof(f13811,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl1),
  inference(resolution,[],[f146,f99])).
fof(f13803,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X12) != s(bool,X13)) ) | $spl384),
  inference(cnf_transformation,[],[f13803_D])).
fof(f13803_D,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X12) != s(bool,X13)) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f13806,plain,(
  s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | $spl389),
  inference(cnf_transformation,[],[f13806_D])).
fof(f13806_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) <=> ~$spl389),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl389])])).
fof(f12573,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) != s(bool,X1)) ) | ($spl14 | $spl166)),
  inference(subsumption_resolution,[],[f12542,f1102])).
fof(f1102,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) != s(bool,X15)) ) | $spl14),
  inference(subsumption_resolution,[],[f1083,f294])).
fof(f294,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f285,f164])).
fof(f164,plain,(
  s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))),
  inference(resolution,[],[f87,f108])).
fof(f285,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))),
  inference(superposition,[],[f164,f116])).
fof(f116,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb',aHL_BOOL_CASES)).
fof(f1083,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X14))) ) | $spl14),
  inference(superposition,[],[f831,f115])).
fof(f115,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f831,plain,(
  ( ! [X45,X44] : (s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X44)) ) | $spl14),
  inference(resolution,[],[f770,f113])).
fof(f770,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41)) ) | $spl14),
  inference(resolution,[],[f181,f109])).
fof(f12542,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) != s(bool,X1) | p(s(bool,X0))) ) | $spl166),
  inference(superposition,[],[f6057,f109])).
fof(f12575,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | ($spl14 | $spl166)),
  inference(subsumption_resolution,[],[f12543,f1102])).
fof(f12543,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X2))) ) | $spl166),
  inference(superposition,[],[f6057,f111])).
fof(f111,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f72])).
fof(f12577,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) | ($spl14 | $spl166)),
  inference(subsumption_resolution,[],[f12547,f1114])).
fof(f1114,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,X45) | s(bool,X44) != s(bool,X45) | s(bool,'T_0') = s(bool,X44)) ) | $spl14),
  inference(resolution,[],[f1102,f112])).
fof(f112,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f73])).
fof(f12547,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | $spl166),
  inference(superposition,[],[f6057,f116])).
fof(f5122,plain,(
  ( ! [X356,X358,X360,X357] : (s(bool,X360) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X356),X357))),s('type.bvlSem.v',X358))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X356),X357),s('type.bvlSem.v',X358))) | s(bool,'T_0') = s(bool,X360)) ) | $spl134),
  inference(cnf_transformation,[],[f5122_D])).
fof(f5122_D,plain,(
  ( ! [X356,X358,X360,X357] : (s(bool,X360) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X356),X357))),s('type.bvlSem.v',X358))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X356),X357),s('type.bvlSem.v',X358))) | s(bool,'T_0') = s(bool,X360)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f13793,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | $spl371),
  inference(cnf_transformation,[],[f13793_D])).
fof(f13793_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) <=> ~$spl371),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl371])])).
fof(f12550,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X20),X21))),s('type.bvlSem.v',X22))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X20),X21),s('type.bvlSem.v',X22)))) ) | $spl166),
  inference(superposition,[],[f6057,f3326])).
fof(f3326,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X18),X19))),s('type.bvlSem.v',X20))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X18),X19),s('type.bvlSem.v',X20)))) )),
  inference(superposition,[],[f97,f116])).
fof(f13809,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X24),X25),s('type.bvlSem.v',X26))) | s(bool,X27) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X24),X25))),s('type.bvlSem.v',X26)))) ) | $spl394),
  inference(cnf_transformation,[],[f13809_D])).
fof(f13809_D,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X24),X25),s('type.bvlSem.v',X26))) | s(bool,X27) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X24),X25))),s('type.bvlSem.v',X26)))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f12551,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X27) | s(bool,X27) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X24),X25))),s('type.bvlSem.v',X26))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X24),X25),s('type.bvlSem.v',X26)))) ) | $spl166),
  inference(superposition,[],[f6057,f3325])).
fof(f3325,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X15),X16))),s('type.bvlSem.v',X17))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X15),X16),s('type.bvlSem.v',X17)))) )),
  inference(superposition,[],[f97,f116])).
fof(f13808,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X32) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl392),
  inference(cnf_transformation,[],[f13808_D])).
fof(f13808_D,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X32) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f12553,plain,(
  ( ! [X32] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X32) | s(bool,X32) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | ($spl3 | $spl166)),
  inference(superposition,[],[f6057,f3815])).
fof(f13807,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X34,bool),X35),s('type.list.list'(X34),X36))) | s(bool,X37) != s(bool,happ(s(fun(X34,bool),X35),s(X34,sK7(X36,X35,X34))))) ) | $spl390),
  inference(cnf_transformation,[],[f13807_D])).
fof(f13807_D,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X34,bool),X35),s('type.list.list'(X34),X36))) | s(bool,X37) != s(bool,happ(s(fun(X34,bool),X35),s(X34,sK7(X36,X35,X34))))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f12555,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X37) | s(bool,X37) != s(bool,happ(s(fun(X34,bool),X35),s(X34,sK7(X36,X35,X34)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X34,bool),X35),s('type.list.list'(X34),X36)))) ) | $spl166),
  inference(superposition,[],[f6057,f3599])).
fof(f3599,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,sK7(X50,X49,X48)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X48,bool),X49),s('type.list.list'(X48),X50)))) )),
  inference(resolution,[],[f3264,f109])).
fof(f3264,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,happ(s(fun(X54,bool),X55),s(X54,sK7(X56,X55,X54))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X54,bool),X55),s('type.list.list'(X54),X56)))) )),
  inference(resolution,[],[f96,f113])).
fof(f13804,plain,(
  ( ! [X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,X38) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl386),
  inference(cnf_transformation,[],[f13804_D])).
fof(f13804_D,plain,(
  ( ! [X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,X38) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f12556,plain,(
  ( ! [X38] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X38) | s(bool,X38) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | ($spl3 | $spl166)),
  inference(superposition,[],[f6057,f2734])).
fof(f10332,plain,(
  ( ! [X213,X211] : (s(bool,X213) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X211,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X211),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X213)) ) | $spl292),
  inference(cnf_transformation,[],[f10332_D])).
fof(f10332_D,plain,(
  ( ! [X213,X211] : (s(bool,X213) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X211,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X211),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X213)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f12559,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X42,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X42),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | $spl166),
  inference(superposition,[],[f6057,f7057])).
fof(f7057,plain,(
  ( ! [X37] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X37,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X37),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f5374,f113])).
fof(f5374,plain,(
  ( ! [X4] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X4,'type.bvlSem.v'))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X4),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f3489,f539])).
fof(f539,plain,(
  ( ! [X6] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X6),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X6))))) )),
  inference(forward_demodulation,[],[f512,f97])).
fof(f512,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X6)))) | ~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X6),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))))) )),
  inference(resolution,[],[f88,f94])).
fof(f94,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f3489,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(X54,sK7(X55,X56,X54)),s(fun(X54,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X54),X55)))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X54,bool),X56),s('type.list.list'(X54),X55)))) )),
  inference(resolution,[],[f95,f113])).
fof(f9715,plain,(
  ( ! [X212,X214] : (s(bool,X214) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X212))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X212),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,'T_0') = s(bool,X214)) ) | $spl276),
  inference(cnf_transformation,[],[f9715_D])).
fof(f9715_D,plain,(
  ( ! [X212,X214] : (s(bool,X214) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X212))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X212),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,'T_0') = s(bool,X214)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f12561,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X47) | s(bool,X47) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X46))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X46),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | $spl166),
  inference(superposition,[],[f6057,f6695])).
fof(f6695,plain,(
  ( ! [X38] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X38))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X38),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f1794,f113])).
fof(f1794,plain,(
  ( ! [X33] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X33)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X33),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f109])).
fof(f8219,plain,(
  ( ! [X362,X365,X361,X363] : (s(bool,X365) != s(bool,'const.bool.IN_2'(s(X361,sK7(X362,X363,X361)),s(fun(X361,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X361),X362))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X361,bool),X363),s('type.list.list'(X361),X362))) | s(bool,'T_0') = s(bool,X365)) ) | $spl226),
  inference(cnf_transformation,[],[f8219_D])).
fof(f8219_D,plain,(
  ( ! [X362,X365,X361,X363] : (s(bool,X365) != s(bool,'const.bool.IN_2'(s(X361,sK7(X362,X363,X361)),s(fun(X361,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X361),X362))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X361,bool),X363),s('type.list.list'(X361),X362))) | s(bool,'T_0') = s(bool,X365)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f12563,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X52) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(X49,sK7(X50,X51,X49)),s(fun(X49,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X49),X50))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X49,bool),X51),s('type.list.list'(X49),X50)))) ) | $spl166),
  inference(superposition,[],[f6057,f5385])).
fof(f5385,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X59,sK7(X61,X60,X59)),s(fun(X59,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X59),X61))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X59,bool),X60),s('type.list.list'(X59),X61)))) )),
  inference(resolution,[],[f3489,f113])).
fof(f153,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) | $spl6),
  inference(cnf_transformation,[],[f153_D])).
fof(f153_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f12592,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))))) ) | ($spl14 | $spl166)),
  inference(subsumption_resolution,[],[f12565,f831])).
fof(f12565,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))))) ) | $spl166),
  inference(superposition,[],[f6057,f109])).
fof(f12594,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))))) ) | ($spl14 | $spl166)),
  inference(subsumption_resolution,[],[f12566,f831])).
fof(f12566,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))))) ) | $spl166),
  inference(superposition,[],[f6057,f111])).
fof(f13792,plain,(
  s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | $spl370),
  inference(cnf_transformation,[],[f13792_D])).
fof(f13792_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f12596,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | ($spl14 | $spl166)),
  inference(subsumption_resolution,[],[f12570,f831])).
fof(f12570,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X12) | s(bool,'T_0') = s(bool,X13) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl166),
  inference(superposition,[],[f6057,f116])).
fof(f13796,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) | $spl374),
  inference(cnf_transformation,[],[f13796_D])).
fof(f13796_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f13801,plain,(
  s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | $spl381),
  inference(cnf_transformation,[],[f13801_D])).
fof(f13801_D,plain,(
  s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) <=> ~$spl381),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl381])])).
fof(f12628,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl192),
  inference(subsumption_resolution,[],[f12597,f113])).
fof(f12597,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X0))) ) | $spl192),
  inference(superposition,[],[f6979,f109])).
fof(f12630,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X1)) ) | $spl192),
  inference(subsumption_resolution,[],[f12598,f113])).
fof(f12598,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X1) | p(s(bool,X1))) ) | $spl192),
  inference(superposition,[],[f6979,f111])).
fof(f12627,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X5)) ) | $spl192),
  inference(duplicate_literal_removal,[],[f12602])).
fof(f12602,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | $spl192),
  inference(superposition,[],[f6979,f116])).
fof(f13802,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X15),X16),s('type.bvlSem.v',X17)))) ) | $spl382),
  inference(cnf_transformation,[],[f13802_D])).
fof(f13802_D,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X15),X16),s('type.bvlSem.v',X17)))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f12633,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X15),X16),s('type.bvlSem.v',X17)))) ) | $spl192),
  inference(subsumption_resolution,[],[f12604,f4897])).
fof(f4897,plain,(
  ( ! [X255,X257,X256] : (s(bool,'T_0') != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X255),X256))),s('type.bvlSem.v',X257))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X255),X256),s('type.bvlSem.v',X257)))) )),
  inference(subsumption_resolution,[],[f4815,f140])).
fof(f4815,plain,(
  ( ! [X255,X257,X256] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X255),X256))),s('type.bvlSem.v',X257))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X255),X256),s('type.bvlSem.v',X257)))) )),
  inference(superposition,[],[f112,f3325])).
fof(f12604,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X15),X16))),s('type.bvlSem.v',X17))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X15),X16),s('type.bvlSem.v',X17)))) ) | $spl192),
  inference(superposition,[],[f6979,f3325])).
fof(f13799,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X21,bool),X22),s('type.list.list'(X21),X23)))) ) | $spl378),
  inference(cnf_transformation,[],[f13799_D])).
fof(f13799_D,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X21,bool),X22),s('type.list.list'(X21),X23)))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f12638,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X21,bool),X22),s('type.list.list'(X21),X23)))) ) | $spl192),
  inference(subsumption_resolution,[],[f12608,f3601])).
fof(f3601,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,happ(s(fun(X54,bool),X55),s(X54,sK7(X56,X55,X54)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X54,bool),X55),s('type.list.list'(X54),X56)))) )),
  inference(resolution,[],[f3264,f112])).
fof(f12608,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,happ(s(fun(X21,bool),X22),s(X21,sK7(X23,X22,X21)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X21,bool),X22),s('type.list.list'(X21),X23)))) ) | $spl192),
  inference(superposition,[],[f6979,f3599])).
fof(f13798,plain,(
  p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | $spl376),
  inference(cnf_transformation,[],[f13798_D])).
fof(f13798_D,plain,(
  p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f12640,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))))) ) | $spl192),
  inference(subsumption_resolution,[],[f12619,f116])).
fof(f12619,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X4) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))))) ) | $spl192),
  inference(superposition,[],[f6979,f109])).
fof(f12642,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))))) ) | $spl192),
  inference(subsumption_resolution,[],[f12620,f116])).
fof(f12620,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))))) ) | $spl192),
  inference(superposition,[],[f6979,f111])).
fof(f13795,plain,(
  s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | $spl372),
  inference(cnf_transformation,[],[f13795_D])).
fof(f13795_D,plain,(
  s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f12644,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl192),
  inference(subsumption_resolution,[],[f12624,f116])).
fof(f12624,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl192),
  inference(superposition,[],[f6979,f116])).
fof(f13221,plain,(
  ( ! [X366,X367,X369] : (s(bool,X369) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X366),s('type.list.list'(bool),X367))) | s(bool,'F_0') = s(bool,sK7(X367,X366,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X366),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X369)) ) | $spl350),
  inference(cnf_transformation,[],[f13221_D])).
fof(f13221_D,plain,(
  ( ! [X366,X367,X369] : (s(bool,X369) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X366),s('type.list.list'(bool),X367))) | s(bool,'F_0') = s(bool,sK7(X367,X366,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X366),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X369)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f13155,plain,(
  ( ! [X553,X551,X552] : (s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X553) | s(bool,X553) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X551),s('type.list.list'(bool),X552))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X551),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X552,X551,bool))) ) | $spl166),
  inference(superposition,[],[f6057,f5140])).
fof(f5140,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X9),s('type.list.list'(bool),X8))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X8,X9,bool))) )),
  inference(superposition,[],[f3599,f116])).
fof(f13546,plain,(
  ( ! [X368,X369,X371] : (s(bool,X371) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X368),s('type.list.list'(bool),X369))) | s(bool,'T_0') = s(bool,sK7(X369,X368,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X368),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X371)) ) | $spl356),
  inference(cnf_transformation,[],[f13546_D])).
fof(f13546_D,plain,(
  ( ! [X368,X369,X371] : (s(bool,X371) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X368),s('type.list.list'(bool),X369))) | s(bool,'T_0') = s(bool,sK7(X369,X368,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X368),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X371)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f13471,plain,(
  ( ! [X553,X555,X554] : (s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X555) | s(bool,X555) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X553),s('type.list.list'(bool),X554))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X553),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X554,X553,bool))) ) | $spl166),
  inference(superposition,[],[f6057,f5141])).
fof(f5141,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X10,X11,bool))) )),
  inference(superposition,[],[f3599,f116])).
fof(f145,plain,(
  p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) | $spl0),
  inference(cnf_transformation,[],[f145_D])).
fof(f145_D,plain,(
  p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2011,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X6),s('type.list.list'('type.bvlSem.v'),X7),s('type.bvlSem.state'(sK0),X8))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X8)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X8))))),s('type.list.list'('type.bvlSem.v'),X7))))) ) | $spl66),
  inference(cnf_transformation,[],[f2011_D])).
fof(f2011_D,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X6),s('type.list.list'('type.bvlSem.v'),X7),s('type.bvlSem.state'(sK0),X8))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X8)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X8))))),s('type.list.list'('type.bvlSem.v'),X7))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f632,plain,(
  ( ! [X4,X5,X3] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X3),s('type.list.list'('type.bvlSem.v'),X4),s('type.bvlSem.state'(sK0),X5))) | p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X5))))),s('type.list.list'('type.bvlSem.v'),X4)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X5))))) )),
  inference(superposition,[],[f92,f89])).
fof(f89,plain,(
  s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',sK3),s('type.bvlSem.state'(sK0),sK1)))))),
  inference(cnf_transformation,[],[f64])).
fof(f92,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1))))) | (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3))))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1)))))) | (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3))))))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1)))))) => (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3))))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : ((p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1)))))) => (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X1)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X7,X8,X9,X10,X12] : ((p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X6),X8)))) & p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X8))))),s('type.list.list'('type.bvlSem.v'),X12)))) & s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X10),s('type.list.list'('type.bvlSem.v'),X12),s('type.bvlSem.state'(X6),X8))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'const.pair.,_2'(s('type.bvlSem.v',X9),s('type.bvlSem.state'(X6),X7)))))) => (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X6),X7)))) & p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X7))))),s('type.bvlSem.v',X9))))))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bvl_to_bviProof.bvl_to_bviProof_29_prove_ltb')).
fof(f2010,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2)))) ) | $spl64),
  inference(cnf_transformation,[],[f2010_D])).
fof(f2010_D,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1874,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) | p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))))) )),
  inference(superposition,[],[f535,f89])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X1),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',X2),s('type.bvlSem.state'(sK0),X0))))) | p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X0))))) )),
  inference(subsumption_resolution,[],[f534,f141])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X0)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X1),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',X2),s('type.bvlSem.state'(sK0),X0)))))) )),
  inference(forward_demodulation,[],[f510,f166])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X0)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X1),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',X2),s('type.bvlSem.state'(sK0),X0))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2))))) )),
  inference(resolution,[],[f88,f92])).
fof(f4086,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),X1),s('type.bvlSem.state'(sK0),X2))) | p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X2))))),s('type.list.list'('type.bvlSem.v'),X1)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X2))))) )),
  inference(superposition,[],[f92,f89])).
fof(f11065,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) | p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))))) )),
  inference(superposition,[],[f535,f89])).
fof(f13791,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),X1),s('type.bvlSem.state'(sK0),sK2))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),X1))))) ) | $spl368),
  inference(cnf_transformation,[],[f13791_D])).
fof(f13791_D,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),X1),s('type.bvlSem.state'(sK0),sK2))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),X1))))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f13607,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),X1),s('type.bvlSem.state'(sK0),sK2))) | p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),X1))))) )),
  inference(superposition,[],[f154,f89])).
fof(f154,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X1),s('type.list.list'('type.bvlSem.v'),X2),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(sK0),X0))))) | p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X0)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),X2))))) )),
  inference(resolution,[],[f87,f92])).
fof(f13785,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f13784,f140])).
fof(f13784,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f13783,f2734])).
fof(f13783,plain,(
  p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f13771,f141])).
fof(f13771,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | ($spl3 | $spl6)),
  inference(superposition,[],[f13702,f6765])).
fof(f6765,plain,(
  s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f6757,f141])).
fof(f6757,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | ($spl3 | $spl6)),
  inference(superposition,[],[f2777,f2671])).
fof(f2671,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | $spl3),
  inference(resolution,[],[f2205,f113])).
fof(f2205,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | $spl3),
  inference(resolution,[],[f148,f95])).
fof(f2777,plain,(
  ( ! [X37] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X37),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X37)))) ) | $spl6),
  inference(resolution,[],[f2269,f113])).
fof(f2269,plain,(
  ( ! [X6] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X6)))) | ~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X6),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))))) ) | $spl6),
  inference(forward_demodulation,[],[f2237,f97])).
fof(f2237,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X6)))) | ~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X6),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))))) ) | $spl6),
  inference(resolution,[],[f153,f94])).
fof(f13702,plain,(
  ( ! [X0] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X0)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X0))))) )),
  inference(resolution,[],[f13701,f118])).
fof(f118,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X4),s('type.bvlSem.v',X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X2))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X1),X4)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X2),s('type.bvlSem.v',X3))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X2),s('type.bvlSem.v',X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X2))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X1),X4)))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X4),s('type.bvlSem.v',X3)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X2),s('type.bvlSem.v',X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X2))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X1),X4))))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X4),s('type.bvlSem.v',X3)))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X2),s('type.bvlSem.v',X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X2))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X1),X4))))))) => p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X4),s('type.bvlSem.v',X3)))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : ((p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X2),s('type.bvlSem.v',X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X2))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X1),X4))))))) => p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X4),s('type.bvlSem.v',X3)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X20] : ! [X21,X13,X22] : ((p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X21),s('type.bvlSem.v',X13)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X6),X21))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X20),X22))))))) => p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X20),X22),s('type.bvlSem.v',X13)))))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bvl_to_bviProof.bvl_to_bviProof_2|49|_prove_ltb')).
fof(f13701,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))))))),
  inference(equality_resolution,[],[f631])).
fof(f631,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),X1),s('type.bvlSem.state'(sK0),X2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))))))) )),
  inference(superposition,[],[f117,f89])).
fof(f117,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1)))))))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1)))))))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1)))))))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1)))))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7,X8,X9,X10,X12] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X10),s('type.list.list'('type.bvlSem.v'),X12),s('type.bvlSem.state'(X6),X8))) = s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X6)),'const.pair.,_2'(s('type.bvlSem.v',X9),s('type.bvlSem.state'(X6),X7))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X8))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X6),X7)))))))))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bvlProps.do_app_refs_SUBSET')).
fof(f13782,plain,(
  ( ! [X12,X10] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X10)))) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X10)))) ) | $spl366),
  inference(cnf_transformation,[],[f13782_D])).
fof(f13782_D,plain,(
  ( ! [X12,X10] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X10)))) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X10)))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f13764,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X48)))) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X48)))) )),
  inference(resolution,[],[f13702,f142])).
fof(f13757,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X38)))) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X38))) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f13702,f105])).
fof(f13755,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X30)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X30)))) )),
  inference(resolution,[],[f13702,f104])).
fof(f13753,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X22)))) )),
  inference(resolution,[],[f13702,f103])).
fof(f13752,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X18)))) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X18))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f13702,f101])).
fof(f13750,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X10)))) )),
  inference(resolution,[],[f13702,f99])).
fof(f13738,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) ) | $spl364),
  inference(cnf_transformation,[],[f13738_D])).
fof(f13738_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f13723,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) )),
  inference(resolution,[],[f13701,f142])).
fof(f13716,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))))) | s(bool,X33) != s(bool,X34)) )),
  inference(resolution,[],[f13701,f105])).
fof(f13713,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) )),
  inference(resolution,[],[f13701,f104])).
fof(f13712,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) )),
  inference(resolution,[],[f13701,f103])).
fof(f13711,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))))) | s(bool,X18) != s(bool,X19)) )),
  inference(resolution,[],[f13701,f101])).
fof(f13709,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) )),
  inference(resolution,[],[f13701,f99])).
fof(f13690,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X11)))) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X14),X15))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X12),X13))))) | ~p(s(bool,X17)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X14),X15)))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X11))))) ) | $spl362),
  inference(cnf_transformation,[],[f13690_D])).
fof(f13690_D,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X11)))) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X14),X15))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X12),X13))))) | ~p(s(bool,X17)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X14),X15)))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X11))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f13677,plain,(
  ( ! [X134,X132,X138,X136,X135,X133,X131,X139,X137] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X131),X132),s('type.bvlSem.v',X133)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X134),X135),s('type.bvlSem.v',X133)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X131),X132))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X136),X137)))))) | p(s(bool,X138)) | ~p(s(bool,X139)) | s(bool,X139) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X136),X137))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X134),X135)))))) )),
  inference(resolution,[],[f3836,f142])).
fof(f3836,plain,(
  ( ! [X6,X4,X2,X8,X7,X5,X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X4),X5))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X2),X3),s('type.bvlSem.v',X6)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X6)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X2),X3))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X4),X5))))))) )),
  inference(resolution,[],[f118,f118])).
fof(f13670,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X79,X86,X84,X82] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X79),X80),s('type.bvlSem.v',X81)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X82),X83),s('type.bvlSem.v',X81)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X79),X80))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X84),X85)))))) | ~p(s(bool,X86)) | ~p(s(bool,X87)) | s(bool,X87) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X84),X85))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X82),X83))))) | s(bool,X86) != s(bool,X88)) )),
  inference(resolution,[],[f3836,f105])).
fof(f13668,plain,(
  ( ! [X61,X59,X68,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X59),X60),s('type.bvlSem.v',X61)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X62),X63),s('type.bvlSem.v',X61)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X59),X60))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X64),X65)))))) | p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X64),X65))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X62),X63)))))) )),
  inference(resolution,[],[f3836,f104])).
fof(f13666,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X46,X44,X42,X40] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X39),X40),s('type.bvlSem.v',X41)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X42),X43),s('type.bvlSem.v',X41)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X39),X40))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X44),X45)))))) | p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X44),X45))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X42),X43)))))) )),
  inference(resolution,[],[f3836,f103])).
fof(f13665,plain,(
  ( ! [X30,X37,X35,X33,X31,X29,X38,X36,X34,X32] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X29),X30),s('type.bvlSem.v',X31)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X32),X33),s('type.bvlSem.v',X31)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X29),X30))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X34),X35)))))) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X37) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X34),X35))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X32),X33))))) | s(bool,X36) != s(bool,X38)) )),
  inference(resolution,[],[f3836,f101])).
fof(f13663,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X18,X16] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X11)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X11)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X14),X15)))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X14),X15))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X12),X13)))))) )),
  inference(resolution,[],[f3836,f99])).
fof(f13645,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))))) ) | $spl360),
  inference(cnf_transformation,[],[f13645_D])).
fof(f13645_D,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f13626,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X56)))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,happ(s(fun(bool,bool),X55),s(bool,'T_0')))) )),
  inference(resolution,[],[f9250,f142])).
fof(f9250,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0))))))) )),
  inference(subsumption_resolution,[],[f9197,f112])).
fof(f9197,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0)))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))))) )),
  inference(resolution,[],[f5396,f3304])).
fof(f3304,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK7(X4,X5,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | ~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0'))))) )),
  inference(superposition,[],[f96,f113])).
fof(f5396,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0)))))) | p(s(bool,sK7(X0,X1,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) )),
  inference(superposition,[],[f3489,f109])).
fof(f13619,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X39)))))) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(bool,bool),X38),s(bool,'T_0'))) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f9250,f105])).
fof(f13617,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X28),s('type.list.list'(bool),X29)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X29)))))) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,happ(s(fun(bool,bool),X28),s(bool,'T_0')))) )),
  inference(resolution,[],[f9250,f104])).
fof(f13615,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X19)))))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0')))) )),
  inference(resolution,[],[f9250,f103])).
fof(f13614,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X14)))))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0'))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f9250,f101])).
fof(f13612,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4)))))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) )),
  inference(resolution,[],[f9250,f99])).
fof(f1751,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,X40)) ) | $spl54),
  inference(cnf_transformation,[],[f1751_D])).
fof(f1751_D,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,X40)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f13555,plain,(
  ( ! [X453,X451,X454,X450] : (s(bool,X453) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X450),s('type.list.list'(bool),X451))) | s(bool,'T_0') = s(bool,sK7(X451,X450,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X450),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X454) | s(bool,X453) != s(bool,X454)) ) | $spl358),
  inference(cnf_transformation,[],[f13555_D])).
fof(f13555_D,plain,(
  ( ! [X453,X451,X454,X450] : (s(bool,X453) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X450),s('type.list.list'(bool),X451))) | s(bool,'T_0') = s(bool,sK7(X451,X450,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X450),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X454) | s(bool,X453) != s(bool,X454)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f13554,plain,(
  ( ! [X453,X451,X452,X454,X450] : (s(bool,X453) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X450),s('type.list.list'(bool),X451))) | s(bool,X453) != s(bool,X454) | s(bool,'F_0') = s(bool,X452) | s(bool,'T_0') = s(bool,X454) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X450),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X451,X450,bool))) ) | ($spl14 | $spl40)),
  inference(subsumption_resolution,[],[f13445,f116])).
fof(f13445,plain,(
  ( ! [X453,X451,X452,X454,X450] : (s(bool,'T_0') != s(bool,X452) | s(bool,X453) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X450),s('type.list.list'(bool),X451))) | s(bool,X453) != s(bool,X454) | s(bool,'F_0') = s(bool,X452) | s(bool,'T_0') = s(bool,X454) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X450),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X451,X450,bool))) ) | ($spl14 | $spl40)),
  inference(superposition,[],[f1953,f5141])).
fof(f1953,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,X66) != s(bool,X67) | s(bool,X64) != s(bool,X66) | s(bool,X64) != s(bool,X65) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X65)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1276,f113])).
fof(f1276,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X78)) | s(bool,X76) != s(bool,X78) | s(bool,X76) != s(bool,X77) | s(bool,X77) != s(bool,X79) | s(bool,'F_0') = s(bool,X79)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f807,f770])).
fof(f13545,plain,(
  ( ! [X368,X370,X369,X371] : (s(bool,'F_0') = s(bool,X370) | s(bool,X371) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X368),s('type.list.list'(bool),X369))) | s(bool,'T_0') = s(bool,X371) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X368),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X369,X368,bool))) ) | $spl14),
  inference(subsumption_resolution,[],[f13424,f116])).
fof(f13424,plain,(
  ( ! [X368,X370,X369,X371] : (s(bool,'T_0') != s(bool,X370) | s(bool,'F_0') = s(bool,X370) | s(bool,X371) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X368),s('type.list.list'(bool),X369))) | s(bool,'T_0') = s(bool,X371) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X368),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X369,X368,bool))) ) | $spl14),
  inference(superposition,[],[f1032,f5141])).
fof(f1032,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f829,f116])).
fof(f829,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,'F_0') = s(bool,X41) | s(bool,X40) != s(bool,X41)) ) | $spl14),
  inference(resolution,[],[f770,f108])).
fof(f8688,plain,(
  ( ! [X24,X27,X25] : (s(bool,'T_0') = s(bool,sK7(X24,X25,bool)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24))) | p(s(bool,X27)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0')))) ) | $spl248),
  inference(cnf_transformation,[],[f8688_D])).
fof(f8688_D,plain,(
  ( ! [X24,X27,X25] : (s(bool,'T_0') = s(bool,sK7(X24,X25,bool)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24))) | p(s(bool,X27)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0')))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f13541,plain,(
  ( ! [X341,X339,X340,X338] : (p(s(bool,X340)) | p(s(bool,X341)) | s(bool,X341) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X338),s('type.list.list'(bool),X339))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X338),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X339,X338,bool))) )),
  inference(subsumption_resolution,[],[f13415,f141])).
fof(f13415,plain,(
  ( ! [X341,X339,X340,X338] : (~p(s(bool,'T_0')) | p(s(bool,X340)) | p(s(bool,X341)) | s(bool,X341) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X338),s('type.list.list'(bool),X339))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X338),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X339,X338,bool))) )),
  inference(superposition,[],[f142,f5141])).
fof(f13537,plain,(
  ( ! [X277,X273,X275,X276,X274] : (p(s(bool,X275)) | ~p(s(bool,X276)) | s(bool,X275) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X273),s('type.list.list'(bool),X274))) | s(bool,X276) != s(bool,X277) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X273),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X274,X273,bool))) )),
  inference(subsumption_resolution,[],[f13391,f136])).
fof(f136,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ((p(s(bool,X7)) & p(s(bool,'T_0'))) <=> p(s(bool,X7)))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.AND_CLAUSES|split|1')).
fof(f13391,plain,(
  ( ! [X277,X273,X275,X276,X274] : (~p(s(bool,'T_0')) | p(s(bool,X275)) | ~p(s(bool,X276)) | s(bool,X275) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X273),s('type.list.list'(bool),X274))) | s(bool,X276) != s(bool,X277) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X273),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X274,X273,bool))) )),
  inference(superposition,[],[f105,f5141])).
fof(f13535,plain,(
  ( ! [X237,X235,X233,X236,X234] : (p(s(bool,X235)) | ~p(s(bool,X236)) | p(s(bool,X237)) | s(bool,X237) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X233),s('type.list.list'(bool),X234))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X233),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X234,X233,bool))) )),
  inference(subsumption_resolution,[],[f13383,f136])).
fof(f13383,plain,(
  ( ! [X237,X235,X233,X236,X234] : (~p(s(bool,'T_0')) | p(s(bool,X235)) | ~p(s(bool,X236)) | p(s(bool,X237)) | s(bool,X237) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X233),s('type.list.list'(bool),X234))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X233),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X234,X233,bool))) )),
  inference(superposition,[],[f104,f5141])).
fof(f13533,plain,(
  ( ! [X212,X210,X208,X211,X209] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X210) != s(bool,X212) | s(bool,X211) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X208),s('type.list.list'(bool),X209))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X208),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X209,X208,bool))) )),
  inference(subsumption_resolution,[],[f13378,f141])).
fof(f13378,plain,(
  ( ! [X212,X210,X208,X211,X209] : (~p(s(bool,'T_0')) | p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X210) != s(bool,X212) | s(bool,X211) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X208),s('type.list.list'(bool),X209))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X208),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X209,X208,bool))) )),
  inference(superposition,[],[f103,f5141])).
fof(f13530,plain,(
  ( ! [X177,X175,X173,X176,X174] : (p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X175) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X173),s('type.list.list'(bool),X174))) | s(bool,X176) != s(bool,X177) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X173),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X174,X173,bool))) )),
  inference(subsumption_resolution,[],[f13371,f136])).
fof(f13371,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,'T_0')) | p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X175) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X173),s('type.list.list'(bool),X174))) | s(bool,X176) != s(bool,X177) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X173),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X174,X173,bool))) )),
  inference(superposition,[],[f101,f5141])).
fof(f13528,plain,(
  ( ! [X142,X140,X138,X141,X139] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X140) != s(bool,X142) | s(bool,X141) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X138),s('type.list.list'(bool),X139))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X138),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X139,X138,bool))) )),
  inference(subsumption_resolution,[],[f13364,f141])).
fof(f13364,plain,(
  ( ! [X142,X140,X138,X141,X139] : (~p(s(bool,'T_0')) | p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X140) != s(bool,X142) | s(bool,X141) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X138),s('type.list.list'(bool),X139))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X138),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X139,X138,bool))) )),
  inference(superposition,[],[f99,f5141])).
fof(f13525,plain,(
  ( ! [X97,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),X97),s(bool,X100)))) | s(bool,'T_0') = s(bool,sK7(X98,X97,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X97),s(bool,'F_0'))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X100),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X98))))))) ) | $spl354),
  inference(cnf_transformation,[],[f13525_D])).
fof(f13525_D,plain,(
  ( ! [X97,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),X97),s(bool,X100)))) | s(bool,'T_0') = s(bool,sK7(X98,X97,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X97),s(bool,'F_0'))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X100),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X98))))))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f415,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,X18)) ) | $spl32),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,X18)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f13524,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'T_0') != s(bool,X99) | p(s(bool,happ(s(fun(bool,bool),X97),s(bool,X100)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X100),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X98)))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X97),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X98,X97,bool))) ) | $spl100),
  inference(subsumption_resolution,[],[f13354,f112])).
fof(f13354,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'T_0') != s(bool,X99) | p(s(bool,happ(s(fun(bool,bool),X97),s(bool,X100)))) | ~p(s(bool,X99)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X100),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X98)))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X97),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X98,X97,bool))) ) | $spl100),
  inference(superposition,[],[f3755,f5141])).
fof(f13312,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f13288,f415])).
fof(f13288,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X30,bool),X32),s('type.list.list'(X30),X31)))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f13286])).
fof(f13286,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X30,bool),X32),s('type.list.list'(X30),X31)))) ) | $spl32),
  inference(superposition,[],[f415,f5385])).
fof(f13311,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f13292,f415])).
fof(f13292,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X27),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f13282])).
fof(f13282,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X27),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | $spl32),
  inference(superposition,[],[f415,f7057])).
fof(f13310,plain,(
  $false | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f13293,f415])).
fof(f13293,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X26),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | ($spl6 | $spl32)),
  inference(trivial_inequality_removal,[],[f13281])).
fof(f13281,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X26),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | ($spl6 | $spl32)),
  inference(superposition,[],[f415,f6736])).
fof(f6736,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X0,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X0),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f2777,f3489])).
fof(f13303,plain,(
  $false | $spl32),
  inference(trivial_inequality_removal,[],[f13265])).
fof(f13265,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl32),
  inference(superposition,[],[f415,f166])).
fof(f13302,plain,(
  $false | ($spl0 | $spl32)),
  inference(trivial_inequality_removal,[],[f13266])).
fof(f13266,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl32)),
  inference(superposition,[],[f415,f227])).
fof(f227,plain,(
  s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | $spl0),
  inference(resolution,[],[f145,f113])).
fof(f13301,plain,(
  $false | $spl32),
  inference(trivial_inequality_removal,[],[f13267])).
fof(f13267,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl32),
  inference(superposition,[],[f415,f523])).
fof(f13300,plain,(
  $false | ($spl6 | $spl32)),
  inference(trivial_inequality_removal,[],[f13268])).
fof(f13268,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl6 | $spl32)),
  inference(superposition,[],[f415,f2248])).
fof(f2248,plain,(
  s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | $spl6),
  inference(resolution,[],[f153,f113])).
fof(f13297,plain,(
  $false | ($spl4 | $spl14 | $spl32)),
  inference(trivial_inequality_removal,[],[f13274])).
fof(f13274,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl14 | $spl32)),
  inference(superposition,[],[f415,f6136])).
fof(f6136,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | ($spl4 | $spl14)),
  inference(trivial_inequality_removal,[],[f6125])).
fof(f6125,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | ($spl4 | $spl14)),
  inference(superposition,[],[f1180,f2023])).
fof(f2023,plain,(
  s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | $spl4),
  inference(resolution,[],[f149,f113])).
fof(f1180,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X14))) | s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X12),X13))),s('type.bvlSem.v',X14)))) ) | $spl14),
  inference(forward_demodulation,[],[f1163,f166])).
fof(f1163,plain,(
  ( ! [X14,X12,X13] : (s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2))) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X14))) | s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X12),X13))),s('type.bvlSem.v',X14)))) ) | $spl14),
  inference(superposition,[],[f1153,f97])).
fof(f1153,plain,(
  ( ! [X24] : (s(bool,X24) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2))) | s(bool,'T_0') = s(bool,X24)) ) | $spl14),
  inference(subsumption_resolution,[],[f1145,f141])).
fof(f1145,plain,(
  ( ! [X24] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X24) | s(bool,X24) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) ) | $spl14),
  inference(superposition,[],[f1102,f166])).
fof(f13296,plain,(
  $false | ($spl3 | $spl6 | $spl14 | $spl32)),
  inference(trivial_inequality_removal,[],[f13276])).
fof(f13276,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl6 | $spl14 | $spl32)),
  inference(superposition,[],[f415,f6930])).
fof(f6930,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | ($spl3 | $spl6 | $spl14)),
  inference(trivial_inequality_removal,[],[f6819])).
fof(f6819,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | ($spl3 | $spl6 | $spl14)),
  inference(superposition,[],[f1180,f6765])).
fof(f13295,plain,(
  $false | ($spl4 | $spl32)),
  inference(trivial_inequality_removal,[],[f13278])).
fof(f13278,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl32)),
  inference(superposition,[],[f415,f2023])).
fof(f13294,plain,(
  $false | ($spl3 | $spl6 | $spl32)),
  inference(trivial_inequality_removal,[],[f13280])).
fof(f13280,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl6 | $spl32)),
  inference(superposition,[],[f415,f6765])).
fof(f13289,plain,(
  $false | ($spl3 | $spl32)),
  inference(trivial_inequality_removal,[],[f13285])).
fof(f13285,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl32)),
  inference(superposition,[],[f415,f2671])).
fof(f13287,plain,(
  $false | $spl32),
  inference(equality_resolution,[],[f415])).
fof(f13230,plain,(
  ( ! [X449,X451,X452,X448] : (s(bool,X451) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X448),s('type.list.list'(bool),X449))) | s(bool,'F_0') = s(bool,sK7(X449,X448,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X448),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X452) | s(bool,X451) != s(bool,X452)) ) | $spl352),
  inference(cnf_transformation,[],[f13230_D])).
fof(f13230_D,plain,(
  ( ! [X449,X451,X452,X448] : (s(bool,X451) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X448),s('type.list.list'(bool),X449))) | s(bool,'F_0') = s(bool,sK7(X449,X448,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X448),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X452) | s(bool,X451) != s(bool,X452)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f13229,plain,(
  ( ! [X449,X451,X452,X448,X450] : (s(bool,X451) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X448),s('type.list.list'(bool),X449))) | s(bool,X451) != s(bool,X452) | s(bool,'F_0') = s(bool,X450) | s(bool,'T_0') = s(bool,X452) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X448),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X449,X448,bool))) ) | ($spl14 | $spl40)),
  inference(subsumption_resolution,[],[f13129,f116])).
fof(f13129,plain,(
  ( ! [X449,X451,X452,X448,X450] : (s(bool,'T_0') != s(bool,X450) | s(bool,X451) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X448),s('type.list.list'(bool),X449))) | s(bool,X451) != s(bool,X452) | s(bool,'F_0') = s(bool,X450) | s(bool,'T_0') = s(bool,X452) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X448),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X449,X448,bool))) ) | ($spl14 | $spl40)),
  inference(superposition,[],[f1953,f5140])).
fof(f13220,plain,(
  ( ! [X366,X368,X367,X369] : (s(bool,'F_0') = s(bool,X368) | s(bool,X369) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X366),s('type.list.list'(bool),X367))) | s(bool,'T_0') = s(bool,X369) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X366),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X367,X366,bool))) ) | $spl14),
  inference(subsumption_resolution,[],[f13108,f116])).
fof(f13108,plain,(
  ( ! [X366,X368,X367,X369] : (s(bool,'T_0') != s(bool,X368) | s(bool,'F_0') = s(bool,X368) | s(bool,X369) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X366),s('type.list.list'(bool),X367))) | s(bool,'T_0') = s(bool,X369) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X366),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X367,X366,bool))) ) | $spl14),
  inference(superposition,[],[f1032,f5140])).
fof(f8648,plain,(
  ( ! [X24,X27,X25] : (s(bool,'F_0') = s(bool,sK7(X24,X25,bool)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24))) | p(s(bool,X27)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0')))) ) | $spl246),
  inference(cnf_transformation,[],[f8648_D])).
fof(f8648_D,plain,(
  ( ! [X24,X27,X25] : (s(bool,'F_0') = s(bool,sK7(X24,X25,bool)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24))) | p(s(bool,X27)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0')))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f13216,plain,(
  ( ! [X337,X339,X336,X338] : (p(s(bool,X338)) | p(s(bool,X339)) | s(bool,X339) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X336),s('type.list.list'(bool),X337))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X336),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X337,X336,bool))) )),
  inference(subsumption_resolution,[],[f13099,f141])).
fof(f13099,plain,(
  ( ! [X337,X339,X336,X338] : (~p(s(bool,'T_0')) | p(s(bool,X338)) | p(s(bool,X339)) | s(bool,X339) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X336),s('type.list.list'(bool),X337))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X336),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X337,X336,bool))) )),
  inference(superposition,[],[f142,f5140])).
fof(f13212,plain,(
  ( ! [X271,X273,X275,X272,X274] : (p(s(bool,X273)) | ~p(s(bool,X274)) | s(bool,X273) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X271),s('type.list.list'(bool),X272))) | s(bool,X274) != s(bool,X275) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X271),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X272,X271,bool))) )),
  inference(subsumption_resolution,[],[f13075,f136])).
fof(f13075,plain,(
  ( ! [X271,X273,X275,X272,X274] : (~p(s(bool,'T_0')) | p(s(bool,X273)) | ~p(s(bool,X274)) | s(bool,X273) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X271),s('type.list.list'(bool),X272))) | s(bool,X274) != s(bool,X275) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X271),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X272,X271,bool))) )),
  inference(superposition,[],[f105,f5140])).
fof(f13210,plain,(
  ( ! [X231,X235,X233,X234,X232] : (p(s(bool,X233)) | ~p(s(bool,X234)) | p(s(bool,X235)) | s(bool,X235) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X231),s('type.list.list'(bool),X232))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X231),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X232,X231,bool))) )),
  inference(subsumption_resolution,[],[f13067,f136])).
fof(f13067,plain,(
  ( ! [X231,X235,X233,X234,X232] : (~p(s(bool,'T_0')) | p(s(bool,X233)) | ~p(s(bool,X234)) | p(s(bool,X235)) | s(bool,X235) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X231),s('type.list.list'(bool),X232))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X231),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X232,X231,bool))) )),
  inference(superposition,[],[f104,f5140])).
fof(f13208,plain,(
  ( ! [X210,X208,X206,X209,X207] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X208) != s(bool,X210) | s(bool,X209) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X206),s('type.list.list'(bool),X207))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X206),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X207,X206,bool))) )),
  inference(subsumption_resolution,[],[f13062,f141])).
fof(f13062,plain,(
  ( ! [X210,X208,X206,X209,X207] : (~p(s(bool,'T_0')) | p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X208) != s(bool,X210) | s(bool,X209) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X206),s('type.list.list'(bool),X207))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X206),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X207,X206,bool))) )),
  inference(superposition,[],[f103,f5140])).
fof(f13205,plain,(
  ( ! [X175,X173,X171,X174,X172] : (p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X171),s('type.list.list'(bool),X172))) | s(bool,X174) != s(bool,X175) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X171),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X172,X171,bool))) )),
  inference(subsumption_resolution,[],[f13055,f136])).
fof(f13055,plain,(
  ( ! [X175,X173,X171,X174,X172] : (~p(s(bool,'T_0')) | p(s(bool,X173)) | ~p(s(bool,X174)) | s(bool,X173) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X171),s('type.list.list'(bool),X172))) | s(bool,X174) != s(bool,X175) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X171),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X172,X171,bool))) )),
  inference(superposition,[],[f101,f5140])).
fof(f13203,plain,(
  ( ! [X140,X138,X136,X139,X137] : (p(s(bool,X138)) | p(s(bool,X139)) | s(bool,X138) != s(bool,X140) | s(bool,X139) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X136),s('type.list.list'(bool),X137))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X136),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X137,X136,bool))) )),
  inference(subsumption_resolution,[],[f13048,f141])).
fof(f13048,plain,(
  ( ! [X140,X138,X136,X139,X137] : (~p(s(bool,'T_0')) | p(s(bool,X138)) | p(s(bool,X139)) | s(bool,X138) != s(bool,X140) | s(bool,X139) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X136),s('type.list.list'(bool),X137))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X136),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X137,X136,bool))) )),
  inference(superposition,[],[f99,f5140])).
fof(f13200,plain,(
  ( ! [X95,X98,X96] : (p(s(bool,happ(s(fun(bool,bool),X95),s(bool,X98)))) | s(bool,'F_0') = s(bool,sK7(X96,X95,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0'))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X98),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X96))))))) ) | $spl348),
  inference(cnf_transformation,[],[f13200_D])).
fof(f13200_D,plain,(
  ( ! [X95,X98,X96] : (p(s(bool,happ(s(fun(bool,bool),X95),s(bool,X98)))) | s(bool,'F_0') = s(bool,sK7(X96,X95,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0'))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X98),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X96))))))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f13199,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,X97) | p(s(bool,happ(s(fun(bool,bool),X95),s(bool,X98)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X98),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X96)))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X96,X95,bool))) ) | $spl100),
  inference(subsumption_resolution,[],[f13038,f112])).
fof(f13038,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,X97) | p(s(bool,happ(s(fun(bool,bool),X95),s(bool,X98)))) | ~p(s(bool,X97)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X98),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X96)))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X96,X95,bool))) ) | $spl100),
  inference(superposition,[],[f3755,f5140])).
fof(f12976,plain,(
  ( ! [X5,X3] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X5)) | s(bool,X5) != s(bool,happ(s(fun('type.bvlSem.v',bool),X3),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl346),
  inference(cnf_transformation,[],[f12976_D])).
fof(f12976_D,plain,(
  ( ! [X5,X3] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X5)) | s(bool,X5) != s(bool,happ(s(fun('type.bvlSem.v',bool),X3),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f12955,plain,(
  ( ! [X43,X41,X42] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X41),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,happ(s(fun('type.bvlSem.v',bool),X41),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f7381,f142])).
fof(f7381,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X0),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X0),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(subsumption_resolution,[],[f7358,f141])).
fof(f7358,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.bvlSem.v',bool),X0),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X0),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(superposition,[],[f3677,f2671])).
fof(f3677,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(X64,X66),s(fun(X64,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X64),X67)))))) | p(s(bool,happ(s(fun(X64,bool),X65),s(X64,X66)))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X64,bool),X65),s('type.list.list'(X64),X67)))) )),
  inference(resolution,[],[f94,f109])).
fof(f12948,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X31),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,happ(s(fun('type.bvlSem.v',bool),X31),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X33) != s(bool,X34)) ) | $spl3),
  inference(resolution,[],[f7381,f105])).
fof(f12945,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X19),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun('type.bvlSem.v',bool),X19),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f7381,f104])).
fof(f12944,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X15),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,happ(s(fun('type.bvlSem.v',bool),X15),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f7381,f103])).
fof(f12943,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X11),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,happ(s(fun('type.bvlSem.v',bool),X11),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f7381,f101])).
fof(f12941,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,happ(s(fun('type.bvlSem.v',bool),X3),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f7381,f99])).
fof(f12848,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X40] : (s(bool,X36) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X37),X38))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X39),X40))))) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X37),X38),s('type.bvlSem.v',X41))) | ~p(s(bool,X43)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X39),X40),s('type.bvlSem.v',X41)))) | ~p(s(bool,X36))) ) | $spl344),
  inference(cnf_transformation,[],[f12848_D])).
fof(f12848_D,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X40] : (s(bool,X36) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X37),X38))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X39),X40))))) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X37),X38),s('type.bvlSem.v',X41))) | ~p(s(bool,X43)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X39),X40),s('type.bvlSem.v',X41)))) | ~p(s(bool,X36))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f12921,plain,(
  ( ! [X123,X121,X118,X116,X122,X120,X119,X117] : (s(bool,X116) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X117),X118),s('type.bvlSem.v',X119))) | ~p(s(bool,X116)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X120),X121),s('type.bvlSem.v',X119)))) | p(s(bool,X122)) | ~p(s(bool,X123)) | s(bool,X123) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X117),X118))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X120),X121)))))) ) | $spl108),
  inference(resolution,[],[f3918,f142])).
fof(f12914,plain,(
  ( ! [X78,X76,X74,X72,X71,X79,X77,X75,X73] : (s(bool,X71) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X72),X73),s('type.bvlSem.v',X74))) | ~p(s(bool,X71)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X75),X76),s('type.bvlSem.v',X74)))) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X72),X73))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X75),X76))))) | s(bool,X77) != s(bool,X79)) ) | $spl108),
  inference(resolution,[],[f3918,f105])).
fof(f12912,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (s(bool,X53) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X54),X55),s('type.bvlSem.v',X56))) | ~p(s(bool,X53)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X57),X58),s('type.bvlSem.v',X56)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X54),X55))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X57),X58)))))) ) | $spl108),
  inference(resolution,[],[f3918,f104])).
fof(f12910,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X42,X40] : (s(bool,X35) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X36),X37),s('type.bvlSem.v',X38))) | ~p(s(bool,X35)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X39),X40),s('type.bvlSem.v',X38)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X36),X37))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X39),X40)))))) ) | $spl108),
  inference(resolution,[],[f3918,f103])).
fof(f12909,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X34,X32] : (s(bool,X26) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X27),X28),s('type.bvlSem.v',X29))) | ~p(s(bool,X26)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X30),X31),s('type.bvlSem.v',X29)))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X27),X28))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X30),X31))))) | s(bool,X32) != s(bool,X34)) ) | $spl108),
  inference(resolution,[],[f3918,f101])).
fof(f12907,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (s(bool,X8) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X11))) | ~p(s(bool,X8)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X11)))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X12),X13)))))) ) | $spl108),
  inference(resolution,[],[f3918,f99])).
fof(f7927,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2))),s('type.list.list'('type.bvlSem.v'),X3))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X2),s('type.bvlSem.v',sK7(X3,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2)),'type.bvlSem.v'))))) ) | $spl222),
  inference(cnf_transformation,[],[f7927_D])).
fof(f7927_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2))),s('type.list.list'('type.bvlSem.v'),X3))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X2),s('type.bvlSem.v',sK7(X3,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2)),'type.bvlSem.v'))))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f12902,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2))),s('type.list.list'('type.bvlSem.v'),X3))) | s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X2),s('type.bvlSem.v',sK7(X3,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2)),'type.bvlSem.v'))))) ) | $spl14),
  inference(subsumption_resolution,[],[f12888,f116])).
fof(f12888,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2))),s('type.list.list'('type.bvlSem.v'),X3))) | s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X2),s('type.bvlSem.v',sK7(X3,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2)),'type.bvlSem.v'))))) ) | $spl14),
  inference(superposition,[],[f3607,f3326])).
fof(f3607,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X79) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK7(X78,X77,X76)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X76,bool),X77),s('type.list.list'(X76),X78))) | s(bool,'F_0') = s(bool,X79)) ) | $spl14),
  inference(resolution,[],[f3264,f770])).
fof(f12828,plain,(
  ( ! [X146,X144,X151,X149,X147,X145,X150,X148] : (s(bool,X144) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X145),X146))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X147),X148))))) | ~p(s(bool,X144)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X147),X148),s('type.bvlSem.v',X149)))) | p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X151) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X145),X146),s('type.bvlSem.v',X149)))) ) | $spl106),
  inference(resolution,[],[f3917,f142])).
fof(f12821,plain,(
  ( ! [X103,X101,X99,X107,X105,X102,X100,X106,X104] : (s(bool,X99) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X100),X101))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X102),X103))))) | ~p(s(bool,X99)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X102),X103),s('type.bvlSem.v',X104)))) | ~p(s(bool,X105)) | ~p(s(bool,X106)) | s(bool,X106) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X100),X101),s('type.bvlSem.v',X104))) | s(bool,X105) != s(bool,X107)) ) | $spl106),
  inference(resolution,[],[f3917,f105])).
fof(f12819,plain,(
  ( ! [X88,X87,X85,X83,X81,X89,X86,X84,X82] : (s(bool,X81) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X82),X83))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X84),X85))))) | ~p(s(bool,X81)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X84),X85),s('type.bvlSem.v',X86)))) | p(s(bool,X87)) | ~p(s(bool,X88)) | ~p(s(bool,X89)) | s(bool,X88) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X82),X83),s('type.bvlSem.v',X86)))) ) | $spl106),
  inference(resolution,[],[f3917,f104])).
fof(f12817,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s(bool,X63) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X64),X65))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X66),X67))))) | ~p(s(bool,X63)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X66),X67),s('type.bvlSem.v',X68)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,X71) | s(bool,X70) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X64),X65),s('type.bvlSem.v',X68)))) ) | $spl106),
  inference(resolution,[],[f3917,f103])).
fof(f12816,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55] : (s(bool,X54) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X55),X56))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X57),X58))))) | ~p(s(bool,X54)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X57),X58),s('type.bvlSem.v',X59)))) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X55),X56),s('type.bvlSem.v',X59))) | s(bool,X60) != s(bool,X62)) ) | $spl106),
  inference(resolution,[],[f3917,f101])).
fof(f12814,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (s(bool,X36) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X37),X38))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X39),X40))))) | ~p(s(bool,X36)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X39),X40),s('type.bvlSem.v',X41)))) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X37),X38),s('type.bvlSem.v',X41)))) ) | $spl106),
  inference(resolution,[],[f3917,f99])).
fof(f12802,plain,(
  ( ! [X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X42),X43),s('type.bvlSem.v',sK7(X44,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X42),X43)),'type.bvlSem.v')))) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X42),X43))),s('type.list.list'('type.bvlSem.v'),X44))) | p(s(bool,X46))) ) | $spl342),
  inference(cnf_transformation,[],[f12802_D])).
fof(f12802_D,plain,(
  ( ! [X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X42),X43),s('type.bvlSem.v',sK7(X44,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X42),X43)),'type.bvlSem.v')))) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X42),X43))),s('type.list.list'('type.bvlSem.v'),X44))) | p(s(bool,X46))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f12775,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X108),X109),s('type.bvlSem.v',sK7(X110,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X108),X109)),'type.bvlSem.v')))) | p(s(bool,X111)) | p(s(bool,X112)) | s(bool,X112) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X108),X109))),s('type.list.list'('type.bvlSem.v'),X110)))) )),
  inference(resolution,[],[f5071,f142])).
fof(f5071,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X14),X15))),s('type.list.list'('type.bvlSem.v'),X16)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X14),X15),s('type.bvlSem.v',sK7(X16,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X14),X15)),'type.bvlSem.v'))))) )),
  inference(trivial_inequality_removal,[],[f4969])).
fof(f4969,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X14),X15))),s('type.list.list'('type.bvlSem.v'),X16)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X14),X15),s('type.bvlSem.v',sK7(X16,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X14),X15)),'type.bvlSem.v'))))) )),
  inference(superposition,[],[f3284,f3326])).
fof(f3284,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,happ(s(fun(X54,bool),X55),s(X54,sK7(X56,X55,X54)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X54,bool),X55),s('type.list.list'(X54),X56))))) )),
  inference(resolution,[],[f96,f112])).
fof(f12768,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X84),X85),s('type.bvlSem.v',sK7(X86,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X84),X85)),'type.bvlSem.v')))) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X84),X85))),s('type.list.list'('type.bvlSem.v'),X86))) | s(bool,X88) != s(bool,X89)) )),
  inference(resolution,[],[f5071,f105])).
fof(f12765,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X66),X67),s('type.bvlSem.v',sK7(X68,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X66),X67)),'type.bvlSem.v')))) | p(s(bool,X69)) | ~p(s(bool,X70)) | p(s(bool,X71)) | s(bool,X71) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X66),X67))),s('type.list.list'('type.bvlSem.v'),X68)))) )),
  inference(resolution,[],[f5071,f104])).
fof(f12764,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X60),X61),s('type.bvlSem.v',sK7(X62,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X60),X61)),'type.bvlSem.v')))) | p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X63) != s(bool,X65) | s(bool,X64) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X60),X61))),s('type.list.list'('type.bvlSem.v'),X62)))) )),
  inference(resolution,[],[f5071,f103])).
fof(f12763,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X54),X55),s('type.bvlSem.v',sK7(X56,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X54),X55)),'type.bvlSem.v')))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X54),X55))),s('type.list.list'('type.bvlSem.v'),X56))) | s(bool,X58) != s(bool,X59)) )),
  inference(resolution,[],[f5071,f101])).
fof(f12761,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X42),X43),s('type.bvlSem.v',sK7(X44,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X42),X43)),'type.bvlSem.v')))) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X42),X43))),s('type.list.list'('type.bvlSem.v'),X44)))) )),
  inference(resolution,[],[f5071,f99])).
fof(f12745,plain,(
  ( ! [X12,X10,X8,X15,X13,X11,X9] : (s(bool,X8) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X11))) | s(bool,X15) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X12),X13))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))))) | ~p(s(bool,X15)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X11)))) | p(s(bool,X8))) ) | $spl340),
  inference(cnf_transformation,[],[f12745_D])).
fof(f12745_D,plain,(
  ( ! [X12,X10,X8,X15,X13,X11,X9] : (s(bool,X8) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X11))) | s(bool,X15) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X12),X13))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))))) | ~p(s(bool,X15)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X11)))) | p(s(bool,X8))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f12732,plain,(
  ( ! [X123,X121,X118,X116,X122,X120,X119,X117] : (s(bool,X116) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X117),X118),s('type.bvlSem.v',X119))) | p(s(bool,X116)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X120),X121),s('type.bvlSem.v',X119)))) | p(s(bool,X122)) | ~p(s(bool,X123)) | s(bool,X123) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X120),X121))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X117),X118)))))) ) | $spl104),
  inference(resolution,[],[f3916,f142])).
fof(f12725,plain,(
  ( ! [X78,X76,X74,X72,X71,X79,X77,X75,X73] : (s(bool,X71) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X72),X73),s('type.bvlSem.v',X74))) | p(s(bool,X71)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X75),X76),s('type.bvlSem.v',X74)))) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X75),X76))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X72),X73))))) | s(bool,X77) != s(bool,X79)) ) | $spl104),
  inference(resolution,[],[f3916,f105])).
fof(f12723,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (s(bool,X53) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X54),X55),s('type.bvlSem.v',X56))) | p(s(bool,X53)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X57),X58),s('type.bvlSem.v',X56)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X57),X58))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X54),X55)))))) ) | $spl104),
  inference(resolution,[],[f3916,f104])).
fof(f12721,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X42,X40] : (s(bool,X35) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X36),X37),s('type.bvlSem.v',X38))) | p(s(bool,X35)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X39),X40),s('type.bvlSem.v',X38)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X39),X40))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X36),X37)))))) ) | $spl104),
  inference(resolution,[],[f3916,f103])).
fof(f12720,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X34,X32] : (s(bool,X26) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X27),X28),s('type.bvlSem.v',X29))) | p(s(bool,X26)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X30),X31),s('type.bvlSem.v',X29)))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X30),X31))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X27),X28))))) | s(bool,X32) != s(bool,X34)) ) | $spl104),
  inference(resolution,[],[f3916,f101])).
fof(f12718,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (s(bool,X8) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X11))) | p(s(bool,X8)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X12),X13),s('type.bvlSem.v',X11)))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X12),X13))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10)))))) ) | $spl104),
  inference(resolution,[],[f3916,f99])).
fof(f12691,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9))),s('type.list.list'('type.bvlSem.v'),X10))) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X8),X9),s('type.bvlSem.v',sK7(X10,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9)),'type.bvlSem.v')))) | ~p(s(bool,X12))) ) | $spl338),
  inference(cnf_transformation,[],[f12691_D])).
fof(f12691_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9))),s('type.list.list'('type.bvlSem.v'),X10))) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X8),X9),s('type.bvlSem.v',sK7(X10,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9)),'type.bvlSem.v')))) | ~p(s(bool,X12))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f12667,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X74),X75))),s('type.list.list'('type.bvlSem.v'),X76))) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X74),X75),s('type.bvlSem.v',sK7(X76,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X74),X75)),'type.bvlSem.v'))))) )),
  inference(resolution,[],[f3616,f142])).
fof(f3616,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',sK7(X2,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X0),X1)),'type.bvlSem.v'))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))),s('type.list.list'('type.bvlSem.v'),X2)))) )),
  inference(superposition,[],[f3264,f97])).
fof(f12660,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X50),X51))),s('type.list.list'('type.bvlSem.v'),X52))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X50),X51),s('type.bvlSem.v',sK7(X52,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X50),X51)),'type.bvlSem.v')))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f3616,f105])).
fof(f12658,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X38),X39))),s('type.list.list'('type.bvlSem.v'),X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X38),X39),s('type.bvlSem.v',sK7(X40,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X38),X39)),'type.bvlSem.v'))))) )),
  inference(resolution,[],[f3616,f104])).
fof(f12656,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X26),X27))),s('type.list.list'('type.bvlSem.v'),X28))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X26),X27),s('type.bvlSem.v',sK7(X28,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X26),X27)),'type.bvlSem.v'))))) )),
  inference(resolution,[],[f3616,f103])).
fof(f12655,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X20),X21))),s('type.list.list'('type.bvlSem.v'),X22))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X20),X21),s('type.bvlSem.v',sK7(X22,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X20),X21)),'type.bvlSem.v')))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f3616,f101])).
fof(f12653,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9))),s('type.list.list'('type.bvlSem.v'),X10))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X8),X9),s('type.bvlSem.v',sK7(X10,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9)),'type.bvlSem.v'))))) )),
  inference(resolution,[],[f3616,f99])).
fof(f12486,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8))))) | ~p(s(bool,X13)) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9)))) ) | $spl336),
  inference(cnf_transformation,[],[f12486_D])).
fof(f12486_D,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8))))) | ~p(s(bool,X13)) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9)))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f12473,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X101),X102),s('type.bvlSem.v',X103)))) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X104),X105),s('type.bvlSem.v',X103))) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X107) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X104),X105))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X101),X102)))))) )),
  inference(resolution,[],[f3887,f142])).
fof(f3887,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X84),X85))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X81),X82)))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X81),X82),s('type.bvlSem.v',X83)))) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X84),X85),s('type.bvlSem.v',X83)))) )),
  inference(resolution,[],[f118,f112])).
fof(f12466,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X63),X64),s('type.bvlSem.v',X65)))) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X66),X67),s('type.bvlSem.v',X65))) | ~p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X69) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X66),X67))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X63),X64))))) | s(bool,X68) != s(bool,X70)) )),
  inference(resolution,[],[f3887,f105])).
fof(f12464,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X47),X48),s('type.bvlSem.v',X49)))) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X50),X51),s('type.bvlSem.v',X49))) | p(s(bool,X52)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X50),X51))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X47),X48)))))) )),
  inference(resolution,[],[f3887,f104])).
fof(f12462,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X31),X32),s('type.bvlSem.v',X33)))) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X34),X35),s('type.bvlSem.v',X33))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X34),X35))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X31),X32)))))) )),
  inference(resolution,[],[f3887,f103])).
fof(f12461,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X23),X24),s('type.bvlSem.v',X25)))) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X26),X27),s('type.bvlSem.v',X25))) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X26),X27))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X23),X24))))) | s(bool,X28) != s(bool,X30)) )),
  inference(resolution,[],[f3887,f101])).
fof(f12459,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,'T_0') != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8)))))) )),
  inference(resolution,[],[f3887,f99])).
fof(f12413,plain,(
  ( ! [X4,X2] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X2),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | s(bool,X4) != s(bool,happ(s(fun('type.bvlSem.v',bool),X2),s('type.bvlSem.v',sK7(sK5,X2,'type.bvlSem.v')))) | ~p(s(bool,X4))) ) | $spl334),
  inference(cnf_transformation,[],[f12413_D])).
fof(f12413_D,plain,(
  ( ! [X4,X2] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X2),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | s(bool,X4) != s(bool,happ(s(fun('type.bvlSem.v',bool),X2),s('type.bvlSem.v',sK7(sK5,X2,'type.bvlSem.v')))) | ~p(s(bool,X4))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f12395,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X40),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,happ(s(fun('type.bvlSem.v',bool),X40),s('type.bvlSem.v',sK7(sK5,X40,'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f7510,f142])).
fof(f7510,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun('type.bvlSem.v',bool),X4),s('type.bvlSem.v',sK7(sK5,X4,'type.bvlSem.v'))))) | p(s(bool,happ(s(fun('type.bvlSem.v',bool),X4),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))))) ) | $spl3),
  inference(resolution,[],[f2660,f96])).
fof(f2660,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X0),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,happ(s(fun('type.bvlSem.v',bool),X0),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))))) ) | $spl3),
  inference(resolution,[],[f2205,f94])).
fof(f12388,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X30),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,happ(s(fun('type.bvlSem.v',bool),X30),s('type.bvlSem.v',sK7(sK5,X30,'type.bvlSem.v')))) | s(bool,X31) != s(bool,X33)) ) | $spl3),
  inference(resolution,[],[f7510,f105])).
fof(f12386,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X22),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun('type.bvlSem.v',bool),X22),s('type.bvlSem.v',sK7(sK5,X22,'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f7510,f104])).
fof(f12384,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X14),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun('type.bvlSem.v',bool),X14),s('type.bvlSem.v',sK7(sK5,X14,'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f7510,f103])).
fof(f12383,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X10),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,happ(s(fun('type.bvlSem.v',bool),X10),s('type.bvlSem.v',sK7(sK5,X10,'type.bvlSem.v')))) | s(bool,X11) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f7510,f101])).
fof(f12381,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X2),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun('type.bvlSem.v',bool),X2),s('type.bvlSem.v',sK7(sK5,X2,'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f7510,f99])).
fof(f12320,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8))))) | ~p(s(bool,X13)) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9)))) ) | $spl332),
  inference(cnf_transformation,[],[f12320_D])).
fof(f12320_D,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8))))) | ~p(s(bool,X13)) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9)))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f12307,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X101),X102),s('type.bvlSem.v',X103)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X104),X105),s('type.bvlSem.v',X103))) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X107) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X104),X105))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X101),X102)))))) )),
  inference(resolution,[],[f3885,f142])).
fof(f3885,plain,(
  ( ! [X74,X72,X71,X75,X73] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X74),X75))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X71),X72)))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X71),X72),s('type.bvlSem.v',X73)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X74),X75),s('type.bvlSem.v',X73)))) )),
  inference(resolution,[],[f118,f109])).
fof(f12300,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X63),X64),s('type.bvlSem.v',X65)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X66),X67),s('type.bvlSem.v',X65))) | ~p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X69) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X66),X67))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X63),X64))))) | s(bool,X68) != s(bool,X70)) )),
  inference(resolution,[],[f3885,f105])).
fof(f12298,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X47),X48),s('type.bvlSem.v',X49)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X50),X51),s('type.bvlSem.v',X49))) | p(s(bool,X52)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X50),X51))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X47),X48)))))) )),
  inference(resolution,[],[f3885,f104])).
fof(f12296,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X31),X32),s('type.bvlSem.v',X33)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X34),X35),s('type.bvlSem.v',X33))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X34),X35))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X31),X32)))))) )),
  inference(resolution,[],[f3885,f103])).
fof(f12295,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X23),X24),s('type.bvlSem.v',X25)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X26),X27),s('type.bvlSem.v',X25))) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X26),X27))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X23),X24))))) | s(bool,X28) != s(bool,X30)) )),
  inference(resolution,[],[f3885,f101])).
fof(f12293,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8)))))) )),
  inference(resolution,[],[f3885,f99])).
fof(f12162,plain,(
  ( ! [X54,X52,X50,X51] : (p(s(bool,'const.bool.IN_2'(s(bool,X50),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X51)))))) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X51))) | p(s(bool,X54)) | p(s(bool,X50)) | p(s(bool,sK7(X51,X52,bool)))) ) | $spl330),
  inference(cnf_transformation,[],[f12162_D])).
fof(f12162_D,plain,(
  ( ! [X54,X52,X50,X51] : (p(s(bool,'const.bool.IN_2'(s(bool,X50),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X51)))))) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X51))) | p(s(bool,X54)) | p(s(bool,X50)) | p(s(bool,sK7(X51,X52,bool)))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f12101,plain,(
  ( ! [X118,X116,X120,X119,X117] : (p(s(bool,'const.bool.IN_2'(s(bool,X116),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X117)))))) | p(s(bool,sK7(X117,X118,bool))) | p(s(bool,X116)) | p(s(bool,X119)) | p(s(bool,X120)) | s(bool,X120) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X118),s('type.list.list'(bool),X117)))) )),
  inference(resolution,[],[f5510,f142])).
fof(f5510,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | p(s(bool,sK7(X4,X5,bool))) | p(s(bool,X3))) )),
  inference(superposition,[],[f3528,f111])).
fof(f3528,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | p(s(bool,sK7(X0,X1,bool)))) )),
  inference(superposition,[],[f95,f109])).
fof(f12094,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,'const.bool.IN_2'(s(bool,X92),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X93)))))) | p(s(bool,sK7(X93,X94,bool))) | p(s(bool,X92)) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X94),s('type.list.list'(bool),X93))) | s(bool,X96) != s(bool,X97)) )),
  inference(resolution,[],[f5510,f105])).
fof(f12091,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (p(s(bool,'const.bool.IN_2'(s(bool,X74),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X75)))))) | p(s(bool,sK7(X75,X76,bool))) | p(s(bool,X74)) | p(s(bool,X77)) | ~p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X79) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X76),s('type.list.list'(bool),X75)))) )),
  inference(resolution,[],[f5510,f104])).
fof(f12090,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (p(s(bool,'const.bool.IN_2'(s(bool,X68),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X69)))))) | p(s(bool,sK7(X69,X70,bool))) | p(s(bool,X68)) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X70),s('type.list.list'(bool),X69)))) )),
  inference(resolution,[],[f5510,f103])).
fof(f12089,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,X62),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X63)))))) | p(s(bool,sK7(X63,X64,bool))) | p(s(bool,X62)) | p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X64),s('type.list.list'(bool),X63))) | s(bool,X66) != s(bool,X67)) )),
  inference(resolution,[],[f5510,f101])).
fof(f12087,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,'const.bool.IN_2'(s(bool,X50),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X51)))))) | p(s(bool,sK7(X51,X52,bool))) | p(s(bool,X50)) | p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X51)))) )),
  inference(resolution,[],[f5510,f99])).
fof(f12161,plain,(
  ( ! [X33,X36,X34,X32] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X32),s('type.list.list'(bool),X33)))) | p(s(bool,X34)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(bool,X34),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X33))))) | p(s(bool,X36)) | p(s(bool,sK7(X33,X32,bool)))) ) | $spl328),
  inference(cnf_transformation,[],[f12161_D])).
fof(f12161_D,plain,(
  ( ! [X33,X36,X34,X32] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X32),s('type.list.list'(bool),X33)))) | p(s(bool,X34)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(bool,X34),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X33))))) | p(s(bool,X36)) | p(s(bool,sK7(X33,X32,bool)))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f12066,plain,(
  ( ! [X101,X99,X102,X100,X98] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X98),s('type.list.list'(bool),X99)))) | p(s(bool,sK7(X99,X98,bool))) | p(s(bool,X100)) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X102) != s(bool,'const.bool.IN_2'(s(bool,X100),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X99)))))) )),
  inference(resolution,[],[f5510,f142])).
fof(f12059,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X74),s('type.list.list'(bool),X75)))) | p(s(bool,sK7(X75,X74,bool))) | p(s(bool,X76)) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,'const.bool.IN_2'(s(bool,X76),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X75))))) | s(bool,X78) != s(bool,X79)) )),
  inference(resolution,[],[f5510,f105])).
fof(f12056,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X56),s('type.list.list'(bool),X57)))) | p(s(bool,sK7(X57,X56,bool))) | p(s(bool,X58)) | p(s(bool,X59)) | ~p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X61) != s(bool,'const.bool.IN_2'(s(bool,X58),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X57)))))) )),
  inference(resolution,[],[f5510,f104])).
fof(f12055,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X50),s('type.list.list'(bool),X51)))) | p(s(bool,sK7(X51,X50,bool))) | p(s(bool,X52)) | p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(bool,X52),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X51)))))) )),
  inference(resolution,[],[f5510,f103])).
fof(f12054,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X44),s('type.list.list'(bool),X45)))) | p(s(bool,sK7(X45,X44,bool))) | p(s(bool,X46)) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X45))))) | s(bool,X48) != s(bool,X49)) )),
  inference(resolution,[],[f5510,f101])).
fof(f12052,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X32),s('type.list.list'(bool),X33)))) | p(s(bool,sK7(X33,X32,bool))) | p(s(bool,X34)) | p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(bool,X34),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X33)))))) )),
  inference(resolution,[],[f5510,f99])).
fof(f11982,plain,(
  ( ! [X4,X7,X5] : (p(s(bool,X4)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X5),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X4) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X5)))) ) | $spl326),
  inference(cnf_transformation,[],[f11982_D])).
fof(f11982_D,plain,(
  ( ! [X4,X7,X5] : (p(s(bool,X4)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X5),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X4) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X5)))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f11968,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X56)) | s(bool,X56) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X57))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X57),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | $spl96),
  inference(resolution,[],[f10917,f142])).
fof(f10917,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) | p(s(bool,X0)) | s(bool,X0) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X1)))) ) | $spl96),
  inference(forward_demodulation,[],[f10916,f97])).
fof(f10916,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X1))) | p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))))) ) | $spl96),
  inference(subsumption_resolution,[],[f10902,f136])).
fof(f10902,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X1))) | p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))))) ) | $spl96),
  inference(superposition,[],[f3753,f523])).
fof(f11961,plain,(
  ( ! [X39,X43,X41,X42,X40] : (p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X40))) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X40),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X41) != s(bool,X43)) ) | $spl96),
  inference(resolution,[],[f10917,f105])).
fof(f11959,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X30))) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X30),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | $spl96),
  inference(resolution,[],[f10917,f104])).
fof(f11957,plain,(
  ( ! [X23,X21,X19,X22,X20] : (p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X20),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | $spl96),
  inference(resolution,[],[f10917,f103])).
fof(f11956,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X15))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X15),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X16) != s(bool,X18)) ) | $spl96),
  inference(resolution,[],[f10917,f101])).
fof(f11954,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X5))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X5),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | $spl96),
  inference(resolution,[],[f10917,f99])).
fof(f11934,plain,(
  ( ! [X4,X7,X5] : (s(bool,X4) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X5))) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X5),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X4))) ) | $spl324),
  inference(cnf_transformation,[],[f11934_D])).
fof(f11934_D,plain,(
  ( ! [X4,X7,X5] : (s(bool,X4) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X5))) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X5),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X4))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f11919,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X56) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X57))) | p(s(bool,X56)) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X57),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | ($spl6 | $spl14)),
  inference(resolution,[],[f2782,f142])).
fof(f2782,plain,(
  ( ! [X45,X44] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X44),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | s(bool,X45) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X44))) | p(s(bool,X45))) ) | ($spl6 | $spl14)),
  inference(resolution,[],[f2269,f181])).
fof(f11912,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X39) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X40))) | p(s(bool,X39)) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X40),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X41) != s(bool,X43)) ) | ($spl6 | $spl14)),
  inference(resolution,[],[f2782,f105])).
fof(f11910,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X29) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X30))) | p(s(bool,X29)) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X30),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | ($spl6 | $spl14)),
  inference(resolution,[],[f2782,f104])).
fof(f11908,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X19) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X20))) | p(s(bool,X19)) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X20),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | ($spl6 | $spl14)),
  inference(resolution,[],[f2782,f103])).
fof(f11907,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X14) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X15))) | p(s(bool,X14)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X15),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X16) != s(bool,X18)) ) | ($spl6 | $spl14)),
  inference(resolution,[],[f2782,f101])).
fof(f11905,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,X4) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X5))) | p(s(bool,X4)) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X5),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | ($spl6 | $spl14)),
  inference(resolution,[],[f2782,f99])).
fof(f11713,plain,(
  ( ! [X30,X28,X31,X29,X34,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X28),X29),s('type.bvlSem.v',X30)))) | s(bool,X34) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X31),X32),s('type.bvlSem.v',X30))) | ~p(s(bool,X34)) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X31),X32))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X28),X29)))))) ) | $spl322),
  inference(cnf_transformation,[],[f11713_D])).
fof(f11713_D,plain,(
  ( ! [X30,X28,X31,X29,X34,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X28),X29),s('type.bvlSem.v',X30)))) | s(bool,X34) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X31),X32),s('type.bvlSem.v',X30))) | ~p(s(bool,X34)) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X31),X32))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X28),X29)))))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f11693,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X122),X123),s('type.bvlSem.v',X124)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X125),X126))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X122),X123))))) | p(s(bool,X127)) | ~p(s(bool,X128)) | s(bool,X128) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X125),X126),s('type.bvlSem.v',X124)))) )),
  inference(resolution,[],[f3865,f142])).
fof(f3865,plain,(
  ( ! [X68,X66,X64,X67,X65] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X67),X68),s('type.bvlSem.v',X66)))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X64),X65),s('type.bvlSem.v',X66)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X67),X68))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X64),X65)))))) )),
  inference(resolution,[],[f118,f109])).
fof(f11686,plain,(
  ( ! [X90,X88,X87,X85,X91,X89,X86,X84] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X84),X85),s('type.bvlSem.v',X86)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X87),X88))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X84),X85))))) | ~p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X87),X88),s('type.bvlSem.v',X86))) | s(bool,X89) != s(bool,X91)) )),
  inference(resolution,[],[f3865,f105])).
fof(f11684,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X68),X69),s('type.bvlSem.v',X70)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X71),X72))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X68),X69))))) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~p(s(bool,X75)) | s(bool,X74) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X71),X72),s('type.bvlSem.v',X70)))) )),
  inference(resolution,[],[f3865,f104])).
fof(f11682,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X52),X53),s('type.bvlSem.v',X54)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X55),X56))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X52),X53))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,X59) | s(bool,X58) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X55),X56),s('type.bvlSem.v',X54)))) )),
  inference(resolution,[],[f3865,f103])).
fof(f11681,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X51,X49] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X44),X45),s('type.bvlSem.v',X46)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X47),X48))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X44),X45))))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X47),X48),s('type.bvlSem.v',X46))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f3865,f101])).
fof(f11679,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X28),X29),s('type.bvlSem.v',X30)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X31),X32))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X28),X29))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X31),X32),s('type.bvlSem.v',X30)))) )),
  inference(resolution,[],[f3865,f99])).
fof(f11586,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,X2) != s(bool,X4) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X0) != s(bool,X2)) ) | $spl320),
  inference(cnf_transformation,[],[f11586_D])).
fof(f11586_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,X2) != s(bool,X4) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f11504,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,X94) != s(bool,X95) | s(bool,X94) != s(bool,X96) | s(bool,'T_0') = s(bool,X97) | s(bool,X97) != s(bool,X98) | s(bool,X96) != s(bool,X98) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X95) != s(bool,X100)) ) | $spl122),
  inference(resolution,[],[f4223,f142])).
fof(f11497,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,'T_0') = s(bool,X59) | s(bool,X59) != s(bool,X60) | s(bool,X58) != s(bool,X60) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X57) != s(bool,X62) | s(bool,X61) != s(bool,X63)) ) | $spl122),
  inference(resolution,[],[f4223,f105])).
fof(f11495,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,X40) != s(bool,X42) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) != s(bool,X44) | s(bool,X42) != s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X41) != s(bool,X46)) ) | $spl122),
  inference(resolution,[],[f4223,f104])).
fof(f11493,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') = s(bool,X27) | s(bool,X27) != s(bool,X28) | s(bool,X26) != s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X25) != s(bool,X30)) ) | $spl122),
  inference(resolution,[],[f4223,f103])).
fof(f11492,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,'T_0') = s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,X18) != s(bool,X20) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X17) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl122),
  inference(resolution,[],[f4223,f101])).
fof(f11490,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X2) != s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X1) != s(bool,X6)) ) | $spl122),
  inference(resolution,[],[f4223,f99])).
fof(f11478,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,sK7(X4,X5,bool))) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))) | ~p(s(bool,X8)) | s(bool,'T_0') = s(bool,X6) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4))))) ) | $spl318),
  inference(cnf_transformation,[],[f11478_D])).
fof(f11478_D,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,sK7(X4,X5,bool))) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))) | ~p(s(bool,X8)) | s(bool,'T_0') = s(bool,X6) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4))))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f11430,plain,(
  ( ! [X70,X74,X72,X71,X73] : (p(s(bool,sK7(X70,X71,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X71),s('type.list.list'(bool),X70)))) | s(bool,'T_0') = s(bool,X72) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(bool,bool),X71),s(bool,X72)))) )),
  inference(resolution,[],[f10651,f142])).
fof(f10651,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X27)))) | p(s(bool,sK7(X29,X28,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X28),s('type.list.list'(bool),X29)))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f10614,f127])).
fof(f127,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : ((p(s(bool,'F_0')) & p(s(bool,X7))) <=> p(s(bool,'F_0')))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.AND_CLAUSES|split|2')).
fof(f10614,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X27)))) | p(s(bool,sK7(X29,X28,bool))) | p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X28),s('type.list.list'(bool),X29)))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f10551,f116])).
fof(f10551,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X13)))) | p(s(bool,sK7(X14,X12,bool))) | p(s(bool,X13)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X12),s('type.list.list'(bool),X14))))) )),
  inference(trivial_inequality_removal,[],[f10547])).
fof(f10547,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X12),s('type.list.list'(bool),X14)))) | p(s(bool,sK7(X14,X12,bool))) | p(s(bool,X13)) | ~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X13))))) )),
  inference(superposition,[],[f5609,f115])).
fof(f5609,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X4),s(bool,X3))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5)))) | p(s(bool,sK7(X5,X4,bool))) | p(s(bool,X3))) )),
  inference(superposition,[],[f3819,f111])).
fof(f3819,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | p(s(bool,sK7(X0,X1,bool)))) )),
  inference(superposition,[],[f3284,f109])).
fof(f11423,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (p(s(bool,sK7(X46,X47,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X47),s('type.list.list'(bool),X46)))) | s(bool,'T_0') = s(bool,X48) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,happ(s(fun(bool,bool),X47),s(bool,X48))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f10651,f105])).
fof(f11421,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,sK7(X34,X35,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34)))) | s(bool,'T_0') = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,happ(s(fun(bool,bool),X35),s(bool,X36)))) )),
  inference(resolution,[],[f10651,f104])).
fof(f11419,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (p(s(bool,sK7(X22,X23,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X22)))) | s(bool,'T_0') = s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,happ(s(fun(bool,bool),X23),s(bool,X24)))) )),
  inference(resolution,[],[f10651,f103])).
fof(f11418,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (p(s(bool,sK7(X16,X17,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X17),s('type.list.list'(bool),X16)))) | s(bool,'T_0') = s(bool,X18) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f10651,f101])).
fof(f11416,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (p(s(bool,sK7(X4,X5,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | s(bool,'T_0') = s(bool,X6) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,X6)))) )),
  inference(resolution,[],[f10651,f99])).
fof(f11413,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))))) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9)))) ) | $spl316),
  inference(cnf_transformation,[],[f11413_D])).
fof(f11413_D,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))))) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9)))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f11400,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X101),X102),s('type.bvlSem.v',X103)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X104),X105),s('type.bvlSem.v',X103))) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X107) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X101),X102))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X104),X105)))))) )),
  inference(resolution,[],[f3847,f142])).
fof(f3847,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X83),X84))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X85),X86)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X83),X84),s('type.bvlSem.v',X87)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X85),X86),s('type.bvlSem.v',X87)))) )),
  inference(resolution,[],[f118,f113])).
fof(f11393,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X63),X64),s('type.bvlSem.v',X65)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X66),X67),s('type.bvlSem.v',X65))) | ~p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X69) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X63),X64))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X66),X67))))) | s(bool,X68) != s(bool,X70)) )),
  inference(resolution,[],[f3847,f105])).
fof(f11391,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X47),X48),s('type.bvlSem.v',X49)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X50),X51),s('type.bvlSem.v',X49))) | p(s(bool,X52)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X47),X48))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X50),X51)))))) )),
  inference(resolution,[],[f3847,f104])).
fof(f11389,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X31),X32),s('type.bvlSem.v',X33)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X34),X35),s('type.bvlSem.v',X33))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X31),X32))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X34),X35)))))) )),
  inference(resolution,[],[f3847,f103])).
fof(f11388,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X23),X24),s('type.bvlSem.v',X25)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X26),X27),s('type.bvlSem.v',X25))) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X23),X24))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X26),X27))))) | s(bool,X28) != s(bool,X30)) )),
  inference(resolution,[],[f3847,f101])).
fof(f11386,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11)))))) )),
  inference(resolution,[],[f3847,f99])).
fof(f11225,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,X2))) | p(s(bool,sK7(X1,X0,bool)))) ) | $spl314),
  inference(cnf_transformation,[],[f11225_D])).
fof(f11225_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,X2))) | p(s(bool,sK7(X1,X0,bool)))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f11117,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X66),s('type.list.list'(bool),X67)))) | p(s(bool,sK7(X67,X66,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X66),s(bool,X68))) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X68) != s(bool,X70)) )),
  inference(resolution,[],[f10550,f142])).
fof(f10550,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X15),s('type.list.list'(bool),X17)))) | p(s(bool,X16)) | p(s(bool,sK7(X17,X15,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,X16)))) )),
  inference(trivial_inequality_removal,[],[f10548])).
fof(f10548,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X15),s('type.list.list'(bool),X17)))) | p(s(bool,sK7(X17,X15,bool))) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,X16)))) )),
  inference(superposition,[],[f5609,f116])).
fof(f11110,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X42),s('type.list.list'(bool),X43)))) | p(s(bool,sK7(X43,X42,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X42),s(bool,X44))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f10550,f105])).
fof(f11107,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X24),s('type.list.list'(bool),X25)))) | p(s(bool,sK7(X25,X24,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,X26))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) )),
  inference(resolution,[],[f10550,f104])).
fof(f11106,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) | p(s(bool,sK7(X19,X18,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,X20))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f10550,f103])).
fof(f11105,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X12),s('type.list.list'(bool),X13)))) | p(s(bool,sK7(X13,X12,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X12),s(bool,X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f10550,f101])).
fof(f11103,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | p(s(bool,sK7(X1,X0,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,X2))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f10550,f99])).
fof(f11224,plain,(
  ( ! [X54,X52,X50,X51] : (p(s(bool,X50)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X51))) | p(s(bool,X54)) | p(s(bool,sK7(X51,X52,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,X50)))) ) | $spl312),
  inference(cnf_transformation,[],[f11224_D])).
fof(f11224_D,plain,(
  ( ! [X54,X52,X50,X51] : (p(s(bool,X50)) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X51))) | p(s(bool,X54)) | p(s(bool,sK7(X51,X52,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,X50)))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f11095,plain,(
  ( ! [X118,X116,X120,X119,X117] : (p(s(bool,X116)) | p(s(bool,sK7(X117,X118,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X118),s(bool,X116))) | p(s(bool,X119)) | p(s(bool,X120)) | s(bool,X120) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X118),s('type.list.list'(bool),X117)))) )),
  inference(resolution,[],[f10550,f142])).
fof(f11088,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,X92)) | p(s(bool,sK7(X93,X94,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X94),s(bool,X92))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X94),s('type.list.list'(bool),X93))) | s(bool,X96) != s(bool,X97)) )),
  inference(resolution,[],[f10550,f105])).
fof(f11085,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (p(s(bool,X74)) | p(s(bool,sK7(X75,X76,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X76),s(bool,X74))) | p(s(bool,X77)) | ~p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X79) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X76),s('type.list.list'(bool),X75)))) )),
  inference(resolution,[],[f10550,f104])).
fof(f11084,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (p(s(bool,X68)) | p(s(bool,sK7(X69,X70,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X70),s(bool,X68))) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X70),s('type.list.list'(bool),X69)))) )),
  inference(resolution,[],[f10550,f103])).
fof(f11083,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (p(s(bool,X62)) | p(s(bool,sK7(X63,X64,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X64),s(bool,X62))) | p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X64),s('type.list.list'(bool),X63))) | s(bool,X66) != s(bool,X67)) )),
  inference(resolution,[],[f10550,f101])).
fof(f11081,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,X50)) | p(s(bool,sK7(X51,X52,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,X50))) | p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X51)))) )),
  inference(resolution,[],[f10550,f99])).
fof(f10977,plain,(
  ( ! [X78,X76,X74,X72,X75,X73] : (s(bool,X72) != s(bool,'const.bool.IN_2'(s(X73,X74),s(fun(X73,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X73),X75))))) | s(bool,X78) != s(bool,'const.list.EVERY_2'(s(fun(X73,bool),X76),s('type.list.list'(X73),X75))) | ~p(s(bool,X78)) | p(s(bool,happ(s(fun(X73,bool),X76),s(X73,X74)))) | ~p(s(bool,X72))) ) | $spl310),
  inference(cnf_transformation,[],[f10977_D])).
fof(f10977_D,plain,(
  ( ! [X78,X76,X74,X72,X75,X73] : (s(bool,X72) != s(bool,'const.bool.IN_2'(s(X73,X74),s(fun(X73,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X73),X75))))) | s(bool,X78) != s(bool,'const.list.EVERY_2'(s(fun(X73,bool),X76),s('type.list.list'(X73),X75))) | ~p(s(bool,X78)) | p(s(bool,happ(s(fun(X73,bool),X76),s(X73,X74)))) | ~p(s(bool,X72))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f10962,plain,(
  ( ! [X167,X171,X169,X166,X172,X170,X168] : (s(bool,X166) != s(bool,'const.bool.IN_2'(s(X167,X168),s(fun(X167,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X167),X169))))) | ~p(s(bool,X166)) | p(s(bool,happ(s(fun(X167,bool),X170),s(X167,X168)))) | p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.list.EVERY_2'(s(fun(X167,bool),X170),s('type.list.list'(X167),X169)))) ) | $spl98),
  inference(resolution,[],[f3754,f142])).
fof(f10955,plain,(
  ( ! [X134,X132,X130,X128,X135,X133,X131,X129] : (s(bool,X128) != s(bool,'const.bool.IN_2'(s(X129,X130),s(fun(X129,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X129),X131))))) | ~p(s(bool,X128)) | p(s(bool,happ(s(fun(X129,bool),X132),s(X129,X130)))) | ~p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X134) != s(bool,'const.list.EVERY_2'(s(fun(X129,bool),X132),s('type.list.list'(X129),X131))) | s(bool,X133) != s(bool,X135)) ) | $spl98),
  inference(resolution,[],[f3754,f105])).
fof(f10953,plain,(
  ( ! [X118,X116,X114,X112,X119,X117,X115,X113] : (s(bool,X112) != s(bool,'const.bool.IN_2'(s(X113,X114),s(fun(X113,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X113),X115))))) | ~p(s(bool,X112)) | p(s(bool,happ(s(fun(X113,bool),X116),s(X113,X114)))) | p(s(bool,X117)) | ~p(s(bool,X118)) | ~p(s(bool,X119)) | s(bool,X118) != s(bool,'const.list.EVERY_2'(s(fun(X113,bool),X116),s('type.list.list'(X113),X115)))) ) | $spl98),
  inference(resolution,[],[f3754,f104])).
fof(f10951,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (s(bool,X96) != s(bool,'const.bool.IN_2'(s(X97,X98),s(fun(X97,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X97),X99))))) | ~p(s(bool,X96)) | p(s(bool,happ(s(fun(X97,bool),X100),s(X97,X98)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,'const.list.EVERY_2'(s(fun(X97,bool),X100),s('type.list.list'(X97),X99)))) ) | $spl98),
  inference(resolution,[],[f3754,f103])).
fof(f10950,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (s(bool,X88) != s(bool,'const.bool.IN_2'(s(X89,X90),s(fun(X89,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X89),X91))))) | ~p(s(bool,X88)) | p(s(bool,happ(s(fun(X89,bool),X92),s(X89,X90)))) | ~p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X94) != s(bool,'const.list.EVERY_2'(s(fun(X89,bool),X92),s('type.list.list'(X89),X91))) | s(bool,X93) != s(bool,X95)) ) | $spl98),
  inference(resolution,[],[f3754,f101])).
fof(f10948,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s(bool,X72) != s(bool,'const.bool.IN_2'(s(X73,X74),s(fun(X73,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X73),X75))))) | ~p(s(bool,X72)) | p(s(bool,happ(s(fun(X73,bool),X76),s(X73,X74)))) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.list.EVERY_2'(s(fun(X73,bool),X76),s('type.list.list'(X73),X75)))) ) | $spl98),
  inference(resolution,[],[f3754,f99])).
fof(f10911,plain,(
  ( ! [X78,X76,X74,X72,X75,X73] : (s(bool,X72) != s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s(bool,X78) != s(bool,'const.list.EVERY_2'(s(fun(X73,bool),X74),s('type.list.list'(X73),X76))) | ~p(s(bool,X78)) | ~p(s(bool,'const.bool.IN_2'(s(X73,X75),s(fun(X73,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X73),X76)))))) | p(s(bool,X72))) ) | $spl308),
  inference(cnf_transformation,[],[f10911_D])).
fof(f10911_D,plain,(
  ( ! [X78,X76,X74,X72,X75,X73] : (s(bool,X72) != s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s(bool,X78) != s(bool,'const.list.EVERY_2'(s(fun(X73,bool),X74),s('type.list.list'(X73),X76))) | ~p(s(bool,X78)) | ~p(s(bool,'const.bool.IN_2'(s(X73,X75),s(fun(X73,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X73),X76)))))) | p(s(bool,X72))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f10896,plain,(
  ( ! [X167,X171,X169,X166,X172,X170,X168] : (s(bool,X166) != s(bool,happ(s(fun(X167,bool),X168),s(X167,X169))) | p(s(bool,X166)) | ~p(s(bool,'const.bool.IN_2'(s(X167,X169),s(fun(X167,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X167),X170)))))) | p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.list.EVERY_2'(s(fun(X167,bool),X168),s('type.list.list'(X167),X170)))) ) | $spl96),
  inference(resolution,[],[f3753,f142])).
fof(f10889,plain,(
  ( ! [X134,X132,X130,X128,X135,X133,X131,X129] : (s(bool,X128) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X131))) | p(s(bool,X128)) | ~p(s(bool,'const.bool.IN_2'(s(X129,X131),s(fun(X129,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X129),X132)))))) | ~p(s(bool,X133)) | ~p(s(bool,X134)) | s(bool,X134) != s(bool,'const.list.EVERY_2'(s(fun(X129,bool),X130),s('type.list.list'(X129),X132))) | s(bool,X133) != s(bool,X135)) ) | $spl96),
  inference(resolution,[],[f3753,f105])).
fof(f10887,plain,(
  ( ! [X118,X116,X114,X112,X119,X117,X115,X113] : (s(bool,X112) != s(bool,happ(s(fun(X113,bool),X114),s(X113,X115))) | p(s(bool,X112)) | ~p(s(bool,'const.bool.IN_2'(s(X113,X115),s(fun(X113,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X113),X116)))))) | p(s(bool,X117)) | ~p(s(bool,X118)) | ~p(s(bool,X119)) | s(bool,X118) != s(bool,'const.list.EVERY_2'(s(fun(X113,bool),X114),s('type.list.list'(X113),X116)))) ) | $spl96),
  inference(resolution,[],[f3753,f104])).
fof(f10885,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (s(bool,X96) != s(bool,happ(s(fun(X97,bool),X98),s(X97,X99))) | p(s(bool,X96)) | ~p(s(bool,'const.bool.IN_2'(s(X97,X99),s(fun(X97,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X97),X100)))))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,'const.list.EVERY_2'(s(fun(X97,bool),X98),s('type.list.list'(X97),X100)))) ) | $spl96),
  inference(resolution,[],[f3753,f103])).
fof(f10884,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (s(bool,X88) != s(bool,happ(s(fun(X89,bool),X90),s(X89,X91))) | p(s(bool,X88)) | ~p(s(bool,'const.bool.IN_2'(s(X89,X91),s(fun(X89,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X89),X92)))))) | ~p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X94) != s(bool,'const.list.EVERY_2'(s(fun(X89,bool),X90),s('type.list.list'(X89),X92))) | s(bool,X93) != s(bool,X95)) ) | $spl96),
  inference(resolution,[],[f3753,f101])).
fof(f10882,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s(bool,X72) != s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | p(s(bool,X72)) | ~p(s(bool,'const.bool.IN_2'(s(X73,X75),s(fun(X73,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X73),X76)))))) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.list.EVERY_2'(s(fun(X73,bool),X74),s('type.list.list'(X73),X76)))) ) | $spl96),
  inference(resolution,[],[f3753,f99])).
fof(f10854,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X4) != s(bool,X6) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl306),
  inference(cnf_transformation,[],[f10854_D])).
fof(f10854_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X4) != s(bool,X6) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f10757,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,X94) != s(bool,X95) | s(bool,X94) != s(bool,X96) | s(bool,X95) != s(bool,X97) | s(bool,X97) != s(bool,X98) | s(bool,'F_0') = s(bool,X96) | p(s(bool,X99)) | p(s(bool,X100)) | s(bool,X98) != s(bool,X100)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1336,f142])).
fof(f1336,plain,(
  ( ! [X78,X76,X79,X77,X75] : (p(s(bool,X79)) | s(bool,X76) != s(bool,X77) | s(bool,X75) != s(bool,X76) | s(bool,X77) != s(bool,X78) | s(bool,X78) != s(bool,X79) | s(bool,'F_0') = s(bool,X75)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f836,f807])).
fof(f836,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X58)) | s(bool,'F_0') = s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X56) != s(bool,X57)) ) | $spl14),
  inference(resolution,[],[f770,f181])).
fof(f10750,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,X59) | s(bool,X59) != s(bool,X60) | s(bool,'F_0') = s(bool,X58) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X60) != s(bool,X61) | s(bool,X62) != s(bool,X63)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1336,f105])).
fof(f10747,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X34) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X36) != s(bool,X39)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1336,f104])).
fof(f10746,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,X27) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X28) != s(bool,X30)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1336,f103])).
fof(f10745,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1336,f101])).
fof(f10743,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X4) != s(bool,X6)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1336,f99])).
fof(f10741,plain,(
  ( ! [X4,X7,X5] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5)))) | s(bool,X7) != s(bool,sK7(X5,X4,bool)) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X5)))))) ) | $spl304),
  inference(cnf_transformation,[],[f10741_D])).
fof(f10741_D,plain,(
  ( ! [X4,X7,X5] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5)))) | s(bool,X7) != s(bool,sK7(X5,X4,bool)) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X5)))))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f10728,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X56),s('type.list.list'(bool),X57)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X57))))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,sK7(X57,X56,bool))) )),
  inference(resolution,[],[f9327,f142])).
fof(f9327,plain,(
  ( ! [X10,X11] : (~p(s(bool,sK7(X10,X11,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10)))))) )),
  inference(trivial_inequality_removal,[],[f9326])).
fof(f9326,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10)))) | ~p(s(bool,sK7(X10,X11,bool))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10)))))) )),
  inference(superposition,[],[f5415,f116])).
fof(f5415,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | ~p(s(bool,sK7(X4,X5,bool)))) )),
  inference(superposition,[],[f3507,f113])).
fof(f3507,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X48,sK7(X50,X49,X48)),s(fun(X48,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X48),X50))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X48,bool),X49),s('type.list.list'(X48),X50))))) )),
  inference(resolution,[],[f95,f108])).
fof(f10721,plain,(
  ( ! [X39,X43,X41,X42,X40] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X40)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X40))))) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,sK7(X40,X39,bool)) | s(bool,X41) != s(bool,X43)) )),
  inference(resolution,[],[f9327,f105])).
fof(f10719,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X29),s('type.list.list'(bool),X30)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X30))))) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,sK7(X30,X29,bool))) )),
  inference(resolution,[],[f9327,f104])).
fof(f10717,plain,(
  ( ! [X23,X21,X19,X22,X20] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X19),s('type.list.list'(bool),X20)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X20))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,sK7(X20,X19,bool))) )),
  inference(resolution,[],[f9327,f103])).
fof(f10716,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X14),s('type.list.list'(bool),X15)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X15))))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,sK7(X15,X14,bool)) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f9327,f101])).
fof(f10714,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X5))))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,sK7(X5,X4,bool))) )),
  inference(resolution,[],[f9327,f99])).
fof(f10710,plain,(
  ( ! [X6,X10,X8,X7] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X6),s('type.list.list'(bool),X7)))) | s(bool,X10) != s(bool,sK7(X7,X6,bool)) | ~p(s(bool,X10)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X6),s(bool,X8)))) ) | $spl302),
  inference(cnf_transformation,[],[f10710_D])).
fof(f10710_D,plain,(
  ( ! [X6,X10,X8,X7] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X6),s('type.list.list'(bool),X7)))) | s(bool,X10) != s(bool,sK7(X7,X6,bool)) | ~p(s(bool,X10)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X6),s(bool,X8)))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f10697,plain,(
  ( ! [X76,X74,X72,X75,X73] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X72),s('type.list.list'(bool),X73)))) | s(bool,X74) != s(bool,happ(s(fun(bool,bool),X72),s(bool,X74))) | ~p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,sK7(X73,X72,bool))) )),
  inference(resolution,[],[f5629,f142])).
fof(f5629,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,sK7(X5,X4,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5)))) | s(bool,X3) != s(bool,happ(s(fun(bool,bool),X4),s(bool,X3))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f3821,f115])).
fof(f3821,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0'))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | ~p(s(bool,sK7(X4,X5,bool)))) )),
  inference(superposition,[],[f3284,f113])).
fof(f10690,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X48),s('type.list.list'(bool),X49)))) | s(bool,X50) != s(bool,happ(s(fun(bool,bool),X48),s(bool,X50))) | ~p(s(bool,X50)) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,sK7(X49,X48,bool)) | s(bool,X51) != s(bool,X53)) )),
  inference(resolution,[],[f5629,f105])).
fof(f10688,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X36),s('type.list.list'(bool),X37)))) | s(bool,X38) != s(bool,happ(s(fun(bool,bool),X36),s(bool,X38))) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,sK7(X37,X36,bool))) )),
  inference(resolution,[],[f5629,f104])).
fof(f10686,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X24),s('type.list.list'(bool),X25)))) | s(bool,X26) != s(bool,happ(s(fun(bool,bool),X24),s(bool,X26))) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,sK7(X25,X24,bool))) )),
  inference(resolution,[],[f5629,f103])).
fof(f10685,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) | s(bool,X20) != s(bool,happ(s(fun(bool,bool),X18),s(bool,X20))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,sK7(X19,X18,bool)) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f5629,f101])).
fof(f10683,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X6),s('type.list.list'(bool),X7)))) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X6),s(bool,X8))) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,sK7(X7,X6,bool))) )),
  inference(resolution,[],[f5629,f99])).
fof(f10647,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,sK7(X4,X5,bool))) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))) | ~p(s(bool,X8)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | p(s(bool,X6))) ) | $spl300),
  inference(cnf_transformation,[],[f10647_D])).
fof(f10647_D,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,sK7(X4,X5,bool))) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))) | ~p(s(bool,X8)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | p(s(bool,X6))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f10600,plain,(
  ( ! [X70,X74,X72,X71,X73] : (p(s(bool,sK7(X70,X71,bool))) | p(s(bool,X72)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X71),s('type.list.list'(bool),X70)))) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(bool,bool),X71),s(bool,X72)))) )),
  inference(resolution,[],[f10551,f142])).
fof(f10593,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (p(s(bool,sK7(X46,X47,bool))) | p(s(bool,X48)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X47),s('type.list.list'(bool),X46)))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,happ(s(fun(bool,bool),X47),s(bool,X48))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f10551,f105])).
fof(f10591,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,sK7(X34,X35,bool))) | p(s(bool,X36)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,happ(s(fun(bool,bool),X35),s(bool,X36)))) )),
  inference(resolution,[],[f10551,f104])).
fof(f10589,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (p(s(bool,sK7(X22,X23,bool))) | p(s(bool,X24)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X22)))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,happ(s(fun(bool,bool),X23),s(bool,X24)))) )),
  inference(resolution,[],[f10551,f103])).
fof(f10588,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (p(s(bool,sK7(X16,X17,bool))) | p(s(bool,X18)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X17),s('type.list.list'(bool),X16)))) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f10551,f101])).
fof(f10586,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (p(s(bool,sK7(X4,X5,bool))) | p(s(bool,X6)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,X6)))) )),
  inference(resolution,[],[f10551,f99])).
fof(f10476,plain,(
  ( ! [X6,X4,X8,X5] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X4),s(bool,X6))) | ~p(s(bool,X8)) | p(s(bool,X6)) | p(s(bool,sK7(X5,X4,bool)))) ) | $spl298),
  inference(cnf_transformation,[],[f10476_D])).
fof(f10476_D,plain,(
  ( ! [X6,X4,X8,X5] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X4),s(bool,X6))) | ~p(s(bool,X8)) | p(s(bool,X6)) | p(s(bool,sK7(X5,X4,bool)))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f10429,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X70),s('type.list.list'(bool),X71))) | p(s(bool,sK7(X71,X70,bool))) | p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(bool,bool),X70),s(bool,X72)))) )),
  inference(resolution,[],[f5563,f142])).
fof(f5563,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X3)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | p(s(bool,sK7(X5,X4,bool))) | p(s(bool,X3))) )),
  inference(superposition,[],[f3610,f111])).
fof(f3610,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0))) | p(s(bool,sK7(X0,X1,bool)))) )),
  inference(superposition,[],[f3264,f109])).
fof(f10422,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X46),s('type.list.list'(bool),X47))) | p(s(bool,sK7(X47,X46,bool))) | p(s(bool,X48)) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,happ(s(fun(bool,bool),X46),s(bool,X48))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f5563,f105])).
fof(f10420,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X34),s('type.list.list'(bool),X35))) | p(s(bool,sK7(X35,X34,bool))) | p(s(bool,X36)) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,happ(s(fun(bool,bool),X34),s(bool,X36)))) )),
  inference(resolution,[],[f5563,f104])).
fof(f10418,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X22),s('type.list.list'(bool),X23))) | p(s(bool,sK7(X23,X22,bool))) | p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,happ(s(fun(bool,bool),X22),s(bool,X24)))) )),
  inference(resolution,[],[f5563,f103])).
fof(f10417,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X16),s('type.list.list'(bool),X17))) | p(s(bool,sK7(X17,X16,bool))) | p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,happ(s(fun(bool,bool),X16),s(bool,X18))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f5563,f101])).
fof(f10415,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | p(s(bool,sK7(X5,X4,bool))) | p(s(bool,X6)) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X4),s(bool,X6)))) )),
  inference(resolution,[],[f5563,f99])).
fof(f10384,plain,(
  ( ! [X14,X12,X13,X16] : (p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X13)))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X14))))) | s(bool,'T_0') = s(bool,X13) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X12),s('type.list.list'(bool),X14))))) ) | $spl296),
  inference(cnf_transformation,[],[f10384_D])).
fof(f10384_D,plain,(
  ( ! [X14,X12,X13,X16] : (p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X13)))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X14))))) | s(bool,'T_0') = s(bool,X13) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X12),s('type.list.list'(bool),X14))))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f10364,plain,(
  ( ! [X80,X78,X81,X79,X82] : (p(s(bool,happ(s(fun(bool,bool),X78),s(bool,X79)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X78),s('type.list.list'(bool),X80)))) | s(bool,'T_0') = s(bool,X79) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X80)))))) )),
  inference(resolution,[],[f3697,f142])).
fof(f3697,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X29)))))) | p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X27)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X28),s('type.list.list'(bool),X29)))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f94,f116])).
fof(f10357,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,happ(s(fun(bool,bool),X54),s(bool,X55)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X54),s('type.list.list'(bool),X56)))) | s(bool,'T_0') = s(bool,X55) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X56))))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f3697,f105])).
fof(f10355,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(bool,bool),X42),s(bool,X43)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X42),s('type.list.list'(bool),X44)))) | s(bool,'T_0') = s(bool,X43) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X44)))))) )),
  inference(resolution,[],[f3697,f104])).
fof(f10353,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),X30),s(bool,X31)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X30),s('type.list.list'(bool),X32)))) | s(bool,'T_0') = s(bool,X31) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X32)))))) )),
  inference(resolution,[],[f3697,f103])).
fof(f10352,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X25)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X24),s('type.list.list'(bool),X26)))) | s(bool,'T_0') = s(bool,X25) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X26))))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f3697,f101])).
fof(f10350,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X13)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X12),s('type.list.list'(bool),X14)))) | s(bool,'T_0') = s(bool,X13) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X14)))))) )),
  inference(resolution,[],[f3697,f99])).
fof(f10341,plain,(
  ( ! [X265,X267,X268] : (s(bool,X267) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X265,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X265),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X268) | s(bool,X267) != s(bool,X268)) ) | $spl294),
  inference(cnf_transformation,[],[f10341_D])).
fof(f10341_D,plain,(
  ( ! [X265,X267,X268] : (s(bool,X267) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X265,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X265),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X268) | s(bool,X267) != s(bool,X268)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f10340,plain,(
  ( ! [X265,X267,X268,X266] : (s(bool,X267) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X265,'type.bvlSem.v')))) | s(bool,X267) != s(bool,X268) | s(bool,'F_0') = s(bool,X266) | s(bool,'T_0') = s(bool,X268) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X265),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | ($spl14 | $spl40)),
  inference(subsumption_resolution,[],[f10278,f116])).
fof(f10278,plain,(
  ( ! [X265,X267,X268,X266] : (s(bool,'T_0') != s(bool,X266) | s(bool,X267) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X265,'type.bvlSem.v')))) | s(bool,X267) != s(bool,X268) | s(bool,'F_0') = s(bool,X266) | s(bool,'T_0') = s(bool,X268) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X265),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | ($spl14 | $spl40)),
  inference(superposition,[],[f1953,f7057])).
fof(f10331,plain,(
  ( ! [X212,X213,X211] : (s(bool,'F_0') = s(bool,X212) | s(bool,X213) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X211,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X213) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X211),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | $spl14),
  inference(subsumption_resolution,[],[f10259,f116])).
fof(f10259,plain,(
  ( ! [X212,X213,X211] : (s(bool,'T_0') != s(bool,X212) | s(bool,'F_0') = s(bool,X212) | s(bool,X213) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X211,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X213) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X211),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | $spl14),
  inference(superposition,[],[f1032,f7057])).
fof(f7075,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK6))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X3,'type.bvlSem.v'))))) ) | $spl198),
  inference(cnf_transformation,[],[f7075_D])).
fof(f7075_D,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK6))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X3,'type.bvlSem.v'))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f10327,plain,(
  ( ! [X191,X192,X190] : (p(s(bool,X191)) | p(s(bool,X192)) | s(bool,X192) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X190,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X190),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(subsumption_resolution,[],[f10250,f141])).
fof(f10250,plain,(
  ( ! [X191,X192,X190] : (~p(s(bool,'T_0')) | p(s(bool,X191)) | p(s(bool,X192)) | s(bool,X192) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X190,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X190),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(superposition,[],[f142,f7057])).
fof(f10325,plain,(
  ( ! [X154,X152,X155,X153] : (p(s(bool,X153)) | ~p(s(bool,X154)) | s(bool,X153) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X152,'type.bvlSem.v')))) | s(bool,X154) != s(bool,X155) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X152),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(subsumption_resolution,[],[f10226,f136])).
fof(f10226,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,'T_0')) | p(s(bool,X153)) | ~p(s(bool,X154)) | s(bool,X153) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X152,'type.bvlSem.v')))) | s(bool,X154) != s(bool,X155) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X152),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(superposition,[],[f105,f7057])).
fof(f10323,plain,(
  ( ! [X123,X121,X122,X120] : (p(s(bool,X121)) | ~p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X120,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X120),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(subsumption_resolution,[],[f10218,f136])).
fof(f10218,plain,(
  ( ! [X123,X121,X122,X120] : (~p(s(bool,'T_0')) | p(s(bool,X121)) | ~p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X120,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X120),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(superposition,[],[f104,f7057])).
fof(f10321,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X100,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X100),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(subsumption_resolution,[],[f10213,f141])).
fof(f10213,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X100,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X100),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(superposition,[],[f103,f7057])).
fof(f10318,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X72,'type.bvlSem.v')))) | s(bool,X74) != s(bool,X75) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X72),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(subsumption_resolution,[],[f10206,f136])).
fof(f10206,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X72,'type.bvlSem.v')))) | s(bool,X74) != s(bool,X75) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X72),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(superposition,[],[f101,f7057])).
fof(f10316,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X44,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X44),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(subsumption_resolution,[],[f10199,f141])).
fof(f10199,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X44,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X44),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(superposition,[],[f99,f7057])).
fof(f10173,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X6))))) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X6))))) ) | $spl290),
  inference(cnf_transformation,[],[f10173_D])).
fof(f10173_D,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X6))))) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X6))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f10156,plain,(
  ( ! [X70,X74,X72,X71,X73] : (p(s(bool,happ(s(fun(bool,bool),X70),s(bool,X71)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X70),s('type.list.list'(bool),X72)))) | s(bool,'F_0') = s(bool,X71) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X72)))))) )),
  inference(resolution,[],[f3694,f142])).
fof(f3694,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X20)))))) | p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X18)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X19),s('type.list.list'(bool),X20)))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f94,f116])).
fof(f10149,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X47)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X46),s('type.list.list'(bool),X48)))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X48))))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f3694,f105])).
fof(f10147,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X35)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X34),s('type.list.list'(bool),X36)))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X36)))))) )),
  inference(resolution,[],[f3694,f104])).
fof(f10145,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (p(s(bool,happ(s(fun(bool,bool),X22),s(bool,X23)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X22),s('type.list.list'(bool),X24)))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X24)))))) )),
  inference(resolution,[],[f3694,f103])).
fof(f10144,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X17)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X16),s('type.list.list'(bool),X18)))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X18))))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f3694,f101])).
fof(f10142,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X6)))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X6)))))) )),
  inference(resolution,[],[f3694,f99])).
fof(f10115,plain,(
  ( ! [X47,X50,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X47)))))) | s(bool,'T_0') = s(bool,X46) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,'F_0')))) | s(bool,X50) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X48),s('type.list.list'(bool),X47))) | ~p(s(bool,X50))) ) | $spl288),
  inference(cnf_transformation,[],[f10115_D])).
fof(f10115_D,plain,(
  ( ! [X47,X50,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X47)))))) | s(bool,'T_0') = s(bool,X46) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,'F_0')))) | s(bool,X50) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X48),s('type.list.list'(bool),X47))) | ~p(s(bool,X50))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f10104,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,'const.bool.IN_2'(s(bool,X112),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X113)))))) | p(s(bool,happ(s(fun(bool,bool),X114),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X112) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X116) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X114),s('type.list.list'(bool),X113)))) )),
  inference(resolution,[],[f3693,f142])).
fof(f3693,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X16),s('type.list.list'(bool),X17)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X15),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X17)))))) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f94,f116])).
fof(f10097,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (~p(s(bool,'const.bool.IN_2'(s(bool,X88),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X89)))))) | p(s(bool,happ(s(fun(bool,bool),X90),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X88) | ~p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X92) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X90),s('type.list.list'(bool),X89))) | s(bool,X91) != s(bool,X93)) )),
  inference(resolution,[],[f3693,f105])).
fof(f10095,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (~p(s(bool,'const.bool.IN_2'(s(bool,X76),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X77)))))) | p(s(bool,happ(s(fun(bool,bool),X78),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X76) | p(s(bool,X79)) | ~p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X78),s('type.list.list'(bool),X77)))) )),
  inference(resolution,[],[f3693,f104])).
fof(f10093,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X64),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X65)))))) | p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,X69) | s(bool,X68) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X66),s('type.list.list'(bool),X65)))) )),
  inference(resolution,[],[f3693,f103])).
fof(f10092,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,X58),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X59)))))) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X58) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X60),s('type.list.list'(bool),X59))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f3693,f101])).
fof(f10090,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X47)))))) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X46) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X48),s('type.list.list'(bool),X47)))) )),
  inference(resolution,[],[f3693,f99])).
fof(f10069,plain,(
  ( ! [X271,X268,X270] : (s(bool,X270) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X268))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X268),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X271) | s(bool,X270) != s(bool,X271)) ) | $spl286),
  inference(cnf_transformation,[],[f10069_D])).
fof(f10069_D,plain,(
  ( ! [X271,X268,X270] : (s(bool,X270) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X268))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X268),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X271) | s(bool,X270) != s(bool,X271)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f10068,plain,(
  ( ! [X269,X271,X268,X270] : (s(bool,X270) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X268))) | s(bool,X270) != s(bool,X271) | s(bool,'F_0') = s(bool,X269) | s(bool,'T_0') = s(bool,X271) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X268),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | ($spl6 | $spl14 | $spl40)),
  inference(subsumption_resolution,[],[f10007,f116])).
fof(f10007,plain,(
  ( ! [X269,X271,X268,X270] : (s(bool,'T_0') != s(bool,X269) | s(bool,X270) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X268))) | s(bool,X270) != s(bool,X271) | s(bool,'F_0') = s(bool,X269) | s(bool,'T_0') = s(bool,X271) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X268),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | ($spl6 | $spl14 | $spl40)),
  inference(superposition,[],[f1953,f6746])).
fof(f6746,plain,(
  ( ! [X34] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X34))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X34),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(resolution,[],[f2777,f109])).
fof(f10060,plain,(
  ( ! [X216,X214] : (s(bool,X216) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X214))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X214),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X216)) ) | $spl284),
  inference(cnf_transformation,[],[f10060_D])).
fof(f10060_D,plain,(
  ( ! [X216,X214] : (s(bool,X216) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X214))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X214),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X216)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f10059,plain,(
  ( ! [X216,X215,X214] : (s(bool,'F_0') = s(bool,X215) | s(bool,X216) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X214))) | s(bool,'T_0') = s(bool,X216) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X214),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f9988,f116])).
fof(f9988,plain,(
  ( ! [X216,X215,X214] : (s(bool,'T_0') != s(bool,X215) | s(bool,'F_0') = s(bool,X215) | s(bool,X216) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X214))) | s(bool,'T_0') = s(bool,X216) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X214),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f1032,f6746])).
fof(f7148,plain,(
  ( ! [X6,X4] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X4),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X4)))) ) | $spl202),
  inference(cnf_transformation,[],[f7148_D])).
fof(f7148_D,plain,(
  ( ! [X6,X4] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X4),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X4)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f10055,plain,(
  ( ! [X194,X195,X193] : (p(s(bool,X194)) | p(s(bool,X195)) | s(bool,X195) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X193))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X193),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9979,f141])).
fof(f9979,plain,(
  ( ! [X194,X195,X193] : (~p(s(bool,'T_0')) | p(s(bool,X194)) | p(s(bool,X195)) | s(bool,X195) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X193))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X193),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(superposition,[],[f142,f6746])).
fof(f10051,plain,(
  ( ! [X158,X156,X157,X155] : (p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X156) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X155))) | s(bool,X157) != s(bool,X158) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X155),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9955,f136])).
fof(f9955,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,'T_0')) | p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X156) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X155))) | s(bool,X157) != s(bool,X158) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X155),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(superposition,[],[f105,f6746])).
fof(f10049,plain,(
  ( ! [X125,X123,X126,X124] : (p(s(bool,X124)) | ~p(s(bool,X125)) | p(s(bool,X126)) | s(bool,X126) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X123))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X123),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9947,f136])).
fof(f9947,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,'T_0')) | p(s(bool,X124)) | ~p(s(bool,X125)) | p(s(bool,X126)) | s(bool,X126) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X123))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X123),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(superposition,[],[f104,f6746])).
fof(f10047,plain,(
  ( ! [X103,X105,X106,X104] : (p(s(bool,X104)) | p(s(bool,X105)) | s(bool,X104) != s(bool,X106) | s(bool,X105) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X103))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X103),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9942,f141])).
fof(f9942,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | p(s(bool,X105)) | s(bool,X104) != s(bool,X106) | s(bool,X105) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X103))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X103),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(superposition,[],[f103,f6746])).
fof(f10044,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X76) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X75))) | s(bool,X77) != s(bool,X78) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X75),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9935,f136])).
fof(f9935,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X76) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X75))) | s(bool,X77) != s(bool,X78) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X75),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(superposition,[],[f101,f6746])).
fof(f10042,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X48) != s(bool,X50) | s(bool,X49) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X47))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X47),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9928,f141])).
fof(f9928,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X48) != s(bool,X50) | s(bool,X49) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X47))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X47),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(superposition,[],[f99,f6746])).
fof(f9896,plain,(
  ( ! [X263,X265,X266] : (s(bool,X265) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X263,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X263),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X266) | s(bool,X265) != s(bool,X266)) ) | $spl282),
  inference(cnf_transformation,[],[f9896_D])).
fof(f9896_D,plain,(
  ( ! [X263,X265,X266] : (s(bool,X265) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X263,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X263),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X266) | s(bool,X265) != s(bool,X266)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f9895,plain,(
  ( ! [X263,X265,X264,X266] : (s(bool,X265) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X263,'type.bvlSem.v')))) | s(bool,X265) != s(bool,X266) | s(bool,'F_0') = s(bool,X264) | s(bool,'T_0') = s(bool,X266) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X263),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | ($spl6 | $spl14 | $spl40)),
  inference(subsumption_resolution,[],[f9832,f116])).
fof(f9832,plain,(
  ( ! [X263,X265,X264,X266] : (s(bool,'T_0') != s(bool,X264) | s(bool,X265) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X263,'type.bvlSem.v')))) | s(bool,X265) != s(bool,X266) | s(bool,'F_0') = s(bool,X264) | s(bool,'T_0') = s(bool,X266) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X263),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | ($spl6 | $spl14 | $spl40)),
  inference(superposition,[],[f1953,f6736])).
fof(f9887,plain,(
  ( ! [X211,X209] : (s(bool,X211) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X209,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X209),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X211)) ) | $spl280),
  inference(cnf_transformation,[],[f9887_D])).
fof(f9887_D,plain,(
  ( ! [X211,X209] : (s(bool,X211) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X209,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X209),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X211)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f9886,plain,(
  ( ! [X210,X211,X209] : (s(bool,'F_0') = s(bool,X210) | s(bool,X211) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X209,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X211) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X209),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f9813,f116])).
fof(f9813,plain,(
  ( ! [X210,X211,X209] : (s(bool,'T_0') != s(bool,X210) | s(bool,'F_0') = s(bool,X210) | s(bool,X211) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X209,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X211) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X209),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f1032,f6736])).
fof(f8266,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X3,'type.bvlSem.v'))))) ) | $spl230),
  inference(cnf_transformation,[],[f8266_D])).
fof(f8266_D,plain,(
  ( ! [X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X3,'type.bvlSem.v'))))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f9882,plain,(
  ( ! [X189,X190,X188] : (p(s(bool,X189)) | p(s(bool,X190)) | s(bool,X190) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X188,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X188),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(subsumption_resolution,[],[f9804,f141])).
fof(f9804,plain,(
  ( ! [X189,X190,X188] : (~p(s(bool,'T_0')) | p(s(bool,X189)) | p(s(bool,X190)) | s(bool,X190) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X188,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X188),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(superposition,[],[f142,f6736])).
fof(f9878,plain,(
  ( ! [X152,X151,X153,X150] : (p(s(bool,X151)) | ~p(s(bool,X152)) | s(bool,X151) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X150,'type.bvlSem.v')))) | s(bool,X152) != s(bool,X153) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X150),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(subsumption_resolution,[],[f9780,f136])).
fof(f9780,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X151)) | ~p(s(bool,X152)) | s(bool,X151) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X150,'type.bvlSem.v')))) | s(bool,X152) != s(bool,X153) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X150),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(superposition,[],[f105,f6736])).
fof(f9876,plain,(
  ( ! [X121,X118,X120,X119] : (p(s(bool,X119)) | ~p(s(bool,X120)) | p(s(bool,X121)) | s(bool,X121) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X118,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X118),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(subsumption_resolution,[],[f9772,f136])).
fof(f9772,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,'T_0')) | p(s(bool,X119)) | ~p(s(bool,X120)) | p(s(bool,X121)) | s(bool,X121) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X118,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X118),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(superposition,[],[f104,f6736])).
fof(f9874,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,X99)) | p(s(bool,X100)) | s(bool,X99) != s(bool,X101) | s(bool,X100) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X98,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X98),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(subsumption_resolution,[],[f9767,f141])).
fof(f9767,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X99)) | p(s(bool,X100)) | s(bool,X99) != s(bool,X101) | s(bool,X100) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X98,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X98),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(superposition,[],[f103,f6736])).
fof(f9871,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X71) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X70,'type.bvlSem.v')))) | s(bool,X72) != s(bool,X73) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X70),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(subsumption_resolution,[],[f9760,f136])).
fof(f9760,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X71) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X70,'type.bvlSem.v')))) | s(bool,X72) != s(bool,X73) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X70),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(superposition,[],[f101,f6736])).
fof(f9869,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X42,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X42),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(subsumption_resolution,[],[f9753,f141])).
fof(f9753,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X42,'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X42),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(superposition,[],[f99,f6736])).
fof(f9724,plain,(
  ( ! [X269,X268,X266] : (s(bool,X268) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X266))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X266),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,'T_0') = s(bool,X269) | s(bool,X268) != s(bool,X269)) ) | $spl278),
  inference(cnf_transformation,[],[f9724_D])).
fof(f9724_D,plain,(
  ( ! [X269,X268,X266] : (s(bool,X268) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X266))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X266),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,'T_0') = s(bool,X269) | s(bool,X268) != s(bool,X269)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f9723,plain,(
  ( ! [X269,X267,X268,X266] : (s(bool,X268) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X266))) | s(bool,X268) != s(bool,X269) | s(bool,'F_0') = s(bool,X267) | s(bool,'T_0') = s(bool,X269) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X266),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | ($spl14 | $spl40)),
  inference(subsumption_resolution,[],[f9664,f116])).
fof(f9664,plain,(
  ( ! [X269,X267,X268,X266] : (s(bool,'T_0') != s(bool,X267) | s(bool,X268) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X266))) | s(bool,X268) != s(bool,X269) | s(bool,'F_0') = s(bool,X267) | s(bool,'T_0') = s(bool,X269) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X266),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | ($spl14 | $spl40)),
  inference(superposition,[],[f1953,f6695])).
fof(f9714,plain,(
  ( ! [X212,X213,X214] : (s(bool,'F_0') = s(bool,X213) | s(bool,X214) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X212))) | s(bool,'T_0') = s(bool,X214) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X212),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | $spl14),
  inference(subsumption_resolution,[],[f9645,f116])).
fof(f9645,plain,(
  ( ! [X212,X213,X214] : (s(bool,'T_0') != s(bool,X213) | s(bool,'F_0') = s(bool,X213) | s(bool,X214) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X212))) | s(bool,'T_0') = s(bool,X214) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X212),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | $spl14),
  inference(superposition,[],[f1032,f6695])).
fof(f6712,plain,(
  ( ! [X6,X4] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X4),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X4)))) ) | $spl186),
  inference(cnf_transformation,[],[f6712_D])).
fof(f6712_D,plain,(
  ( ! [X6,X4] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X4),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X4)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f9710,plain,(
  ( ! [X191,X192,X193] : (p(s(bool,X192)) | p(s(bool,X193)) | s(bool,X193) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X191))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X191),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(subsumption_resolution,[],[f9636,f141])).
fof(f9636,plain,(
  ( ! [X191,X192,X193] : (~p(s(bool,'T_0')) | p(s(bool,X192)) | p(s(bool,X193)) | s(bool,X193) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X191))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X191),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(superposition,[],[f142,f6695])).
fof(f9708,plain,(
  ( ! [X156,X154,X155,X153] : (p(s(bool,X154)) | ~p(s(bool,X155)) | s(bool,X154) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X153))) | s(bool,X155) != s(bool,X156) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X153),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(subsumption_resolution,[],[f9612,f136])).
fof(f9612,plain,(
  ( ! [X156,X154,X155,X153] : (~p(s(bool,'T_0')) | p(s(bool,X154)) | ~p(s(bool,X155)) | s(bool,X154) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X153))) | s(bool,X155) != s(bool,X156) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X153),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(superposition,[],[f105,f6695])).
fof(f9706,plain,(
  ( ! [X123,X121,X124,X122] : (p(s(bool,X122)) | ~p(s(bool,X123)) | p(s(bool,X124)) | s(bool,X124) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X121))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X121),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(subsumption_resolution,[],[f9604,f136])).
fof(f9604,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,X122)) | ~p(s(bool,X123)) | p(s(bool,X124)) | s(bool,X124) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X121))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X121),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(superposition,[],[f104,f6695])).
fof(f9704,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,X103) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X101))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X101),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(subsumption_resolution,[],[f9599,f141])).
fof(f9599,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,X103) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X101))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X101),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(superposition,[],[f103,f6695])).
fof(f9701,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X74)) | ~p(s(bool,X75)) | s(bool,X74) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X73))) | s(bool,X75) != s(bool,X76) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X73),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(subsumption_resolution,[],[f9592,f136])).
fof(f9592,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X74)) | ~p(s(bool,X75)) | s(bool,X74) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X73))) | s(bool,X75) != s(bool,X76) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X73),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(superposition,[],[f101,f6695])).
fof(f9699,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X45))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X45),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(subsumption_resolution,[],[f9585,f141])).
fof(f9585,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X45))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X45),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(superposition,[],[f99,f6695])).
fof(f9557,plain,(
  ( ! [X47,X50,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X47)))))) | s(bool,'F_0') = s(bool,X46) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0')))) | s(bool,X50) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X48),s('type.list.list'(bool),X47))) | ~p(s(bool,X50))) ) | $spl274),
  inference(cnf_transformation,[],[f9557_D])).
fof(f9557_D,plain,(
  ( ! [X47,X50,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X47)))))) | s(bool,'F_0') = s(bool,X46) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0')))) | s(bool,X50) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X48),s('type.list.list'(bool),X47))) | ~p(s(bool,X50))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f9546,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,'const.bool.IN_2'(s(bool,X112),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X113)))))) | p(s(bool,happ(s(fun(bool,bool),X114),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X112) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X116) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X114),s('type.list.list'(bool),X113)))) )),
  inference(resolution,[],[f3692,f142])).
fof(f3692,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X14)))))) | p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f94,f116])).
fof(f9539,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (~p(s(bool,'const.bool.IN_2'(s(bool,X88),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X89)))))) | p(s(bool,happ(s(fun(bool,bool),X90),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X88) | ~p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X92) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X90),s('type.list.list'(bool),X89))) | s(bool,X91) != s(bool,X93)) )),
  inference(resolution,[],[f3692,f105])).
fof(f9537,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (~p(s(bool,'const.bool.IN_2'(s(bool,X76),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X77)))))) | p(s(bool,happ(s(fun(bool,bool),X78),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X76) | p(s(bool,X79)) | ~p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X78),s('type.list.list'(bool),X77)))) )),
  inference(resolution,[],[f3692,f104])).
fof(f9535,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X64),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X65)))))) | p(s(bool,happ(s(fun(bool,bool),X66),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,X69) | s(bool,X68) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X66),s('type.list.list'(bool),X65)))) )),
  inference(resolution,[],[f3692,f103])).
fof(f9534,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,X58),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X59)))))) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X58) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X60),s('type.list.list'(bool),X59))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f3692,f101])).
fof(f9532,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X47)))))) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X46) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X48),s('type.list.list'(bool),X47)))) )),
  inference(resolution,[],[f3692,f99])).
fof(f9475,plain,(
  ( ! [X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(X20,X22),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X23))))) | ~p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(X20,sK7(X23,X21,X20)),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X23))))))) ) | $spl272),
  inference(cnf_transformation,[],[f9475_D])).
fof(f9475_D,plain,(
  ( ! [X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(X20,X22),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X23))))) | ~p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(X20,sK7(X23,X21,X20)),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X23))))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f9432,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (p(s(bool,happ(s(fun(X100,bool),X101),s(X100,X102)))) | p(s(bool,'const.bool.IN_2'(s(X100,sK7(X103,X101,X100)),s(fun(X100,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X100),X103)))))) | p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X105) != s(bool,'const.bool.IN_2'(s(X100,X102),s(fun(X100,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X100),X103)))))) )),
  inference(resolution,[],[f3667,f142])).
fof(f3667,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X3,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) )),
  inference(resolution,[],[f94,f95])).
fof(f9425,plain,(
  ( ! [X70,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X69,bool),X70),s(X69,X71)))) | p(s(bool,'const.bool.IN_2'(s(X69,sK7(X72,X70,X69)),s(fun(X69,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X69),X72)))))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,'const.bool.IN_2'(s(X69,X71),s(fun(X69,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X69),X72))))) | s(bool,X73) != s(bool,X75)) )),
  inference(resolution,[],[f3667,f105])).
fof(f9423,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (p(s(bool,happ(s(fun(X55,bool),X56),s(X55,X57)))) | p(s(bool,'const.bool.IN_2'(s(X55,sK7(X58,X56,X55)),s(fun(X55,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X55),X58)))))) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.bool.IN_2'(s(X55,X57),s(fun(X55,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X55),X58)))))) )),
  inference(resolution,[],[f3667,f104])).
fof(f9421,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (p(s(bool,happ(s(fun(X41,bool),X42),s(X41,X43)))) | p(s(bool,'const.bool.IN_2'(s(X41,sK7(X44,X42,X41)),s(fun(X41,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X41),X44)))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(X41,X43),s(fun(X41,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X41),X44)))))) )),
  inference(resolution,[],[f3667,f103])).
fof(f9420,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X34,bool),X35),s(X34,X36)))) | p(s(bool,'const.bool.IN_2'(s(X34,sK7(X37,X35,X34)),s(fun(X34,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X34),X37)))))) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X34,X36),s(fun(X34,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X34),X37))))) | s(bool,X38) != s(bool,X40)) )),
  inference(resolution,[],[f3667,f101])).
fof(f9418,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | p(s(bool,'const.bool.IN_2'(s(X20,sK7(X23,X21,X20)),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X23)))))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(X20,X22),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X23)))))) )),
  inference(resolution,[],[f3667,f99])).
fof(f9408,plain,(
  ( ! [X10,X8,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X8),X9),s('type.bvlSem.v',X10)))) | s(bool,X13) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X10),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X11))))) | ~p(s(bool,X13)) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9))),s('type.list.list'('type.bvlSem.v'),X11))))) ) | $spl270),
  inference(cnf_transformation,[],[f9408_D])).
fof(f9408_D,plain,(
  ( ! [X10,X8,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X8),X9),s('type.bvlSem.v',X10)))) | s(bool,X13) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X10),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X11))))) | ~p(s(bool,X13)) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9))),s('type.list.list'('type.bvlSem.v'),X11))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f9395,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X88),X89),s('type.bvlSem.v',X90)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X88),X89))),s('type.list.list'('type.bvlSem.v'),X91)))) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X90),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X91)))))) )),
  inference(resolution,[],[f3709,f142])).
fof(f3709,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X3)))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))),s('type.list.list'('type.bvlSem.v'),X3))))) )),
  inference(superposition,[],[f94,f97])).
fof(f9388,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X57),X58),s('type.bvlSem.v',X59)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X57),X58))),s('type.list.list'('type.bvlSem.v'),X60)))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X59),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X60))))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f3709,f105])).
fof(f9386,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X43),X44),s('type.bvlSem.v',X45)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X43),X44))),s('type.list.list'('type.bvlSem.v'),X46)))) | p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X45),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X46)))))) )),
  inference(resolution,[],[f3709,f104])).
fof(f9384,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X29),X30),s('type.bvlSem.v',X31)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X29),X30))),s('type.list.list'('type.bvlSem.v'),X32)))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X31),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X32)))))) )),
  inference(resolution,[],[f3709,f103])).
fof(f9383,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X22),X23),s('type.bvlSem.v',X24)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X22),X23))),s('type.list.list'('type.bvlSem.v'),X25)))) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X24),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X25))))) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f3709,f101])).
fof(f9381,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X8),X9),s('type.bvlSem.v',X10)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X8),X9))),s('type.list.list'('type.bvlSem.v'),X11)))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X10),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X11)))))) )),
  inference(resolution,[],[f3709,f99])).
fof(f9366,plain,(
  ( ! [X24,X27,X25] : (p(s(bool,sK7(X24,X25,bool))) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24))) | p(s(bool,X27)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X24)))))) ) | $spl268),
  inference(cnf_transformation,[],[f9366_D])).
fof(f9366_D,plain,(
  ( ! [X24,X27,X25] : (p(s(bool,sK7(X24,X25,bool))) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24))) | p(s(bool,X27)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X24)))))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f9352,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,sK7(X76,X77,bool))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X76))))) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X79) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X77),s('type.list.list'(bool),X76)))) )),
  inference(resolution,[],[f5476,f142])).
fof(f5476,plain,(
  ( ! [X47,X48] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X47),s('type.list.list'(bool),X48)))) | p(s(bool,sK7(X48,X47,bool))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X48)))))) )),
  inference(resolution,[],[f3528,f113])).
fof(f9345,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,sK7(X59,X60,bool))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X59))))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X60),s('type.list.list'(bool),X59))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f5476,f105])).
fof(f9342,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,sK7(X44,X45,bool))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X44))))) | p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X45),s('type.list.list'(bool),X44)))) )),
  inference(resolution,[],[f5476,f104])).
fof(f9341,plain,(
  ( ! [X39,X43,X41,X42,X40] : (p(s(bool,sK7(X39,X40,bool))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X39))))) | p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X40),s('type.list.list'(bool),X39)))) )),
  inference(resolution,[],[f5476,f103])).
fof(f9340,plain,(
  ( ! [X37,X35,X38,X36,X34] : (p(s(bool,sK7(X34,X35,bool))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X34))))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34))) | s(bool,X37) != s(bool,X38)) )),
  inference(resolution,[],[f5476,f101])).
fof(f9338,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,sK7(X24,X25,bool))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X24))))) | p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24)))) )),
  inference(resolution,[],[f5476,f99])).
fof(f9293,plain,(
  ( ! [X4,X7,X5] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | s(bool,X7) != s(bool,sK7(X5,X4,bool)) | ~p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X5))))))) ) | $spl266),
  inference(cnf_transformation,[],[f9293_D])).
fof(f9293_D,plain,(
  ( ! [X4,X7,X5] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | s(bool,X7) != s(bool,sK7(X5,X4,bool)) | ~p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X5))))))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f9280,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X56),s('type.list.list'(bool),X57))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X57)))))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,sK7(X57,X56,bool))) )),
  inference(resolution,[],[f5398,f142])).
fof(f5398,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK7(X4,X5,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))))) )),
  inference(superposition,[],[f3489,f113])).
fof(f9273,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X40))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X40)))))) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,sK7(X40,X39,bool)) | s(bool,X41) != s(bool,X43)) )),
  inference(resolution,[],[f5398,f105])).
fof(f9271,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X29),s('type.list.list'(bool),X30))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X30)))))) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,sK7(X30,X29,bool))) )),
  inference(resolution,[],[f5398,f104])).
fof(f9269,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X19),s('type.list.list'(bool),X20))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X20)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,sK7(X20,X19,bool))) )),
  inference(resolution,[],[f5398,f103])).
fof(f9268,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X14),s('type.list.list'(bool),X15))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X15)))))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,sK7(X15,X14,bool)) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f5398,f101])).
fof(f9266,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X5)))))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,sK7(X5,X4,bool))) )),
  inference(resolution,[],[f5398,f99])).
fof(f9255,plain,(
  ( ! [X19,X17,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X16)))))) | s(bool,X19) != s(bool,sK7(X16,X17,bool)) | p(s(bool,X19)) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X17),s('type.list.list'(bool),X16)))) ) | $spl264),
  inference(cnf_transformation,[],[f9255_D])).
fof(f9255_D,plain,(
  ( ! [X19,X17,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X16)))))) | s(bool,X19) != s(bool,sK7(X16,X17,bool)) | p(s(bool,X19)) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X17),s('type.list.list'(bool),X16)))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f9218,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X68)))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X69),s('type.list.list'(bool),X68))) | p(s(bool,X70)) | p(s(bool,X71)) | s(bool,X71) != s(bool,sK7(X68,X69,bool))) )),
  inference(resolution,[],[f5396,f142])).
fof(f9211,plain,(
  ( ! [X54,X52,X55,X53,X51] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X51)))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X51))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,sK7(X51,X52,bool)) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f5396,f105])).
fof(f9208,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X36)))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X37),s('type.list.list'(bool),X36))) | p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) != s(bool,sK7(X36,X37,bool))) )),
  inference(resolution,[],[f5396,f104])).
fof(f9207,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X31)))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X32),s('type.list.list'(bool),X31))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,sK7(X31,X32,bool))) )),
  inference(resolution,[],[f5396,f103])).
fof(f9206,plain,(
  ( ! [X30,X28,X26,X29,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X26)))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X27),s('type.list.list'(bool),X26))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,sK7(X26,X27,bool)) | s(bool,X29) != s(bool,X30)) )),
  inference(resolution,[],[f5396,f101])).
fof(f9204,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X16)))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X17),s('type.list.list'(bool),X16))) | p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,sK7(X16,X17,bool))) )),
  inference(resolution,[],[f5396,f99])).
fof(f9249,plain,(
  ( ! [X10,X13,X11] : (p(s(bool,sK7(X10,X11,bool))) | p(s(bool,X13)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10)))) ) | $spl262),
  inference(cnf_transformation,[],[f9249_D])).
fof(f9249_D,plain,(
  ( ! [X10,X13,X11] : (p(s(bool,sK7(X10,X11,bool))) | p(s(bool,X13)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f9189,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,sK7(X62,X63,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X63),s('type.list.list'(bool),X62))) | p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X65) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X62)))))) )),
  inference(resolution,[],[f5396,f142])).
fof(f9182,plain,(
  ( ! [X47,X45,X48,X46,X49] : (p(s(bool,sK7(X45,X46,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X46),s('type.list.list'(bool),X45))) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X45))))) | s(bool,X48) != s(bool,X49)) )),
  inference(resolution,[],[f5396,f105])).
fof(f9179,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,sK7(X30,X31,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X31),s('type.list.list'(bool),X30))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X30)))))) )),
  inference(resolution,[],[f5396,f104])).
fof(f9178,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,sK7(X25,X26,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X26),s('type.list.list'(bool),X25))) | p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X25)))))) )),
  inference(resolution,[],[f5396,f103])).
fof(f9177,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,sK7(X20,X21,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X21),s('type.list.list'(bool),X20))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X20))))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f5396,f101])).
fof(f9175,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,sK7(X10,X11,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10)))))) )),
  inference(resolution,[],[f5396,f99])).
fof(f9052,plain,(
  ( ! [X6,X4,X8,X5] : (s(bool,X4) != s(bool,sK7(X5,X6,bool)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0'))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X6),s('type.list.list'(bool),X5)))) | ~p(s(bool,X4))) ) | $spl258),
  inference(cnf_transformation,[],[f9052_D])).
fof(f9052_D,plain,(
  ( ! [X6,X4,X8,X5] : (s(bool,X4) != s(bool,sK7(X5,X6,bool)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0'))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X6),s('type.list.list'(bool),X5)))) | ~p(s(bool,X4))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f9159,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X69),s('type.list.list'(bool),X70)))) | ~p(s(bool,X71)) | s(bool,X71) != s(bool,happ(s(fun(bool,bool),X69),s(bool,'T_0'))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,sK7(X70,X69,bool))) ) | $spl88),
  inference(resolution,[],[f4708,f142])).
fof(f4708,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,sK7(X6,X7,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X7),s('type.list.list'(bool),X6)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0')))) ) | $spl88),
  inference(superposition,[],[f3315,f113])).
fof(f9152,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X45),s('type.list.list'(bool),X46)))) | ~p(s(bool,X47)) | s(bool,X47) != s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0'))) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,sK7(X46,X45,bool)) | s(bool,X48) != s(bool,X50)) ) | $spl88),
  inference(resolution,[],[f4708,f105])).
fof(f9150,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X33),s('type.list.list'(bool),X34)))) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(bool,bool),X33),s(bool,'T_0'))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,sK7(X34,X33,bool))) ) | $spl88),
  inference(resolution,[],[f4708,f104])).
fof(f9148,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X21),s('type.list.list'(bool),X22)))) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,happ(s(fun(bool,bool),X21),s(bool,'T_0'))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,sK7(X22,X21,bool))) ) | $spl88),
  inference(resolution,[],[f4708,f103])).
fof(f9147,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X15),s('type.list.list'(bool),X16)))) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,happ(s(fun(bool,bool),X15),s(bool,'T_0'))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,sK7(X16,X15,bool)) | s(bool,X18) != s(bool,X20)) ) | $spl88),
  inference(resolution,[],[f4708,f101])).
fof(f9145,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | ~p(s(bool,X5)) | s(bool,X5) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,sK7(X4,X3,bool))) ) | $spl88),
  inference(resolution,[],[f4708,f99])).
fof(f9081,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,X3) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | s(bool,X7) != s(bool,sK7(X5,X4,bool)) | ~p(s(bool,X7)) | ~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | p(s(bool,X3))) ) | $spl260),
  inference(cnf_transformation,[],[f9081_D])).
fof(f9081_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,X3) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | s(bool,X7) != s(bool,sK7(X5,X4,bool)) | ~p(s(bool,X7)) | ~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | p(s(bool,X3))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f9069,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X70),s('type.list.list'(bool),X71))) | p(s(bool,X69)) | ~p(s(bool,happ(s(fun(bool,bool),X70),s(bool,'T_0')))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,sK7(X71,X70,bool))) ) | $spl86),
  inference(resolution,[],[f4649,f142])).
fof(f4649,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,sK7(X6,X7,bool))) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X7),s('type.list.list'(bool),X6))) | p(s(bool,X8)) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0'))))) ) | $spl86),
  inference(superposition,[],[f3314,f113])).
fof(f9062,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,X45) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X46),s('type.list.list'(bool),X47))) | p(s(bool,X45)) | ~p(s(bool,happ(s(fun(bool,bool),X46),s(bool,'T_0')))) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,sK7(X47,X46,bool)) | s(bool,X48) != s(bool,X50)) ) | $spl86),
  inference(resolution,[],[f4649,f105])).
fof(f9060,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,X33) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X34),s('type.list.list'(bool),X35))) | p(s(bool,X33)) | ~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0')))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,sK7(X35,X34,bool))) ) | $spl86),
  inference(resolution,[],[f4649,f104])).
fof(f9058,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X22),s('type.list.list'(bool),X23))) | p(s(bool,X21)) | ~p(s(bool,happ(s(fun(bool,bool),X22),s(bool,'T_0')))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,sK7(X23,X22,bool))) ) | $spl86),
  inference(resolution,[],[f4649,f103])).
fof(f9057,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X16),s('type.list.list'(bool),X17))) | p(s(bool,X15)) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'T_0')))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,sK7(X17,X16,bool)) | s(bool,X18) != s(bool,X20)) ) | $spl86),
  inference(resolution,[],[f4649,f101])).
fof(f9055,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,X3) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X5))) | p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,sK7(X5,X4,bool))) ) | $spl86),
  inference(resolution,[],[f4649,f99])).
fof(f9035,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X70) != s(bool,sK7(X71,X72,bool)) | ~p(s(bool,X70)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X72),s('type.list.list'(bool),X71)))) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(bool,bool),X72),s(bool,'T_0')))) ) | $spl112),
  inference(resolution,[],[f3981,f142])).
fof(f9028,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,X46) != s(bool,sK7(X47,X48,bool)) | ~p(s(bool,X46)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X48),s('type.list.list'(bool),X47)))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) | s(bool,X49) != s(bool,X51)) ) | $spl112),
  inference(resolution,[],[f3981,f105])).
fof(f9026,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,X34) != s(bool,sK7(X35,X36,bool)) | ~p(s(bool,X34)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X36),s('type.list.list'(bool),X35)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,happ(s(fun(bool,bool),X36),s(bool,'T_0')))) ) | $spl112),
  inference(resolution,[],[f3981,f104])).
fof(f9024,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,X22) != s(bool,sK7(X23,X24,bool)) | ~p(s(bool,X22)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X24),s('type.list.list'(bool),X23)))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0')))) ) | $spl112),
  inference(resolution,[],[f3981,f103])).
fof(f9023,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,X16) != s(bool,sK7(X17,X18,bool)) | ~p(s(bool,X16)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X17)))) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0'))) | s(bool,X19) != s(bool,X21)) ) | $spl112),
  inference(resolution,[],[f3981,f101])).
fof(f9021,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(bool,X4) != s(bool,sK7(X5,X6,bool)) | ~p(s(bool,X4)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X6),s('type.list.list'(bool),X5)))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0')))) ) | $spl112),
  inference(resolution,[],[f3981,f99])).
fof(f9012,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,sK7(X4,X5,bool))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0'))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) ) | $spl256),
  inference(cnf_transformation,[],[f9012_D])).
fof(f9012_D,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,sK7(X4,X5,bool))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0'))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f8992,plain,(
  ( ! [X70,X74,X72,X71,X73] : (p(s(bool,sK7(X70,X71,bool))) | s(bool,X72) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X71),s('type.list.list'(bool),X70))) | p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(bool,bool),X71),s(bool,'F_0')))) ) | $spl14),
  inference(resolution,[],[f3941,f142])).
fof(f3941,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'F_0')))) | p(s(bool,sK7(X60,X59,bool))) | s(bool,X61) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X59),s('type.list.list'(bool),X60))) | p(s(bool,X61))) ) | $spl14),
  inference(resolution,[],[f3302,f181])).
fof(f3302,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,sK7(X0,X1,bool)))) )),
  inference(superposition,[],[f96,f109])).
fof(f8985,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (p(s(bool,sK7(X46,X47,bool))) | s(bool,X48) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X47),s('type.list.list'(bool),X46))) | p(s(bool,X48)) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,happ(s(fun(bool,bool),X47),s(bool,'F_0'))) | s(bool,X49) != s(bool,X51)) ) | $spl14),
  inference(resolution,[],[f3941,f105])).
fof(f8983,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,sK7(X34,X35,bool))) | s(bool,X36) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34))) | p(s(bool,X36)) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,happ(s(fun(bool,bool),X35),s(bool,'F_0')))) ) | $spl14),
  inference(resolution,[],[f3941,f104])).
fof(f8981,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (p(s(bool,sK7(X22,X23,bool))) | s(bool,X24) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X22))) | p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,happ(s(fun(bool,bool),X23),s(bool,'F_0')))) ) | $spl14),
  inference(resolution,[],[f3941,f103])).
fof(f8980,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (p(s(bool,sK7(X16,X17,bool))) | s(bool,X18) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X17),s('type.list.list'(bool),X16))) | p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,X19) != s(bool,X21)) ) | $spl14),
  inference(resolution,[],[f3941,f101])).
fof(f8978,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (p(s(bool,sK7(X4,X5,bool))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4))) | p(s(bool,X6)) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0')))) ) | $spl14),
  inference(resolution,[],[f3941,f99])).
fof(f8958,plain,(
  ( ! [X159,X160] : (s(bool,X159) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X160) | s(bool,X159) != s(bool,X160)) ) | $spl254),
  inference(cnf_transformation,[],[f8958_D])).
fof(f8958_D,plain,(
  ( ! [X159,X160] : (s(bool,X159) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X160) | s(bool,X159) != s(bool,X160)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f8957,plain,(
  ( ! [X158,X159,X160] : (s(bool,X159) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X159) != s(bool,X160) | s(bool,'F_0') = s(bool,X158) | s(bool,'T_0') = s(bool,X160)) ) | ($spl3 | $spl6 | $spl14 | $spl40)),
  inference(subsumption_resolution,[],[f8846,f116])).
fof(f8846,plain,(
  ( ! [X158,X159,X160] : (s(bool,'T_0') != s(bool,X158) | s(bool,X159) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X159) != s(bool,X160) | s(bool,'F_0') = s(bool,X158) | s(bool,'T_0') = s(bool,X160)) ) | ($spl3 | $spl6 | $spl14 | $spl40)),
  inference(superposition,[],[f1953,f6930])).
fof(f8941,plain,(
  ( ! [X124] : (s(bool,X124) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X124)) ) | $spl252),
  inference(cnf_transformation,[],[f8941_D])).
fof(f8941_D,plain,(
  ( ! [X124] : (s(bool,X124) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X124)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f8940,plain,(
  ( ! [X123,X124] : (s(bool,'F_0') = s(bool,X123) | s(bool,X124) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl3 | $spl6 | $spl14)),
  inference(subsumption_resolution,[],[f8827,f116])).
fof(f8827,plain,(
  ( ! [X123,X124] : (s(bool,'T_0') != s(bool,X123) | s(bool,'F_0') = s(bool,X123) | s(bool,X124) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl3 | $spl6 | $spl14)),
  inference(superposition,[],[f1032,f6930])).
fof(f8748,plain,(
  ( ! [X6,X10,X8,X7,X5] : (s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) | s(bool,X10) != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK7(X8,X7,X6)))) | ~p(s(bool,X10)) | s(bool,'T_0') = s(bool,X5)) ) | $spl250),
  inference(cnf_transformation,[],[f8748_D])).
fof(f8748_D,plain,(
  ( ! [X6,X10,X8,X7,X5] : (s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) | s(bool,X10) != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK7(X8,X7,X6)))) | ~p(s(bool,X10)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f8726,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X87),s('type.list.list'(X86),X88))) | s(bool,'T_0') = s(bool,X85) | p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,happ(s(fun(X86,bool),X87),s(X86,sK7(X88,X87,X86))))) ) | $spl14),
  inference(resolution,[],[f1126,f142])).
fof(f1126,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,happ(s(fun(X80,bool),X81),s(X80,sK7(X82,X81,X80))))) | s(bool,X79) != s(bool,'const.list.EVERY_2'(s(fun(X80,bool),X81),s('type.list.list'(X80),X82))) | s(bool,'T_0') = s(bool,X79)) ) | $spl14),
  inference(resolution,[],[f1102,f96])).
fof(f8719,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(X55,bool),X56),s('type.list.list'(X55),X57))) | s(bool,'T_0') = s(bool,X54) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,happ(s(fun(X55,bool),X56),s(X55,sK7(X57,X56,X55)))) | s(bool,X58) != s(bool,X60)) ) | $spl14),
  inference(resolution,[],[f1126,f105])).
fof(f8717,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(X41,bool),X42),s('type.list.list'(X41),X43))) | s(bool,'T_0') = s(bool,X40) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,happ(s(fun(X41,bool),X42),s(X41,sK7(X43,X42,X41))))) ) | $spl14),
  inference(resolution,[],[f1126,f104])).
fof(f8715,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(X27,bool),X28),s('type.list.list'(X27),X29))) | s(bool,'T_0') = s(bool,X26) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(X27,bool),X28),s(X27,sK7(X29,X28,X27))))) ) | $spl14),
  inference(resolution,[],[f1126,f103])).
fof(f8714,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(X20,bool),X21),s('type.list.list'(X20),X22))) | s(bool,'T_0') = s(bool,X19) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK7(X22,X21,X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl14),
  inference(resolution,[],[f1126,f101])).
fof(f8712,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK7(X8,X7,X6))))) ) | $spl14),
  inference(resolution,[],[f1126,f99])).
fof(f8674,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,sK7(X76,X77,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X77),s(bool,'F_0'))) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X79) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X77),s('type.list.list'(bool),X76)))) )),
  inference(resolution,[],[f7961,f142])).
fof(f7961,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9)))) | s(bool,'T_0') = s(bool,sK7(X9,X8,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) )),
  inference(trivial_inequality_removal,[],[f7959])).
fof(f7959,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9)))) | s(bool,'T_0') = s(bool,sK7(X9,X8,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) )),
  inference(superposition,[],[f3824,f116])).
fof(f3824,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0'))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10)))) | s(bool,'T_0') = s(bool,sK7(X10,X11,bool))) )),
  inference(superposition,[],[f3284,f116])).
fof(f8667,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,'T_0') = s(bool,sK7(X59,X60,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X60),s(bool,'F_0'))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X60),s('type.list.list'(bool),X59))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f7961,f105])).
fof(f8664,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'T_0') = s(bool,sK7(X44,X45,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X45),s(bool,'F_0'))) | p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X45),s('type.list.list'(bool),X44)))) )),
  inference(resolution,[],[f7961,f104])).
fof(f8663,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,sK7(X39,X40,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X40),s(bool,'F_0'))) | p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X40),s('type.list.list'(bool),X39)))) )),
  inference(resolution,[],[f7961,f103])).
fof(f8662,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,sK7(X34,X35,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X35),s(bool,'F_0'))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34))) | s(bool,X37) != s(bool,X38)) )),
  inference(resolution,[],[f7961,f101])).
fof(f8660,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,sK7(X24,X25,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0'))) | p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24)))) )),
  inference(resolution,[],[f7961,f99])).
fof(f8634,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') = s(bool,sK7(X76,X77,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X77),s(bool,'T_0'))) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X79) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X77),s('type.list.list'(bool),X76)))) )),
  inference(resolution,[],[f7943,f142])).
fof(f7943,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9)))) | s(bool,'F_0') = s(bool,sK7(X9,X8,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) )),
  inference(trivial_inequality_removal,[],[f7941])).
fof(f7941,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9)))) | s(bool,'F_0') = s(bool,sK7(X9,X8,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) )),
  inference(superposition,[],[f3823,f116])).
fof(f3823,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0'))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X9),s('type.list.list'(bool),X8)))) | s(bool,'F_0') = s(bool,sK7(X8,X9,bool))) )),
  inference(superposition,[],[f3284,f116])).
fof(f8627,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,'F_0') = s(bool,sK7(X59,X60,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0'))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X60),s('type.list.list'(bool),X59))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f7943,f105])).
fof(f8624,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'F_0') = s(bool,sK7(X44,X45,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0'))) | p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X45),s('type.list.list'(bool),X44)))) )),
  inference(resolution,[],[f7943,f104])).
fof(f8623,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,sK7(X39,X40,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X40),s(bool,'T_0'))) | p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X40),s('type.list.list'(bool),X39)))) )),
  inference(resolution,[],[f7943,f103])).
fof(f8622,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'F_0') = s(bool,sK7(X34,X35,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0'))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34))) | s(bool,X37) != s(bool,X38)) )),
  inference(resolution,[],[f7943,f101])).
fof(f8620,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,sK7(X24,X25,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0'))) | p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X24)))) )),
  inference(resolution,[],[f7943,f99])).
fof(f8608,plain,(
  ( ! [X12,X10,X8,X9] : (p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10))))) | p(s(bool,X9)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X10))))) ) | $spl244),
  inference(cnf_transformation,[],[f8608_D])).
fof(f8608_D,plain,(
  ( ! [X12,X10,X8,X9] : (p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10))))) | p(s(bool,X9)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X10))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f8588,plain,(
  ( ! [X78,X76,X74,X77,X75] : (p(s(bool,happ(s(fun(bool,bool),X74),s(bool,X75)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X74),s('type.list.list'(bool),X76)))) | p(s(bool,X75)) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X76)))))) )),
  inference(resolution,[],[f3698,f142])).
fof(f3698,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X32)))))) | p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X30)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X31),s('type.list.list'(bool),X32)))) | p(s(bool,X30))) )),
  inference(superposition,[],[f94,f111])).
fof(f8581,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(bool,bool),X50),s(bool,X51)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X50),s('type.list.list'(bool),X52)))) | p(s(bool,X51)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X52))))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f3698,f105])).
fof(f8579,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),X38),s(bool,X39)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X40)))) | p(s(bool,X39)) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X40)))))) )),
  inference(resolution,[],[f3698,f104])).
fof(f8577,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X27)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X26),s('type.list.list'(bool),X28)))) | p(s(bool,X27)) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X28)))))) )),
  inference(resolution,[],[f3698,f103])).
fof(f8576,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X21)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X20),s('type.list.list'(bool),X22)))) | p(s(bool,X21)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X22))))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f3698,f101])).
fof(f8574,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X10)))) | p(s(bool,X9)) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10)))))) )),
  inference(resolution,[],[f3698,f99])).
fof(f8546,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X6))))) | ~p(s(bool,X5)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X6))))) ) | $spl242),
  inference(cnf_transformation,[],[f8546_D])).
fof(f8546_D,plain,(
  ( ! [X6,X4,X8,X5] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X6))))) | ~p(s(bool,X5)) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X6))))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f8529,plain,(
  ( ! [X70,X74,X72,X71,X73] : (p(s(bool,happ(s(fun(bool,bool),X70),s(bool,X71)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X70),s('type.list.list'(bool),X72)))) | ~p(s(bool,X71)) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X72)))))) )),
  inference(resolution,[],[f3695,f142])).
fof(f3695,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X23)))))) | p(s(bool,happ(s(fun(bool,bool),X22),s(bool,X21)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X22),s('type.list.list'(bool),X23)))) | ~p(s(bool,X21))) )),
  inference(superposition,[],[f94,f115])).
fof(f8522,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X47)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X46),s('type.list.list'(bool),X48)))) | ~p(s(bool,X47)) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X48))))) | s(bool,X49) != s(bool,X51)) )),
  inference(resolution,[],[f3695,f105])).
fof(f8520,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X35)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X34),s('type.list.list'(bool),X36)))) | ~p(s(bool,X35)) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X36)))))) )),
  inference(resolution,[],[f3695,f104])).
fof(f8518,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (p(s(bool,happ(s(fun(bool,bool),X22),s(bool,X23)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X22),s('type.list.list'(bool),X24)))) | ~p(s(bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X24)))))) )),
  inference(resolution,[],[f3695,f103])).
fof(f8517,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X17)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X16),s('type.list.list'(bool),X18)))) | ~p(s(bool,X17)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X18))))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f3695,f101])).
fof(f8515,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X6)))) | ~p(s(bool,X5)) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X6)))))) )),
  inference(resolution,[],[f3695,f99])).
fof(f8492,plain,(
  ( ! [X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X2),s('type.list.list'(bool),X3))) | s(bool,X5) != s(bool,sK7(X3,X2,bool)) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) ) | $spl240),
  inference(cnf_transformation,[],[f8492_D])).
fof(f8492_D,plain,(
  ( ! [X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X2),s('type.list.list'(bool),X3))) | s(bool,X5) != s(bool,sK7(X3,X2,bool)) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f8480,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X54),s('type.list.list'(bool),X55))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X54),s(bool,'T_0'))) | p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X57) != s(bool,sK7(X55,X54,bool))) )),
  inference(resolution,[],[f5138,f142])).
fof(f5138,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK7(X4,X5,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0')))) )),
  inference(superposition,[],[f3599,f113])).
fof(f8473,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X37),s('type.list.list'(bool),X38))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X37),s(bool,'T_0'))) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,sK7(X38,X37,bool)) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f5138,f105])).
fof(f8471,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X27),s('type.list.list'(bool),X28))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X27),s(bool,'T_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,sK7(X28,X27,bool))) )),
  inference(resolution,[],[f5138,f104])).
fof(f8469,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X17),s('type.list.list'(bool),X18))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,sK7(X18,X17,bool))) )),
  inference(resolution,[],[f5138,f103])).
fof(f8468,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X12),s('type.list.list'(bool),X13))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X12),s(bool,'T_0'))) | ~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X15) != s(bool,sK7(X13,X12,bool)) | s(bool,X14) != s(bool,X16)) )),
  inference(resolution,[],[f5138,f101])).
fof(f8466,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X2),s('type.list.list'(bool),X3))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0'))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,sK7(X3,X2,bool))) )),
  inference(resolution,[],[f5138,f99])).
fof(f8461,plain,(
  ( ! [X8,X11,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9))) | s(bool,X11) != s(bool,sK7(X9,X8,bool)) | p(s(bool,X11)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) ) | $spl238),
  inference(cnf_transformation,[],[f8461_D])).
fof(f8461_D,plain,(
  ( ! [X8,X11,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9))) | s(bool,X11) != s(bool,sK7(X9,X8,bool)) | p(s(bool,X11)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f8443,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X60),s('type.list.list'(bool),X61))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X60),s(bool,'F_0'))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X63) != s(bool,sK7(X61,X60,bool))) )),
  inference(resolution,[],[f5136,f142])).
fof(f5136,plain,(
  ( ! [X0,X1] : (p(s(bool,sK7(X0,X1,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) )),
  inference(superposition,[],[f3599,f109])).
fof(f8436,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X43),s('type.list.list'(bool),X44))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X43),s(bool,'F_0'))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,sK7(X44,X43,bool)) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f5136,f105])).
fof(f8433,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X28),s('type.list.list'(bool),X29))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X28),s(bool,'F_0'))) | p(s(bool,X30)) | ~p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) != s(bool,sK7(X29,X28,bool))) )),
  inference(resolution,[],[f5136,f104])).
fof(f8432,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X24))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X23),s(bool,'F_0'))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,sK7(X24,X23,bool))) )),
  inference(resolution,[],[f5136,f103])).
fof(f8431,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,'F_0'))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,sK7(X19,X18,bool)) | s(bool,X21) != s(bool,X22)) )),
  inference(resolution,[],[f5136,f101])).
fof(f8429,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,sK7(X9,X8,bool))) )),
  inference(resolution,[],[f5136,f99])).
fof(f8404,plain,(
  ( ! [X39,X37,X41,X38] : (~p(s(bool,'const.bool.IN_2'(s(bool,X37),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X38)))))) | ~p(s(bool,X37)) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X38))) | ~p(s(bool,X41))) ) | $spl236),
  inference(cnf_transformation,[],[f8404_D])).
fof(f8404_D,plain,(
  ( ! [X39,X37,X41,X38] : (~p(s(bool,'const.bool.IN_2'(s(bool,X37),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X38)))))) | ~p(s(bool,X37)) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X38))) | ~p(s(bool,X41))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f8393,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,'const.bool.IN_2'(s(bool,X103),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X104)))))) | p(s(bool,happ(s(fun(bool,bool),X105),s(bool,'T_0')))) | ~p(s(bool,X103)) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X107) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X105),s('type.list.list'(bool),X104)))) )),
  inference(resolution,[],[f3690,f142])).
fof(f3690,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X7),s('type.list.list'(bool),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8)))))) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0')))) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f94,f113])).
fof(f8386,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,X79),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X80)))))) | p(s(bool,happ(s(fun(bool,bool),X81),s(bool,'T_0')))) | ~p(s(bool,X79)) | ~p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X83) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X81),s('type.list.list'(bool),X80))) | s(bool,X82) != s(bool,X84)) )),
  inference(resolution,[],[f3690,f105])).
fof(f8384,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X68)))))) | p(s(bool,happ(s(fun(bool,bool),X69),s(bool,'T_0')))) | ~p(s(bool,X67)) | p(s(bool,X70)) | ~p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X71) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X69),s('type.list.list'(bool),X68)))) )),
  inference(resolution,[],[f3690,f104])).
fof(f8382,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X56)))))) | p(s(bool,happ(s(fun(bool,bool),X57),s(bool,'T_0')))) | ~p(s(bool,X55)) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,X60) | s(bool,X59) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X57),s('type.list.list'(bool),X56)))) )),
  inference(resolution,[],[f3690,f103])).
fof(f8381,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X49),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X50)))))) | p(s(bool,happ(s(fun(bool,bool),X51),s(bool,'T_0')))) | ~p(s(bool,X49)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X51),s('type.list.list'(bool),X50))) | s(bool,X52) != s(bool,X54)) )),
  inference(resolution,[],[f3690,f101])).
fof(f8379,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,X37),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X38)))))) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | ~p(s(bool,X37)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X38)))) )),
  inference(resolution,[],[f3690,f99])).
fof(f8362,plain,(
  ( ! [X45,X43,X48,X46,X44] : (s(bool,'F_0') = s(bool,X43) | s(bool,X48) != s(bool,'const.list.EVERY_2'(s(fun(X44,bool),X45),s('type.list.list'(X44),X46))) | p(s(bool,X48)) | s(bool,X43) != s(bool,happ(s(fun(X44,bool),X45),s(X44,sK7(X46,X45,X44))))) ) | $spl234),
  inference(cnf_transformation,[],[f8362_D])).
fof(f8362_D,plain,(
  ( ! [X45,X43,X48,X46,X44] : (s(bool,'F_0') = s(bool,X43) | s(bool,X48) != s(bool,'const.list.EVERY_2'(s(fun(X44,bool),X45),s('type.list.list'(X44),X46))) | p(s(bool,X48)) | s(bool,X43) != s(bool,happ(s(fun(X44,bool),X45),s(X44,sK7(X46,X45,X44))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f8344,plain,(
  ( ! [X127,X125,X123,X128,X126,X124] : (s(bool,'F_0') = s(bool,X123) | s(bool,X123) != s(bool,happ(s(fun(X124,bool),X125),s(X124,sK7(X126,X125,X124)))) | p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X128) != s(bool,'const.list.EVERY_2'(s(fun(X124,bool),X125),s('type.list.list'(X124),X126)))) ) | $spl14),
  inference(resolution,[],[f845,f142])).
fof(f845,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,'const.list.EVERY_2'(s(fun(X97,bool),X98),s('type.list.list'(X97),X99)))) | s(bool,'F_0') = s(bool,X100) | s(bool,X100) != s(bool,happ(s(fun(X97,bool),X98),s(X97,sK7(X99,X98,X97))))) ) | $spl14),
  inference(resolution,[],[f770,f96])).
fof(f8337,plain,(
  ( ! [X94,X92,X97,X95,X93,X98,X96] : (s(bool,'F_0') = s(bool,X92) | s(bool,X92) != s(bool,happ(s(fun(X93,bool),X94),s(X93,sK7(X95,X94,X93)))) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,'const.list.EVERY_2'(s(fun(X93,bool),X94),s('type.list.list'(X93),X95))) | s(bool,X97) != s(bool,X98)) ) | $spl14),
  inference(resolution,[],[f845,f105])).
fof(f8334,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s(bool,'F_0') = s(bool,X71) | s(bool,X71) != s(bool,happ(s(fun(X72,bool),X73),s(X72,sK7(X74,X73,X72)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.list.EVERY_2'(s(fun(X72,bool),X73),s('type.list.list'(X72),X74)))) ) | $spl14),
  inference(resolution,[],[f845,f104])).
fof(f8333,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (s(bool,'F_0') = s(bool,X64) | s(bool,X64) != s(bool,happ(s(fun(X65,bool),X66),s(X65,sK7(X67,X66,X65)))) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,'const.list.EVERY_2'(s(fun(X65,bool),X66),s('type.list.list'(X65),X67)))) ) | $spl14),
  inference(resolution,[],[f845,f103])).
fof(f8332,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (s(bool,'F_0') = s(bool,X57) | s(bool,X57) != s(bool,happ(s(fun(X58,bool),X59),s(X58,sK7(X60,X59,X58)))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.list.EVERY_2'(s(fun(X58,bool),X59),s('type.list.list'(X58),X60))) | s(bool,X62) != s(bool,X63)) ) | $spl14),
  inference(resolution,[],[f845,f101])).
fof(f8330,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'F_0') = s(bool,X43) | s(bool,X43) != s(bool,happ(s(fun(X44,bool),X45),s(X44,sK7(X46,X45,X44)))) | p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.list.EVERY_2'(s(fun(X44,bool),X45),s('type.list.list'(X44),X46)))) ) | $spl14),
  inference(resolution,[],[f845,f99])).
fof(f8306,plain,(
  ( ! [X37,X35,X33,X34] : (~p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X34)))))) | p(s(bool,X33)) | p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'F_0')))) | s(bool,X37) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34))) | ~p(s(bool,X37))) ) | $spl232),
  inference(cnf_transformation,[],[f8306_D])).
fof(f8306_D,plain,(
  ( ! [X37,X35,X33,X34] : (~p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X34)))))) | p(s(bool,X33)) | p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'F_0')))) | s(bool,X37) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34))) | ~p(s(bool,X37))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f8295,plain,(
  ( ! [X103,X101,X99,X102,X100] : (~p(s(bool,'const.bool.IN_2'(s(bool,X99),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X100)))))) | p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'F_0')))) | p(s(bool,X99)) | p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X103) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X101),s('type.list.list'(bool),X100)))) )),
  inference(resolution,[],[f3688,f142])).
fof(f3688,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X2)))))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,X0))) )),
  inference(superposition,[],[f94,f109])).
fof(f8288,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (~p(s(bool,'const.bool.IN_2'(s(bool,X75),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X76)))))) | p(s(bool,happ(s(fun(bool,bool),X77),s(bool,'F_0')))) | p(s(bool,X75)) | ~p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X79) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X77),s('type.list.list'(bool),X76))) | s(bool,X78) != s(bool,X80)) )),
  inference(resolution,[],[f3688,f105])).
fof(f8286,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,X63),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X64)))))) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,'F_0')))) | p(s(bool,X63)) | p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X65),s('type.list.list'(bool),X64)))) )),
  inference(resolution,[],[f3688,f104])).
fof(f8284,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X52)))))) | p(s(bool,happ(s(fun(bool,bool),X53),s(bool,'F_0')))) | p(s(bool,X51)) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,X56) | s(bool,X55) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X53),s('type.list.list'(bool),X52)))) )),
  inference(resolution,[],[f3688,f103])).
fof(f8283,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X45),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X46)))))) | p(s(bool,happ(s(fun(bool,bool),X47),s(bool,'F_0')))) | p(s(bool,X45)) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X47),s('type.list.list'(bool),X46))) | s(bool,X48) != s(bool,X50)) )),
  inference(resolution,[],[f3688,f101])).
fof(f8281,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (~p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X34)))))) | p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'F_0')))) | p(s(bool,X33)) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X34)))) )),
  inference(resolution,[],[f3688,f99])).
fof(f8250,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X41),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X41,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f8191,f142])).
fof(f8191,plain,(
  ( ! [X22] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X22,'type.bvlSem.v'))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X22),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(trivial_inequality_removal,[],[f8076])).
fof(f8076,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X22,'type.bvlSem.v'))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X22),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(superposition,[],[f7204,f5385])).
fof(f7204,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X1))))) ) | $spl6),
  inference(forward_demodulation,[],[f7203,f97])).
fof(f7203,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X1)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f7192,f141])).
fof(f7192,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X1)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(superposition,[],[f3658,f2248])).
fof(f3658,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X68,bool),X69),s('type.list.list'(X68),X71)))) | p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X68,X70),s(fun(X68,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X68),X71)))))) )),
  inference(resolution,[],[f94,f112])).
fof(f8243,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X31),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X31,'type.bvlSem.v')))) | s(bool,X33) != s(bool,X34)) ) | $spl6),
  inference(resolution,[],[f8191,f105])).
fof(f8240,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X19),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X19,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f8191,f104])).
fof(f8239,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X15),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X15,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f8191,f103])).
fof(f8238,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X11),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X11,'type.bvlSem.v')))) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f8191,f101])).
fof(f8236,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X3,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f8191,f99])).
fof(f8228,plain,(
  ( ! [X453,X455,X457,X454,X458] : (s(bool,X457) != s(bool,'const.bool.IN_2'(s(X453,sK7(X454,X455,X453)),s(fun(X453,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X453),X454))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X453,bool),X455),s('type.list.list'(X453),X454))) | s(bool,'T_0') = s(bool,X458) | s(bool,X457) != s(bool,X458)) ) | $spl228),
  inference(cnf_transformation,[],[f8228_D])).
fof(f8228_D,plain,(
  ( ! [X453,X455,X457,X454,X458] : (s(bool,X457) != s(bool,'const.bool.IN_2'(s(X453,sK7(X454,X455,X453)),s(fun(X453,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X453),X454))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X453,bool),X455),s('type.list.list'(X453),X454))) | s(bool,'T_0') = s(bool,X458) | s(bool,X457) != s(bool,X458)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f8227,plain,(
  ( ! [X453,X455,X457,X454,X456,X458] : (s(bool,X457) != s(bool,'const.bool.IN_2'(s(X453,sK7(X454,X455,X453)),s(fun(X453,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X453),X454))))) | s(bool,X457) != s(bool,X458) | s(bool,'F_0') = s(bool,X456) | s(bool,'T_0') = s(bool,X458) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X453,bool),X455),s('type.list.list'(X453),X454)))) ) | ($spl14 | $spl40)),
  inference(subsumption_resolution,[],[f8165,f116])).
fof(f8165,plain,(
  ( ! [X453,X455,X457,X454,X456,X458] : (s(bool,'T_0') != s(bool,X456) | s(bool,X457) != s(bool,'const.bool.IN_2'(s(X453,sK7(X454,X455,X453)),s(fun(X453,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X453),X454))))) | s(bool,X457) != s(bool,X458) | s(bool,'F_0') = s(bool,X456) | s(bool,'T_0') = s(bool,X458) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X453,bool),X455),s('type.list.list'(X453),X454)))) ) | ($spl14 | $spl40)),
  inference(superposition,[],[f1953,f5385])).
fof(f8218,plain,(
  ( ! [X364,X362,X365,X361,X363] : (s(bool,'F_0') = s(bool,X364) | s(bool,X365) != s(bool,'const.bool.IN_2'(s(X361,sK7(X362,X363,X361)),s(fun(X361,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X361),X362))))) | s(bool,'T_0') = s(bool,X365) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X361,bool),X363),s('type.list.list'(X361),X362)))) ) | $spl14),
  inference(subsumption_resolution,[],[f8146,f116])).
fof(f8146,plain,(
  ( ! [X364,X362,X365,X361,X363] : (s(bool,'T_0') != s(bool,X364) | s(bool,'F_0') = s(bool,X364) | s(bool,X365) != s(bool,'const.bool.IN_2'(s(X361,sK7(X362,X363,X361)),s(fun(X361,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X361),X362))))) | s(bool,'T_0') = s(bool,X365) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X361,bool),X363),s('type.list.list'(X361),X362)))) ) | $spl14),
  inference(superposition,[],[f1032,f5385])).
fof(f5409,plain,(
  ( ! [X6,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X5,bool),X6),s('type.list.list'(X5),X7))) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(X5,sK7(X7,X6,X5)),s(fun(X5,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X5),X7))))) | p(s(bool,X9))) ) | $spl136),
  inference(cnf_transformation,[],[f5409_D])).
fof(f5409_D,plain,(
  ( ! [X6,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X5,bool),X6),s('type.list.list'(X5),X7))) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(X5,sK7(X7,X6,X5)),s(fun(X5,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X5),X7))))) | p(s(bool,X9))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f8214,plain,(
  ( ! [X325,X323,X324,X326,X322] : (p(s(bool,X325)) | p(s(bool,X326)) | s(bool,X326) != s(bool,'const.bool.IN_2'(s(X322,sK7(X323,X324,X322)),s(fun(X322,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X322),X323))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X322,bool),X324),s('type.list.list'(X322),X323)))) )),
  inference(subsumption_resolution,[],[f8137,f141])).
fof(f8137,plain,(
  ( ! [X325,X323,X324,X326,X322] : (~p(s(bool,'T_0')) | p(s(bool,X325)) | p(s(bool,X326)) | s(bool,X326) != s(bool,'const.bool.IN_2'(s(X322,sK7(X323,X324,X322)),s(fun(X322,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X322),X323))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X322,bool),X324),s('type.list.list'(X322),X323)))) )),
  inference(superposition,[],[f142,f5385])).
fof(f8212,plain,(
  ( ! [X231,X235,X233,X230,X234,X232] : (p(s(bool,X233)) | ~p(s(bool,X234)) | s(bool,X233) != s(bool,'const.bool.IN_2'(s(X230,sK7(X231,X232,X230)),s(fun(X230,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X230),X231))))) | s(bool,X234) != s(bool,X235) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X230,bool),X232),s('type.list.list'(X230),X231)))) )),
  inference(subsumption_resolution,[],[f8113,f136])).
fof(f8113,plain,(
  ( ! [X231,X235,X233,X230,X234,X232] : (~p(s(bool,'T_0')) | p(s(bool,X233)) | ~p(s(bool,X234)) | s(bool,X233) != s(bool,'const.bool.IN_2'(s(X230,sK7(X231,X232,X230)),s(fun(X230,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X230),X231))))) | s(bool,X234) != s(bool,X235) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X230,bool),X232),s('type.list.list'(X230),X231)))) )),
  inference(superposition,[],[f105,f5385])).
fof(f8210,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (p(s(bool,X185)) | ~p(s(bool,X186)) | p(s(bool,X187)) | s(bool,X187) != s(bool,'const.bool.IN_2'(s(X182,sK7(X183,X184,X182)),s(fun(X182,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X182),X183))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X182,bool),X184),s('type.list.list'(X182),X183)))) )),
  inference(subsumption_resolution,[],[f8105,f136])).
fof(f8105,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (~p(s(bool,'T_0')) | p(s(bool,X185)) | ~p(s(bool,X186)) | p(s(bool,X187)) | s(bool,X187) != s(bool,'const.bool.IN_2'(s(X182,sK7(X183,X184,X182)),s(fun(X182,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X182),X183))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X182,bool),X184),s('type.list.list'(X182),X183)))) )),
  inference(superposition,[],[f104,f5385])).
fof(f8208,plain,(
  ( ! [X156,X154,X152,X157,X155,X153] : (p(s(bool,X155)) | p(s(bool,X156)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,'const.bool.IN_2'(s(X152,sK7(X153,X154,X152)),s(fun(X152,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X152),X153))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X152,bool),X154),s('type.list.list'(X152),X153)))) )),
  inference(subsumption_resolution,[],[f8100,f141])).
fof(f8100,plain,(
  ( ! [X156,X154,X152,X157,X155,X153] : (~p(s(bool,'T_0')) | p(s(bool,X155)) | p(s(bool,X156)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,'const.bool.IN_2'(s(X152,sK7(X153,X154,X152)),s(fun(X152,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X152),X153))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X152,bool),X154),s('type.list.list'(X152),X153)))) )),
  inference(superposition,[],[f103,f5385])).
fof(f8205,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X113) != s(bool,'const.bool.IN_2'(s(X110,sK7(X111,X112,X110)),s(fun(X110,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X110),X111))))) | s(bool,X114) != s(bool,X115) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X110,bool),X112),s('type.list.list'(X110),X111)))) )),
  inference(subsumption_resolution,[],[f8093,f136])).
fof(f8093,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X113) != s(bool,'const.bool.IN_2'(s(X110,sK7(X111,X112,X110)),s(fun(X110,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X110),X111))))) | s(bool,X114) != s(bool,X115) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X110,bool),X112),s('type.list.list'(X110),X111)))) )),
  inference(superposition,[],[f101,f5385])).
fof(f8203,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,'const.bool.IN_2'(s(X68,sK7(X69,X70,X68)),s(fun(X68,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X68),X69))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X68,bool),X70),s('type.list.list'(X68),X69)))) )),
  inference(subsumption_resolution,[],[f8086,f141])).
fof(f8086,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,'const.bool.IN_2'(s(X68,sK7(X69,X70,X68)),s(fun(X68,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X68),X69))))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X68,bool),X70),s('type.list.list'(X68),X69)))) )),
  inference(superposition,[],[f99,f5385])).
fof(f8040,plain,(
  ( ! [X6,X8,X7,X11,X9] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(X6,sK7(X9,X7,X6)),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8))))) | ~p(s(bool,X11)) | p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X9))))) ) | $spl224),
  inference(cnf_transformation,[],[f8040_D])).
fof(f8040_D,plain,(
  ( ! [X6,X8,X7,X11,X9] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(X6,sK7(X9,X7,X6)),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8))))) | ~p(s(bool,X11)) | p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X9))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f8021,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X87),s('type.list.list'(X86),X88)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X87),s('type.list.list'(X86),X89)))) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X91) != s(bool,'const.bool.IN_2'(s(X86,sK7(X89,X87,X86)),s(fun(X86,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X86),X88)))))) )),
  inference(resolution,[],[f3626,f142])).
fof(f3626,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(X4,sK7(X5,X6,X4)),s(fun(X4,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X4),X7)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X6),s('type.list.list'(X4),X7)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X6),s('type.list.list'(X4),X5))))) )),
  inference(resolution,[],[f94,f96])).
fof(f8014,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X55,bool),X56),s('type.list.list'(X55),X57)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X55,bool),X56),s('type.list.list'(X55),X58)))) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.bool.IN_2'(s(X55,sK7(X58,X56,X55)),s(fun(X55,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X55),X57))))) | s(bool,X59) != s(bool,X61)) )),
  inference(resolution,[],[f3626,f105])).
fof(f8012,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X41,bool),X42),s('type.list.list'(X41),X43)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X41,bool),X42),s('type.list.list'(X41),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(X41,sK7(X44,X42,X41)),s(fun(X41,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X41),X43)))))) )),
  inference(resolution,[],[f3626,f104])).
fof(f8010,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X32] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X27,bool),X28),s('type.list.list'(X27),X29)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X27,bool),X28),s('type.list.list'(X27),X30)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(X27,sK7(X30,X28,X27)),s(fun(X27,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X27),X29)))))) )),
  inference(resolution,[],[f3626,f103])).
fof(f8009,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X20,bool),X21),s('type.list.list'(X20),X22)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X20,bool),X21),s('type.list.list'(X20),X23)))) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(X20,sK7(X23,X21,X20)),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X22))))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f3626,f101])).
fof(f8007,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X9)))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(X6,sK7(X9,X7,X6)),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) )),
  inference(resolution,[],[f3626,f99])).
fof(f7926,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2))),s('type.list.list'('type.bvlSem.v'),X3))) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X2),s('type.bvlSem.v',sK7(X3,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2)),'type.bvlSem.v'))))) ) | $spl94),
  inference(subsumption_resolution,[],[f7902,f113])).
fof(f7902,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2))),s('type.list.list'('type.bvlSem.v'),X3))) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X1),X2),s('type.bvlSem.v',sK7(X3,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X1),X2)),'type.bvlSem.v'))))) ) | $spl94),
  inference(superposition,[],[f3623,f3326])).
fof(f7874,plain,(
  ( ! [X6,X4,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool))) ) | $spl220),
  inference(cnf_transformation,[],[f7874_D])).
fof(f7874_D,plain,(
  ( ! [X6,X4,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f7854,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56))) | s(bool,'T_0') = s(bool,sK7(X56,X55,bool)) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,happ(s(fun(bool,bool),X55),s(bool,'F_0')))) )),
  inference(resolution,[],[f3615,f142])).
fof(f3615,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10))) | s(bool,'T_0') = s(bool,sK7(X10,X11,bool))) )),
  inference(superposition,[],[f3264,f116])).
fof(f7847,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39))) | s(bool,'T_0') = s(bool,sK7(X39,X38,bool)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(bool,bool),X38),s(bool,'F_0'))) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f3615,f105])).
fof(f7845,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X28),s('type.list.list'(bool),X29))) | s(bool,'T_0') = s(bool,sK7(X29,X28,bool)) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,happ(s(fun(bool,bool),X28),s(bool,'F_0')))) )),
  inference(resolution,[],[f3615,f104])).
fof(f7843,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19))) | s(bool,'T_0') = s(bool,sK7(X19,X18,bool)) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),X18),s(bool,'F_0')))) )),
  inference(resolution,[],[f3615,f103])).
fof(f7842,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14))) | s(bool,'T_0') = s(bool,sK7(X14,X13,bool)) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0'))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3615,f101])).
fof(f7840,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool)) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) )),
  inference(resolution,[],[f3615,f99])).
fof(f7837,plain,(
  ( ! [X6,X4,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) | $spl218),
  inference(cnf_transformation,[],[f7837_D])).
fof(f7837_D,plain,(
  ( ! [X6,X4,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f7820,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56))) | s(bool,'F_0') = s(bool,sK7(X56,X55,bool)) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,happ(s(fun(bool,bool),X55),s(bool,'T_0')))) )),
  inference(resolution,[],[f3614,f142])).
fof(f3614,plain,(
  ( ! [X8,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X9),s('type.list.list'(bool),X8))) | s(bool,'F_0') = s(bool,sK7(X8,X9,bool))) )),
  inference(superposition,[],[f3264,f116])).
fof(f7813,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39))) | s(bool,'F_0') = s(bool,sK7(X39,X38,bool)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(bool,bool),X38),s(bool,'T_0'))) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f3614,f105])).
fof(f7811,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X28),s('type.list.list'(bool),X29))) | s(bool,'F_0') = s(bool,sK7(X29,X28,bool)) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,happ(s(fun(bool,bool),X28),s(bool,'T_0')))) )),
  inference(resolution,[],[f3614,f104])).
fof(f7809,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19))) | s(bool,'F_0') = s(bool,sK7(X19,X18,bool)) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0')))) )),
  inference(resolution,[],[f3614,f103])).
fof(f7808,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14))) | s(bool,'F_0') = s(bool,sK7(X14,X13,bool)) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0'))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3614,f101])).
fof(f7806,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool)) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) )),
  inference(resolution,[],[f3614,f99])).
fof(f7592,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | $spl216),
  inference(cnf_transformation,[],[f7592_D])).
fof(f7592_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f7590,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f7563,f108])).
fof(f7563,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'F_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | p(s(bool,X2))) ) | ($spl3 | $spl14)),
  inference(superposition,[],[f3800,f2215])).
fof(f7538,plain,(
  ( ! [X7,X5] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X5),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X5),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl214),
  inference(cnf_transformation,[],[f7538_D])).
fof(f7538_D,plain,(
  ( ! [X7,X5] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X5),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X5),s('type.list.list'('type.bvlSem.v'),sK5)))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f7525,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X43),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X45) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X43),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(resolution,[],[f2660,f142])).
fof(f7518,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X33),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X33),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X34) != s(bool,X36)) ) | $spl3),
  inference(resolution,[],[f2660,f105])).
fof(f7516,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X25),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X25),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(resolution,[],[f2660,f104])).
fof(f7514,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X17),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X17),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(resolution,[],[f2660,f103])).
fof(f7513,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X13),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | ~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X13),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X14) != s(bool,X16)) ) | $spl3),
  inference(resolution,[],[f2660,f101])).
fof(f7511,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),X5),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X5),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(resolution,[],[f2660,f99])).
fof(f7461,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X0) != s(bool,X2)) ) | $spl212),
  inference(cnf_transformation,[],[f7461_D])).
fof(f7461_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f7397,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,X94) != s(bool,X95) | s(bool,X94) != s(bool,X96) | s(bool,X97) != s(bool,X98) | p(s(bool,X98)) | s(bool,X96) != s(bool,X97) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X95) != s(bool,X100)) ) | $spl60),
  inference(resolution,[],[f1924,f142])).
fof(f7390,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X59) != s(bool,X60) | p(s(bool,X60)) | s(bool,X58) != s(bool,X59) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X57) != s(bool,X62) | s(bool,X61) != s(bool,X63)) ) | $spl60),
  inference(resolution,[],[f1924,f105])).
fof(f7388,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,X40) != s(bool,X42) | s(bool,X43) != s(bool,X44) | p(s(bool,X44)) | s(bool,X42) != s(bool,X43) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X41) != s(bool,X46)) ) | $spl60),
  inference(resolution,[],[f1924,f104])).
fof(f7386,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X27) != s(bool,X28) | p(s(bool,X28)) | s(bool,X26) != s(bool,X27) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X25) != s(bool,X30)) ) | $spl60),
  inference(resolution,[],[f1924,f103])).
fof(f7385,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X18) != s(bool,X19) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X17) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl60),
  inference(resolution,[],[f1924,f101])).
fof(f7383,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X2) != s(bool,X3) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X1) != s(bool,X6)) ) | $spl60),
  inference(resolution,[],[f1924,f99])).
fof(f7365,plain,(
  ( ! [X19,X17,X22,X20,X18] : (p(s(bool,happ(s(fun(X17,bool),X18),s(X17,X19)))) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(X17,X19),s(fun(X17,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X17),X20))))) | ~p(s(bool,X22)) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X17,bool),X18),s('type.list.list'(X17),X20)))) ) | $spl210),
  inference(cnf_transformation,[],[f7365_D])).
fof(f7365_D,plain,(
  ( ! [X19,X17,X22,X20,X18] : (p(s(bool,happ(s(fun(X17,bool),X18),s(X17,X19)))) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(X17,X19),s(fun(X17,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X17),X20))))) | ~p(s(bool,X22)) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X17,bool),X18),s('type.list.list'(X17),X20)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f7326,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (p(s(bool,happ(s(fun(X97,bool),X98),s(X97,X99)))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X97,bool),X98),s('type.list.list'(X97),X100))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X102) != s(bool,'const.bool.IN_2'(s(X97,X99),s(fun(X97,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X97),X100)))))) )),
  inference(resolution,[],[f3677,f142])).
fof(f7319,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (p(s(bool,happ(s(fun(X66,bool),X67),s(X66,X68)))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X66,bool),X67),s('type.list.list'(X66),X69))) | ~p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X71) != s(bool,'const.bool.IN_2'(s(X66,X68),s(fun(X66,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X66),X69))))) | s(bool,X70) != s(bool,X72)) )),
  inference(resolution,[],[f3677,f105])).
fof(f7317,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X52,bool),X53),s('type.list.list'(X52),X55))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.bool.IN_2'(s(X52,X54),s(fun(X52,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X52),X55)))))) )),
  inference(resolution,[],[f3677,f104])).
fof(f7315,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X38,bool),X39),s('type.list.list'(X38),X41))) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,X43) != s(bool,'const.bool.IN_2'(s(X38,X40),s(fun(X38,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X38),X41)))))) )),
  inference(resolution,[],[f3677,f103])).
fof(f7314,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X31,bool),X32),s('type.list.list'(X31),X34))) | ~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(X31,X33),s(fun(X31,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X31),X34))))) | s(bool,X35) != s(bool,X37)) )),
  inference(resolution,[],[f3677,f101])).
fof(f7312,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18] : (p(s(bool,happ(s(fun(X17,bool),X18),s(X17,X19)))) | s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(X17,bool),X18),s('type.list.list'(X17),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(X17,X19),s(fun(X17,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X17),X20)))))) )),
  inference(resolution,[],[f3677,f99])).
fof(f7303,plain,(
  ( ! [X19,X17] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X17,'type.bvlSem.v')))) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X17),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl208),
  inference(cnf_transformation,[],[f7303_D])).
fof(f7303_D,plain,(
  ( ! [X19,X17] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X17,'type.bvlSem.v')))) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X17),s('type.list.list'('type.bvlSem.v'),sK5)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f7288,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X55,'type.bvlSem.v')))) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X57) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X55),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f6526,f142])).
fof(f6526,plain,(
  ( ! [X37] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X37),s('type.list.list'('type.bvlSem.v'),sK5)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X37,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f6512,f113])).
fof(f6512,plain,(
  ( ! [X1] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X1,'type.bvlSem.v'))))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X1),s('type.list.list'('type.bvlSem.v'),sK5))))) ) | $spl6),
  inference(forward_demodulation,[],[f6511,f97])).
fof(f6511,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,X1,'type.bvlSem.v'))))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X1),s('type.list.list'('type.bvlSem.v'),sK5))))) ) | $spl6),
  inference(subsumption_resolution,[],[f6500,f141])).
fof(f6500,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',sK7(sK5,X1,'type.bvlSem.v'))))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X1),s('type.list.list'('type.bvlSem.v'),sK5))))) ) | $spl6),
  inference(superposition,[],[f3647,f2248])).
fof(f3647,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X3,X0))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2))))) )),
  inference(resolution,[],[f94,f95])).
fof(f7281,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X45,'type.bvlSem.v')))) | p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X45),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X47) != s(bool,X48)) ) | $spl6),
  inference(resolution,[],[f6526,f105])).
fof(f7278,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X33,'type.bvlSem.v')))) | p(s(bool,X34)) | ~p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X33),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f6526,f104])).
fof(f7277,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X29,'type.bvlSem.v')))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X29),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f6526,f103])).
fof(f7276,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X25,'type.bvlSem.v')))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X25),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X27) != s(bool,X28)) ) | $spl6),
  inference(resolution,[],[f6526,f101])).
fof(f7274,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X17,'type.bvlSem.v')))) | p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X17),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f6526,f99])).
fof(f7253,plain,(
  ( ! [X19,X17] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X17,'type.bvlSem.v')))) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X17),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | $spl206),
  inference(cnf_transformation,[],[f7253_D])).
fof(f7253_D,plain,(
  ( ! [X19,X17] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X17,'type.bvlSem.v')))) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X17),s('type.list.list'('type.bvlSem.v'),sK6)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f7239,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X55,'type.bvlSem.v')))) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X57) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X55),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f4382,f142])).
fof(f4382,plain,(
  ( ! [X37] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X37),s('type.list.list'('type.bvlSem.v'),sK6)))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X37,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f1785,f113])).
fof(f1785,plain,(
  ( ! [X0] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X0,'type.bvlSem.v'))))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X0),s('type.list.list'('type.bvlSem.v'),sK6))))) )),
  inference(resolution,[],[f539,f95])).
fof(f7232,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X45,'type.bvlSem.v')))) | p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X45),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,X47) != s(bool,X48)) )),
  inference(resolution,[],[f4382,f105])).
fof(f7229,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X33,'type.bvlSem.v')))) | p(s(bool,X34)) | ~p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X33),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f4382,f104])).
fof(f7228,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X29,'type.bvlSem.v')))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X29),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f4382,f103])).
fof(f7227,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X25,'type.bvlSem.v')))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X25),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f4382,f101])).
fof(f7225,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X17,'type.bvlSem.v')))) | p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X17),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f4382,f99])).
fof(f7200,plain,(
  ( ! [X39,X37,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X39)))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X40))) | ~p(s(bool,X42)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X37,X39),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X40)))))) ) | $spl204),
  inference(cnf_transformation,[],[f7200_D])).
fof(f7200_D,plain,(
  ( ! [X39,X37,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X39)))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X40))) | ~p(s(bool,X42)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X37,X39),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X40)))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7186,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X117,X119),s(fun(X117,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X117),X120))))) | p(s(bool,X121)) | ~p(s(bool,X122)) | s(bool,X122) != s(bool,'const.list.EVERY_2'(s(fun(X117,bool),X118),s('type.list.list'(X117),X120)))) )),
  inference(resolution,[],[f3658,f142])).
fof(f7179,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (p(s(bool,happ(s(fun(X86,bool),X87),s(X86,X88)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X86,X88),s(fun(X86,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X86),X89))))) | ~p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X91) != s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X87),s('type.list.list'(X86),X89))) | s(bool,X90) != s(bool,X92)) )),
  inference(resolution,[],[f3658,f105])).
fof(f7177,plain,(
  ( ! [X78,X76,X74,X72,X77,X75,X73] : (p(s(bool,happ(s(fun(X72,bool),X73),s(X72,X74)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X72,X74),s(fun(X72,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X72),X75))))) | p(s(bool,X76)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,'const.list.EVERY_2'(s(fun(X72,bool),X73),s('type.list.list'(X72),X75)))) )),
  inference(resolution,[],[f3658,f104])).
fof(f7175,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (p(s(bool,happ(s(fun(X58,bool),X59),s(X58,X60)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X58,X60),s(fun(X58,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X58),X61))))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.list.EVERY_2'(s(fun(X58,bool),X59),s('type.list.list'(X58),X61)))) )),
  inference(resolution,[],[f3658,f103])).
fof(f7174,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (p(s(bool,happ(s(fun(X51,bool),X52),s(X51,X53)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X51,X53),s(fun(X51,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X51),X54))))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),X54))) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f3658,f101])).
fof(f7172,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X39)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X37,X39),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X40))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X40)))) )),
  inference(resolution,[],[f3658,f99])).
fof(f7134,plain,(
  ( ! [X43,X44,X42] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X42),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X42)))) ) | $spl6),
  inference(resolution,[],[f7043,f142])).
fof(f7043,plain,(
  ( ! [X1] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(forward_demodulation,[],[f7042,f97])).
fof(f7042,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f7031,f141])).
fof(f7031,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(superposition,[],[f3656,f2248])).
fof(f3656,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X60,bool),X61),s('type.list.list'(X60),X63)))) | p(s(bool,happ(s(fun(X60,bool),X61),s(X60,X62)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X60,X62),s(fun(X60,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X60),X63)))))) )),
  inference(resolution,[],[f94,f109])).
fof(f7127,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X32),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X32))) | s(bool,X34) != s(bool,X35)) ) | $spl6),
  inference(resolution,[],[f7043,f105])).
fof(f7124,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X20),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X20)))) ) | $spl6),
  inference(resolution,[],[f7043,f104])).
fof(f7123,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X16),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X16)))) ) | $spl6),
  inference(resolution,[],[f7043,f103])).
fof(f7122,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X12),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X12))) | s(bool,X14) != s(bool,X15)) ) | $spl6),
  inference(resolution,[],[f7043,f101])).
fof(f7120,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X4),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X4)))) ) | $spl6),
  inference(resolution,[],[f7043,f99])).
fof(f7116,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X2))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) | $spl200),
  inference(cnf_transformation,[],[f7116_D])).
fof(f7116_D,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X2))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f7105,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X40),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f6808,f142])).
fof(f6808,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X0),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) | s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X0)))) )),
  inference(forward_demodulation,[],[f6807,f97])).
fof(f6807,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X0),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X0)))) )),
  inference(subsumption_resolution,[],[f6797,f136])).
fof(f6797,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X0),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.bvlSem.v',X0)))) )),
  inference(superposition,[],[f3637,f523])).
fof(f3637,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X72,bool),X75),s('type.list.list'(X72),X74)))) | ~p(s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X72),X74)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),X75),s(X72,X73)))) )),
  inference(resolution,[],[f94,f113])).
fof(f7098,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X30))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X30),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f6808,f105])).
fof(f7096,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X22))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X22),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f6808,f104])).
fof(f7094,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X14),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f6808,f103])).
fof(f7093,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X10))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X10),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f6808,f101])).
fof(f7091,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f6808,f99])).
fof(f7061,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X41),s('type.list.list'('type.bvlSem.v'),sK6))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X41,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f5374,f142])).
fof(f7054,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X31),s('type.list.list'('type.bvlSem.v'),sK6))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X31,'type.bvlSem.v')))) | s(bool,X33) != s(bool,X34)) )),
  inference(resolution,[],[f5374,f105])).
fof(f7051,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X19),s('type.list.list'('type.bvlSem.v'),sK6))) | p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X19,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f5374,f104])).
fof(f7050,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X15),s('type.list.list'('type.bvlSem.v'),sK6))) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X15,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f5374,f103])).
fof(f7049,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X11),s('type.list.list'('type.bvlSem.v'),sK6))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X11,'type.bvlSem.v')))) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f5374,f101])).
fof(f7047,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK6))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X3,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f5374,f99])).
fof(f7039,plain,(
  ( ! [X39,X37,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X39)))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X40))) | ~p(s(bool,X42)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X37,X39),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X40)))))) ) | $spl196),
  inference(cnf_transformation,[],[f7039_D])).
fof(f7039_D,plain,(
  ( ! [X39,X37,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X39)))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X40))) | ~p(s(bool,X42)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X37,X39),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X40)))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f7025,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X117,X119),s(fun(X117,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X117),X120))))) | p(s(bool,X121)) | ~p(s(bool,X122)) | s(bool,X122) != s(bool,'const.list.EVERY_2'(s(fun(X117,bool),X118),s('type.list.list'(X117),X120)))) )),
  inference(resolution,[],[f3656,f142])).
fof(f7018,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (p(s(bool,happ(s(fun(X86,bool),X87),s(X86,X88)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X86,X88),s(fun(X86,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X86),X89))))) | ~p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X91) != s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X87),s('type.list.list'(X86),X89))) | s(bool,X90) != s(bool,X92)) )),
  inference(resolution,[],[f3656,f105])).
fof(f7016,plain,(
  ( ! [X78,X76,X74,X72,X77,X75,X73] : (p(s(bool,happ(s(fun(X72,bool),X73),s(X72,X74)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X72,X74),s(fun(X72,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X72),X75))))) | p(s(bool,X76)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,'const.list.EVERY_2'(s(fun(X72,bool),X73),s('type.list.list'(X72),X75)))) )),
  inference(resolution,[],[f3656,f104])).
fof(f7014,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (p(s(bool,happ(s(fun(X58,bool),X59),s(X58,X60)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X58,X60),s(fun(X58,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X58),X61))))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.list.EVERY_2'(s(fun(X58,bool),X59),s('type.list.list'(X58),X61)))) )),
  inference(resolution,[],[f3656,f103])).
fof(f7013,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (p(s(bool,happ(s(fun(X51,bool),X52),s(X51,X53)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X51,X53),s(fun(X51,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X51),X54))))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),X54))) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f3656,f101])).
fof(f7011,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X39)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X37,X39),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X40))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X40)))) )),
  inference(resolution,[],[f3656,f99])).
fof(f6990,plain,(
  ( ! [X158,X159] : (s(bool,X158) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X159) | s(bool,X158) != s(bool,X159)) ) | $spl194),
  inference(cnf_transformation,[],[f6990_D])).
fof(f6990_D,plain,(
  ( ! [X158,X159] : (s(bool,X158) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X159) | s(bool,X158) != s(bool,X159)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f6989,plain,(
  ( ! [X158,X159,X157] : (s(bool,X158) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X158) != s(bool,X159) | s(bool,'F_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X159)) ) | ($spl3 | $spl6 | $spl14 | $spl40)),
  inference(subsumption_resolution,[],[f6907,f116])).
fof(f6907,plain,(
  ( ! [X158,X159,X157] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X158) != s(bool,X159) | s(bool,'F_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X159)) ) | ($spl3 | $spl6 | $spl14 | $spl40)),
  inference(superposition,[],[f1953,f6765])).
fof(f6979,plain,(
  ( ! [X127] : (s(bool,X127) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X127)) ) | $spl192),
  inference(cnf_transformation,[],[f6979_D])).
fof(f6979_D,plain,(
  ( ! [X127] : (s(bool,X127) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X127)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f6978,plain,(
  ( ! [X127,X126] : (s(bool,'F_0') = s(bool,X126) | s(bool,X127) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X127)) ) | ($spl3 | $spl6 | $spl14)),
  inference(subsumption_resolution,[],[f6889,f116])).
fof(f6889,plain,(
  ( ! [X127,X126] : (s(bool,'T_0') != s(bool,X126) | s(bool,'F_0') = s(bool,X126) | s(bool,X127) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,'T_0') = s(bool,X127)) ) | ($spl3 | $spl6 | $spl14)),
  inference(superposition,[],[f1032,f6765])).
fof(f6806,plain,(
  ( ! [X39,X37,X38,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X37,X38),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X39)))))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39))) | ~p(s(bool,X42)) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X40),s(X37,X38)))) ) | $spl190),
  inference(cnf_transformation,[],[f6806_D])).
fof(f6806_D,plain,(
  ( ! [X39,X37,X38,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X37,X38),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X39)))))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39))) | ~p(s(bool,X42)) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X40),s(X37,X38)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f6792,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (~p(s(bool,'const.bool.IN_2'(s(X117,X118),s(fun(X117,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X117),X119)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X117,bool),X120),s(X117,X118))) | p(s(bool,X121)) | ~p(s(bool,X122)) | s(bool,X122) != s(bool,'const.list.EVERY_2'(s(fun(X117,bool),X120),s('type.list.list'(X117),X119)))) )),
  inference(resolution,[],[f3637,f142])).
fof(f6785,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (~p(s(bool,'const.bool.IN_2'(s(X86,X87),s(fun(X86,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X86),X88)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X86,bool),X89),s(X86,X87))) | ~p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X91) != s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X89),s('type.list.list'(X86),X88))) | s(bool,X90) != s(bool,X92)) )),
  inference(resolution,[],[f3637,f105])).
fof(f6783,plain,(
  ( ! [X78,X76,X74,X72,X77,X75,X73] : (~p(s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X72),X74)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),X75),s(X72,X73))) | p(s(bool,X76)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,'const.list.EVERY_2'(s(fun(X72,bool),X75),s('type.list.list'(X72),X74)))) )),
  inference(resolution,[],[f3637,f104])).
fof(f6781,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (~p(s(bool,'const.bool.IN_2'(s(X58,X59),s(fun(X58,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X58),X60)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X58,bool),X61),s(X58,X59))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.list.EVERY_2'(s(fun(X58,bool),X61),s('type.list.list'(X58),X60)))) )),
  inference(resolution,[],[f3637,f103])).
fof(f6780,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(X51,X52),s(fun(X51,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X51),X53)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X51,bool),X54),s(X51,X52))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.list.EVERY_2'(s(fun(X51,bool),X54),s('type.list.list'(X51),X53))) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f3637,f101])).
fof(f6778,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X37,X38),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X39)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X40),s(X37,X38))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39)))) )),
  inference(resolution,[],[f3637,f99])).
fof(f6764,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X2))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl188),
  inference(cnf_transformation,[],[f6764_D])).
fof(f6764_D,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X2))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f6752,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X40))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X40),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(resolution,[],[f2777,f142])).
fof(f6745,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X30))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X30),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X31) != s(bool,X33)) ) | $spl6),
  inference(resolution,[],[f2777,f105])).
fof(f6743,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X22))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X22),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(resolution,[],[f2777,f104])).
fof(f6741,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X14),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(resolution,[],[f2777,f103])).
fof(f6740,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X10))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X10),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X11) != s(bool,X13)) ) | $spl6),
  inference(resolution,[],[f2777,f101])).
fof(f6738,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl6),
  inference(resolution,[],[f2777,f99])).
fof(f6699,plain,(
  ( ! [X43,X44,X42] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X42),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X42)))) )),
  inference(resolution,[],[f1794,f142])).
fof(f6692,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X32),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X32))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f1794,f105])).
fof(f6689,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X20),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X20)))) )),
  inference(resolution,[],[f1794,f104])).
fof(f6688,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X16),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X16)))) )),
  inference(resolution,[],[f1794,f103])).
fof(f6687,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X12),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X12))) | s(bool,X14) != s(bool,X15)) )),
  inference(resolution,[],[f1794,f101])).
fof(f6685,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X4),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X4)))) )),
  inference(resolution,[],[f1794,f99])).
fof(f6677,plain,(
  ( ! [X39,X37,X38,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X37,X38),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X39)))))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39))) | ~p(s(bool,X42)) | s(bool,'F_0') != s(bool,happ(s(fun(X37,bool),X40),s(X37,X38)))) ) | $spl184),
  inference(cnf_transformation,[],[f6677_D])).
fof(f6677_D,plain,(
  ( ! [X39,X37,X38,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X37,X38),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X39)))))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39))) | ~p(s(bool,X42)) | s(bool,'F_0') != s(bool,happ(s(fun(X37,bool),X40),s(X37,X38)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f6663,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (~p(s(bool,'const.bool.IN_2'(s(X117,X118),s(fun(X117,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X117),X119)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X117,bool),X120),s(X117,X118))) | p(s(bool,X121)) | ~p(s(bool,X122)) | s(bool,X122) != s(bool,'const.list.EVERY_2'(s(fun(X117,bool),X120),s('type.list.list'(X117),X119)))) )),
  inference(resolution,[],[f3635,f142])).
fof(f3635,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X64,bool),X67),s('type.list.list'(X64),X66)))) | ~p(s(bool,'const.bool.IN_2'(s(X64,X65),s(fun(X64,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X64),X66)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X64,bool),X67),s(X64,X65)))) )),
  inference(resolution,[],[f94,f108])).
fof(f6656,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (~p(s(bool,'const.bool.IN_2'(s(X86,X87),s(fun(X86,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X86),X88)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X86,bool),X89),s(X86,X87))) | ~p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X91) != s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X89),s('type.list.list'(X86),X88))) | s(bool,X90) != s(bool,X92)) )),
  inference(resolution,[],[f3635,f105])).
fof(f6654,plain,(
  ( ! [X78,X76,X74,X72,X77,X75,X73] : (~p(s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X72),X74)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X72,bool),X75),s(X72,X73))) | p(s(bool,X76)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,'const.list.EVERY_2'(s(fun(X72,bool),X75),s('type.list.list'(X72),X74)))) )),
  inference(resolution,[],[f3635,f104])).
fof(f6652,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (~p(s(bool,'const.bool.IN_2'(s(X58,X59),s(fun(X58,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X58),X60)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X58,bool),X61),s(X58,X59))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.list.EVERY_2'(s(fun(X58,bool),X61),s('type.list.list'(X58),X60)))) )),
  inference(resolution,[],[f3635,f103])).
fof(f6651,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(X51,X52),s(fun(X51,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X51),X53)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X51,bool),X54),s(X51,X52))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.list.EVERY_2'(s(fun(X51,bool),X54),s('type.list.list'(X51),X53))) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f3635,f101])).
fof(f6649,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X37,X38),s(fun(X37,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X37),X39)))))) | s(bool,'F_0') != s(bool,happ(s(fun(X37,bool),X40),s(X37,X38))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39)))) )),
  inference(resolution,[],[f3635,f99])).
fof(f6585,plain,(
  ( ! [X6,X10,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X5),X7)))))) | s(bool,X10) != s(bool,happ(s(fun(X5,bool),X8),s(X5,sK7(X7,X8,X5)))) | ~p(s(bool,X10)) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,X6))))) ) | $spl182),
  inference(cnf_transformation,[],[f6585_D])).
fof(f6585_D,plain,(
  ( ! [X6,X10,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X5),X7)))))) | s(bool,X10) != s(bool,happ(s(fun(X5,bool),X8),s(X5,sK7(X7,X8,X5)))) | ~p(s(bool,X10)) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,X6))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f6564,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,'const.bool.IN_2'(s(X85,X86),s(fun(X85,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X85),X87)))))) | p(s(bool,happ(s(fun(X85,bool),X88),s(X85,X86)))) | p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,happ(s(fun(X85,bool),X88),s(X85,sK7(X87,X88,X85))))) )),
  inference(resolution,[],[f3668,f142])).
fof(f3668,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK7(X7,X5,X4))))) | ~p(s(bool,'const.bool.IN_2'(s(X4,X6),s(fun(X4,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X4),X7)))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) )),
  inference(resolution,[],[f94,f96])).
fof(f6557,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(X54,X55),s(fun(X54,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X54),X56)))))) | p(s(bool,happ(s(fun(X54,bool),X57),s(X54,X55)))) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,happ(s(fun(X54,bool),X57),s(X54,sK7(X56,X57,X54)))) | s(bool,X58) != s(bool,X60)) )),
  inference(resolution,[],[f3668,f105])).
fof(f6555,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X40,X41),s(fun(X40,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X40),X42)))))) | p(s(bool,happ(s(fun(X40,bool),X43),s(X40,X41)))) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,happ(s(fun(X40,bool),X43),s(X40,sK7(X42,X43,X40))))) )),
  inference(resolution,[],[f3668,f104])).
fof(f6553,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (~p(s(bool,'const.bool.IN_2'(s(X26,X27),s(fun(X26,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X26),X28)))))) | p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X27)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(X26,bool),X29),s(X26,sK7(X28,X29,X26))))) )),
  inference(resolution,[],[f3668,f103])).
fof(f6552,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(X19,X20),s(fun(X19,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X19),X21)))))) | p(s(bool,happ(s(fun(X19,bool),X22),s(X19,X20)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X19,bool),X22),s(X19,sK7(X21,X22,X19)))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f3668,f101])).
fof(f6550,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X5),X7)))))) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,X6)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,happ(s(fun(X5,bool),X8),s(X5,sK7(X7,X8,X5))))) )),
  inference(resolution,[],[f3668,f99])).
fof(f6545,plain,(
  ( ! [X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X3,'type.bvlSem.v'))))) ) | $spl180),
  inference(cnf_transformation,[],[f6545_D])).
fof(f6545_D,plain,(
  ( ! [X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X3,'type.bvlSem.v'))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f6530,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X41),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X41,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f6512,f142])).
fof(f6523,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X31),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X31,'type.bvlSem.v')))) | s(bool,X33) != s(bool,X34)) ) | $spl6),
  inference(resolution,[],[f6512,f105])).
fof(f6520,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X19),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X19,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f6512,f104])).
fof(f6519,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X15),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X15,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f6512,f103])).
fof(f6518,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X11),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X11,'type.bvlSem.v')))) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f6512,f101])).
fof(f6516,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK5)))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK5,X3,'type.bvlSem.v'))))) ) | $spl6),
  inference(resolution,[],[f6512,f99])).
fof(f6508,plain,(
  ( ! [X39,X37,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK7(X39,X40,X37))))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X39))) | ~p(s(bool,X42)) | p(s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39))))) ) | $spl178),
  inference(cnf_transformation,[],[f6508_D])).
fof(f6508_D,plain,(
  ( ! [X39,X37,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK7(X39,X40,X37))))) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X39))) | ~p(s(bool,X42)) | p(s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f6494,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK7(X119,X120,X117))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X117,bool),X120),s('type.list.list'(X117),X119)))) | p(s(bool,X121)) | ~p(s(bool,X122)) | s(bool,X122) != s(bool,'const.list.EVERY_2'(s(fun(X117,bool),X118),s('type.list.list'(X117),X119)))) )),
  inference(resolution,[],[f3647,f142])).
fof(f6487,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (p(s(bool,happ(s(fun(X86,bool),X87),s(X86,sK7(X88,X89,X86))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X89),s('type.list.list'(X86),X88)))) | ~p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X91) != s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X87),s('type.list.list'(X86),X88))) | s(bool,X90) != s(bool,X92)) )),
  inference(resolution,[],[f3647,f105])).
fof(f6485,plain,(
  ( ! [X78,X76,X74,X72,X77,X75,X73] : (p(s(bool,happ(s(fun(X72,bool),X73),s(X72,sK7(X74,X75,X72))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X72,bool),X75),s('type.list.list'(X72),X74)))) | p(s(bool,X76)) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,'const.list.EVERY_2'(s(fun(X72,bool),X73),s('type.list.list'(X72),X74)))) )),
  inference(resolution,[],[f3647,f104])).
fof(f6483,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (p(s(bool,happ(s(fun(X58,bool),X59),s(X58,sK7(X60,X61,X58))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X58,bool),X61),s('type.list.list'(X58),X60)))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.list.EVERY_2'(s(fun(X58,bool),X59),s('type.list.list'(X58),X60)))) )),
  inference(resolution,[],[f3647,f103])).
fof(f6482,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (p(s(bool,happ(s(fun(X51,bool),X52),s(X51,sK7(X53,X54,X51))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X51,bool),X54),s('type.list.list'(X51),X53)))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),X53))) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f3647,f101])).
fof(f6480,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK7(X39,X40,X37))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X40),s('type.list.list'(X37),X39)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.list.EVERY_2'(s(fun(X37,bool),X38),s('type.list.list'(X37),X39)))) )),
  inference(resolution,[],[f3647,f99])).
fof(f6467,plain,(
  ( ! [X4,X7,X5] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X4),X5),s('type.bvlSem.v',sK3)))) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X4),X5))))) | ~p(s(bool,X7))) ) | $spl176),
  inference(cnf_transformation,[],[f6467_D])).
fof(f6467_D,plain,(
  ( ! [X4,X7,X5] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X4),X5),s('type.bvlSem.v',sK3)))) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X4),X5))))) | ~p(s(bool,X7))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f6456,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X56),X57),s('type.bvlSem.v',sK3)))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X56),X57)))))) ) | $spl4),
  inference(resolution,[],[f2012,f142])).
fof(f2012,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X1)))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',sK3))))) ) | $spl4),
  inference(resolution,[],[f149,f118])).
fof(f6449,plain,(
  ( ! [X39,X43,X41,X42,X40] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X39),X40),s('type.bvlSem.v',sK3)))) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X39),X40))))) | s(bool,X41) != s(bool,X43)) ) | $spl4),
  inference(resolution,[],[f2012,f105])).
fof(f6447,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X29),X30),s('type.bvlSem.v',sK3)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X29),X30)))))) ) | $spl4),
  inference(resolution,[],[f2012,f104])).
fof(f6445,plain,(
  ( ! [X23,X21,X19,X22,X20] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X19),X20),s('type.bvlSem.v',sK3)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X19),X20)))))) ) | $spl4),
  inference(resolution,[],[f2012,f103])).
fof(f6444,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X14),X15),s('type.bvlSem.v',sK3)))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X14),X15))))) | s(bool,X16) != s(bool,X18)) ) | $spl4),
  inference(resolution,[],[f2012,f101])).
fof(f6442,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X4),X5),s('type.bvlSem.v',sK3)))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X4),X5)))))) ) | $spl4),
  inference(resolution,[],[f2012,f99])).
fof(f6332,plain,(
  ( ! [X156,X155] : (s(bool,X155) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X156) | s(bool,X155) != s(bool,X156)) ) | $spl174),
  inference(cnf_transformation,[],[f6332_D])).
fof(f6332_D,plain,(
  ( ! [X156,X155] : (s(bool,X155) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X156) | s(bool,X155) != s(bool,X156)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f6331,plain,(
  ( ! [X156,X154,X155] : (s(bool,X155) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | s(bool,X155) != s(bool,X156) | s(bool,'F_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X156)) ) | ($spl4 | $spl14 | $spl40)),
  inference(subsumption_resolution,[],[f6232,f116])).
fof(f6232,plain,(
  ( ! [X156,X154,X155] : (s(bool,'T_0') != s(bool,X154) | s(bool,X155) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | s(bool,X155) != s(bool,X156) | s(bool,'F_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X156)) ) | ($spl4 | $spl14 | $spl40)),
  inference(superposition,[],[f1953,f6136])).
fof(f6318,plain,(
  ( ! [X124] : (s(bool,X124) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X124)) ) | $spl172),
  inference(cnf_transformation,[],[f6318_D])).
fof(f6318_D,plain,(
  ( ! [X124] : (s(bool,X124) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X124)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f6317,plain,(
  ( ! [X123,X124] : (s(bool,'F_0') = s(bool,X123) | s(bool,X124) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl4 | $spl14)),
  inference(subsumption_resolution,[],[f6215,f116])).
fof(f6215,plain,(
  ( ! [X123,X124] : (s(bool,'T_0') != s(bool,X123) | s(bool,'F_0') = s(bool,X123) | s(bool,X124) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl4 | $spl14)),
  inference(superposition,[],[f1032,f6136])).
fof(f6061,plain,(
  ( ! [X94,X95] : (s(bool,X94) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X95) | s(bool,X94) != s(bool,X95)) ) | $spl170),
  inference(cnf_transformation,[],[f6061_D])).
fof(f6061_D,plain,(
  ( ! [X94,X95] : (s(bool,X94) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X95) | s(bool,X94) != s(bool,X95)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f6060,plain,(
  ( ! [X94,X95,X93] : (s(bool,X94) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X94) != s(bool,X95) | s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X95)) ) | ($spl3 | $spl14 | $spl40)),
  inference(subsumption_resolution,[],[f6024,f116])).
fof(f6024,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,X93) | s(bool,X94) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X94) != s(bool,X95) | s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X95)) ) | ($spl3 | $spl14 | $spl40)),
  inference(superposition,[],[f1953,f2671])).
fof(f6059,plain,(
  ( ! [X68,X64,X67,X65,X63] : (s(bool,X67) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X63),X64))),s('type.bvlSem.v',X65))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X63),X64),s('type.bvlSem.v',X65))) | s(bool,'T_0') = s(bool,X68) | s(bool,X67) != s(bool,X68)) ) | $spl168),
  inference(cnf_transformation,[],[f6059_D])).
fof(f6059_D,plain,(
  ( ! [X68,X64,X67,X65,X63] : (s(bool,X67) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X63),X64))),s('type.bvlSem.v',X65))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X63),X64),s('type.bvlSem.v',X65))) | s(bool,'T_0') = s(bool,X68) | s(bool,X67) != s(bool,X68)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f6058,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X67) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X63),X64))),s('type.bvlSem.v',X65))) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X63),X64),s('type.bvlSem.v',X65)))) ) | ($spl14 | $spl40)),
  inference(subsumption_resolution,[],[f6018,f116])).
fof(f6018,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X63),X64))),s('type.bvlSem.v',X65))) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X63),X64),s('type.bvlSem.v',X65)))) ) | ($spl14 | $spl40)),
  inference(superposition,[],[f1953,f3326])).
fof(f6057,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X59) | s(bool,X58) != s(bool,X59)) ) | $spl166),
  inference(cnf_transformation,[],[f6057_D])).
fof(f6057_D,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X59) | s(bool,X58) != s(bool,X59)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f6056,plain,(
  ( ! [X59,X57,X58] : (s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X58) != s(bool,X59) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X59)) ) | ($spl6 | $spl14 | $spl40)),
  inference(subsumption_resolution,[],[f6016,f116])).
fof(f6016,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') != s(bool,X57) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X58) != s(bool,X59) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X59)) ) | ($spl6 | $spl14 | $spl40)),
  inference(superposition,[],[f1953,f2248])).
fof(f6055,plain,(
  ( ! [X56,X55] : (s(bool,X55) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X56) | s(bool,X55) != s(bool,X56)) ) | $spl164),
  inference(cnf_transformation,[],[f6055_D])).
fof(f6055_D,plain,(
  ( ! [X56,X55] : (s(bool,X55) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X56) | s(bool,X55) != s(bool,X56)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f6054,plain,(
  ( ! [X54,X56,X55] : (s(bool,X55) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X56)) ) | ($spl14 | $spl40)),
  inference(subsumption_resolution,[],[f6015,f116])).
fof(f6015,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,X54) | s(bool,X55) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X56)) ) | ($spl14 | $spl40)),
  inference(superposition,[],[f1953,f523])).
fof(f5992,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,sK7(X3,X4,bool))) ) | $spl162),
  inference(cnf_transformation,[],[f5992_D])).
fof(f5992_D,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,sK7(X3,X4,bool))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f5963,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X55)))))) | s(bool,'T_0') = s(bool,sK7(X55,X56,bool)) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X56),s('type.list.list'(bool),X55)))) )),
  inference(resolution,[],[f3533,f142])).
fof(f3533,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X10)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10)))) | s(bool,'T_0') = s(bool,sK7(X10,X11,bool))) )),
  inference(superposition,[],[f95,f116])).
fof(f5956,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X38)))))) | s(bool,'T_0') = s(bool,sK7(X38,X39,bool)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X38))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3533,f105])).
fof(f5953,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X23)))))) | s(bool,'T_0') = s(bool,sK7(X23,X24,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X24),s('type.list.list'(bool),X23)))) )),
  inference(resolution,[],[f3533,f104])).
fof(f5952,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X18)))))) | s(bool,'T_0') = s(bool,sK7(X18,X19,bool)) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X19),s('type.list.list'(bool),X18)))) )),
  inference(resolution,[],[f3533,f103])).
fof(f5951,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X13)))))) | s(bool,'T_0') = s(bool,sK7(X13,X14,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X14),s('type.list.list'(bool),X13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3533,f101])).
fof(f5949,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | s(bool,'T_0') = s(bool,sK7(X3,X4,bool)) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3)))) )),
  inference(resolution,[],[f3533,f99])).
fof(f5991,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool))) ) | $spl160),
  inference(cnf_transformation,[],[f5991_D])).
fof(f5991_D,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f5941,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56)))) | s(bool,'T_0') = s(bool,sK7(X56,X55,bool)) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X56)))))) )),
  inference(resolution,[],[f3533,f142])).
fof(f5934,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39)))) | s(bool,'T_0') = s(bool,sK7(X39,X38,bool)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X39))))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3533,f105])).
fof(f5931,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X24)))) | s(bool,'T_0') = s(bool,sK7(X24,X23,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X24)))))) )),
  inference(resolution,[],[f3533,f104])).
fof(f5930,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) | s(bool,'T_0') = s(bool,sK7(X19,X18,bool)) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X19)))))) )),
  inference(resolution,[],[f3533,f103])).
fof(f5929,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14)))) | s(bool,'T_0') = s(bool,sK7(X14,X13,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X14))))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3533,f101])).
fof(f5927,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool)) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4)))))) )),
  inference(resolution,[],[f3533,f99])).
fof(f5883,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | s(bool,X3) != s(bool,sK7(X1,X0,bool)) | ~p(s(bool,X3)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl158),
  inference(cnf_transformation,[],[f5883_D])).
fof(f5883_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | s(bool,X3) != s(bool,sK7(X1,X0,bool)) | ~p(s(bool,X3)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f5872,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X53)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,'T_0'))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,sK7(X53,X52,bool))) )),
  inference(resolution,[],[f5640,f142])).
fof(f5640,plain,(
  ( ! [X8,X9] : (~p(s(bool,sK7(X9,X8,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) )),
  inference(trivial_inequality_removal,[],[f5638])).
fof(f5638,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9)))) | ~p(s(bool,sK7(X9,X8,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) )),
  inference(superposition,[],[f3821,f116])).
fof(f5865,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X36)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0'))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,sK7(X36,X35,bool)) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f5640,f105])).
fof(f5863,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X26)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0'))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,sK7(X26,X25,bool))) )),
  inference(resolution,[],[f5640,f104])).
fof(f5861,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X15),s('type.list.list'(bool),X16)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,'T_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,sK7(X16,X15,bool))) )),
  inference(resolution,[],[f5640,f103])).
fof(f5860,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X10),s('type.list.list'(bool),X11)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0'))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,sK7(X11,X10,bool)) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f5640,f101])).
fof(f5858,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,sK7(X1,X0,bool))) )),
  inference(resolution,[],[f5640,f99])).
fof(f5854,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,sK7(X3,X4,bool))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3))) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) ) | $spl156),
  inference(cnf_transformation,[],[f5854_D])).
fof(f5854_D,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,sK7(X3,X4,bool))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3))) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f5841,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,sK7(X55,X56,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X56),s(bool,'F_0'))) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X56),s('type.list.list'(bool),X55)))) )),
  inference(resolution,[],[f5623,f142])).
fof(f5623,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9)))) | p(s(bool,sK7(X9,X8,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) )),
  inference(trivial_inequality_removal,[],[f5621])).
fof(f5621,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X8),s('type.list.list'(bool),X9)))) | p(s(bool,sK7(X9,X8,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) )),
  inference(superposition,[],[f3819,f116])).
fof(f5834,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,sK7(X38,X39,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X39),s(bool,'F_0'))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X38))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f5623,f105])).
fof(f5831,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,sK7(X23,X24,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,'F_0'))) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X24),s('type.list.list'(bool),X23)))) )),
  inference(resolution,[],[f5623,f104])).
fof(f5830,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,sK7(X18,X19,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X19),s(bool,'F_0'))) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X19),s('type.list.list'(bool),X18)))) )),
  inference(resolution,[],[f5623,f103])).
fof(f5829,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,sK7(X13,X14,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,'F_0'))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X14),s('type.list.list'(bool),X13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f5623,f101])).
fof(f5827,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,sK7(X3,X4,bool))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0'))) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3)))) )),
  inference(resolution,[],[f5623,f99])).
fof(f5819,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3))) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,sK7(X3,X4,bool))) ) | $spl154),
  inference(cnf_transformation,[],[f5819_D])).
fof(f5819_D,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3))) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,sK7(X3,X4,bool))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f5793,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X55)))))) | s(bool,'F_0') = s(bool,sK7(X55,X56,bool)) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X56),s('type.list.list'(bool),X55)))) )),
  inference(resolution,[],[f3532,f142])).
fof(f3532,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X9),s('type.list.list'(bool),X8)))) | s(bool,'F_0') = s(bool,sK7(X8,X9,bool))) )),
  inference(superposition,[],[f95,f116])).
fof(f5786,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X38)))))) | s(bool,'F_0') = s(bool,sK7(X38,X39,bool)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X38))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3532,f105])).
fof(f5783,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X23)))))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X24),s('type.list.list'(bool),X23)))) )),
  inference(resolution,[],[f3532,f104])).
fof(f5782,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X18)))))) | s(bool,'F_0') = s(bool,sK7(X18,X19,bool)) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X19),s('type.list.list'(bool),X18)))) )),
  inference(resolution,[],[f3532,f103])).
fof(f5781,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X13)))))) | s(bool,'F_0') = s(bool,sK7(X13,X14,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X14),s('type.list.list'(bool),X13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3532,f101])).
fof(f5779,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | s(bool,'F_0') = s(bool,sK7(X3,X4,bool)) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3)))) )),
  inference(resolution,[],[f3532,f99])).
fof(f5818,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) | $spl152),
  inference(cnf_transformation,[],[f5818_D])).
fof(f5818_D,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f5771,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56)))) | s(bool,'F_0') = s(bool,sK7(X56,X55,bool)) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X56)))))) )),
  inference(resolution,[],[f3532,f142])).
fof(f5764,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39)))) | s(bool,'F_0') = s(bool,sK7(X39,X38,bool)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X39))))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3532,f105])).
fof(f5761,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X24)))) | s(bool,'F_0') = s(bool,sK7(X24,X23,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X24)))))) )),
  inference(resolution,[],[f3532,f104])).
fof(f5760,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) | s(bool,'F_0') = s(bool,sK7(X19,X18,bool)) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X19)))))) )),
  inference(resolution,[],[f3532,f103])).
fof(f5759,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14)))) | s(bool,'F_0') = s(bool,sK7(X14,X13,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X14))))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3532,f101])).
fof(f5757,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool)) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4)))))) )),
  inference(resolution,[],[f3532,f99])).
fof(f5719,plain,(
  ( ! [X6,X10,X8,X7] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7))),s('type.list.list'('type.bvlSem.v'),X8)))) | s(bool,X10) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X7),s('type.bvlSem.v',sK7(X8,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7)),'type.bvlSem.v')))) | ~p(s(bool,X10))) ) | $spl150),
  inference(cnf_transformation,[],[f5719_D])).
fof(f5719_D,plain,(
  ( ! [X6,X10,X8,X7] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7))),s('type.list.list'('type.bvlSem.v'),X8)))) | s(bool,X10) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X7),s('type.bvlSem.v',sK7(X8,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7)),'type.bvlSem.v')))) | ~p(s(bool,X10))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f5705,plain,(
  ( ! [X76,X74,X72,X75,X73] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X72),X73))),s('type.list.list'('type.bvlSem.v'),X74)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X72),X73),s('type.bvlSem.v',sK7(X74,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X72),X73)),'type.bvlSem.v'))))) )),
  inference(resolution,[],[f3327,f142])).
fof(f3327,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',sK7(X2,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X0),X1)),'type.bvlSem.v'))))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))),s('type.list.list'('type.bvlSem.v'),X2))))) )),
  inference(superposition,[],[f96,f97])).
fof(f5698,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X48),X49))),s('type.list.list'('type.bvlSem.v'),X50)))) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X48),X49),s('type.bvlSem.v',sK7(X50,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X48),X49)),'type.bvlSem.v')))) | s(bool,X51) != s(bool,X53)) )),
  inference(resolution,[],[f3327,f105])).
fof(f5696,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X36),X37))),s('type.list.list'('type.bvlSem.v'),X38)))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X36),X37),s('type.bvlSem.v',sK7(X38,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X36),X37)),'type.bvlSem.v'))))) )),
  inference(resolution,[],[f3327,f104])).
fof(f5694,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X24),X25))),s('type.list.list'('type.bvlSem.v'),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X24),X25),s('type.bvlSem.v',sK7(X26,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X24),X25)),'type.bvlSem.v'))))) )),
  inference(resolution,[],[f3327,f103])).
fof(f5693,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X18),X19))),s('type.list.list'('type.bvlSem.v'),X20)))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X18),X19),s('type.bvlSem.v',sK7(X20,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X18),X19)),'type.bvlSem.v')))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f3327,f101])).
fof(f5691,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7))),s('type.list.list'('type.bvlSem.v'),X8)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X7),s('type.bvlSem.v',sK7(X8,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7)),'type.bvlSem.v'))))) )),
  inference(resolution,[],[f3327,f99])).
fof(f5668,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | s(bool,X3) != s(bool,sK7(X1,X0,bool)) | ~p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X1))))))) ) | $spl148),
  inference(cnf_transformation,[],[f5668_D])).
fof(f5668_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | s(bool,X3) != s(bool,sK7(X1,X0,bool)) | ~p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X1))))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f5657,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X53)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X53)))))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,sK7(X53,X52,bool))) )),
  inference(resolution,[],[f3530,f142])).
fof(f3530,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK7(X4,X5,bool))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))))) )),
  inference(superposition,[],[f95,f113])).
fof(f5650,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X36)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X36)))))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,sK7(X36,X35,bool)) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f3530,f105])).
fof(f5648,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X26)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X26)))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,sK7(X26,X25,bool))) )),
  inference(resolution,[],[f3530,f104])).
fof(f5646,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X15),s('type.list.list'(bool),X16)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X16)))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,sK7(X16,X15,bool))) )),
  inference(resolution,[],[f3530,f103])).
fof(f5645,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X10),s('type.list.list'(bool),X11)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X11)))))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,sK7(X11,X10,bool)) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f3530,f101])).
fof(f5643,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X1)))))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,sK7(X1,X0,bool))) )),
  inference(resolution,[],[f3530,f99])).
fof(f5606,plain,(
  ( ! [X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1))) | s(bool,X3) != s(bool,sK7(X1,X0,bool)) | ~p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl146),
  inference(cnf_transformation,[],[f5606_D])).
fof(f5606_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1))) | s(bool,X3) != s(bool,sK7(X1,X0,bool)) | ~p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f5595,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X53))) | ~p(s(bool,happ(s(fun(bool,bool),X52),s(bool,'T_0')))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,sK7(X53,X52,bool))) )),
  inference(resolution,[],[f3612,f142])).
fof(f3612,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK7(X4,X5,bool))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X5),s('type.list.list'(bool),X4))) | ~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0'))))) )),
  inference(superposition,[],[f3264,f113])).
fof(f5588,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X36))) | ~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,sK7(X36,X35,bool)) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f3612,f105])).
fof(f5586,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X26))) | ~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0')))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,sK7(X26,X25,bool))) )),
  inference(resolution,[],[f3612,f104])).
fof(f5584,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X15),s('type.list.list'(bool),X16))) | ~p(s(bool,happ(s(fun(bool,bool),X15),s(bool,'T_0')))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,sK7(X16,X15,bool))) )),
  inference(resolution,[],[f3612,f103])).
fof(f5583,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X10),s('type.list.list'(bool),X11))) | ~p(s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0')))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,sK7(X11,X10,bool)) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f3612,f101])).
fof(f5581,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,sK7(X1,X0,bool))) )),
  inference(resolution,[],[f3612,f99])).
fof(f5577,plain,(
  ( ! [X6,X4,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | p(s(bool,sK7(X4,X3,bool)))) ) | $spl144),
  inference(cnf_transformation,[],[f5577_D])).
fof(f5577_D,plain,(
  ( ! [X6,X4,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | p(s(bool,sK7(X4,X3,bool)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f5557,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56))) | p(s(bool,sK7(X56,X55,bool))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,happ(s(fun(bool,bool),X55),s(bool,'F_0')))) )),
  inference(resolution,[],[f3610,f142])).
fof(f5550,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39))) | p(s(bool,sK7(X39,X38,bool))) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(bool,bool),X38),s(bool,'F_0'))) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f3610,f105])).
fof(f5548,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X28),s('type.list.list'(bool),X29))) | p(s(bool,sK7(X29,X28,bool))) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,happ(s(fun(bool,bool),X28),s(bool,'F_0')))) )),
  inference(resolution,[],[f3610,f104])).
fof(f5546,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19))) | p(s(bool,sK7(X19,X18,bool))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),X18),s(bool,'F_0')))) )),
  inference(resolution,[],[f3610,f103])).
fof(f5545,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14))) | p(s(bool,sK7(X14,X13,bool))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0'))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3610,f101])).
fof(f5543,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | p(s(bool,sK7(X4,X3,bool))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) )),
  inference(resolution,[],[f3610,f99])).
fof(f5531,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3))) | p(s(bool,X6)) | p(s(bool,sK7(X3,X4,bool)))) ) | $spl142),
  inference(cnf_transformation,[],[f5531_D])).
fof(f5531_D,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3))) | p(s(bool,X6)) | p(s(bool,sK7(X3,X4,bool)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f5502,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X55)))))) | p(s(bool,sK7(X55,X56,bool))) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X56),s('type.list.list'(bool),X55)))) )),
  inference(resolution,[],[f3528,f142])).
fof(f5495,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X38)))))) | p(s(bool,sK7(X38,X39,bool))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X39),s('type.list.list'(bool),X38))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3528,f105])).
fof(f5492,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X23)))))) | p(s(bool,sK7(X23,X24,bool))) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X24),s('type.list.list'(bool),X23)))) )),
  inference(resolution,[],[f3528,f104])).
fof(f5491,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X18)))))) | p(s(bool,sK7(X18,X19,bool))) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X19),s('type.list.list'(bool),X18)))) )),
  inference(resolution,[],[f3528,f103])).
fof(f5490,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X13)))))) | p(s(bool,sK7(X13,X14,bool))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X14),s('type.list.list'(bool),X13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3528,f101])).
fof(f5488,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X3)))))) | p(s(bool,sK7(X3,X4,bool))) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),X3)))) )),
  inference(resolution,[],[f3528,f99])).
fof(f5530,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))) | p(s(bool,sK7(X4,X3,bool)))) ) | $spl140),
  inference(cnf_transformation,[],[f5530_D])).
fof(f5530_D,plain,(
  ( ! [X6,X4,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4))))) | p(s(bool,sK7(X4,X3,bool)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f5480,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56)))) | p(s(bool,sK7(X56,X55,bool))) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X56)))))) )),
  inference(resolution,[],[f3528,f142])).
fof(f5473,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39)))) | p(s(bool,sK7(X39,X38,bool))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X39))))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3528,f105])).
fof(f5470,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X24)))) | p(s(bool,sK7(X24,X23,bool))) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X24)))))) )),
  inference(resolution,[],[f3528,f104])).
fof(f5469,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) | p(s(bool,sK7(X19,X18,bool))) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X19)))))) )),
  inference(resolution,[],[f3528,f103])).
fof(f5468,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14)))) | p(s(bool,sK7(X14,X13,bool))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X14))))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3528,f101])).
fof(f5466,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) | p(s(bool,sK7(X4,X3,bool))) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4)))))) )),
  inference(resolution,[],[f3528,f99])).
fof(f5461,plain,(
  ( ! [X14,X12,X10,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X10,sK7(X11,X12,X10)),s(fun(X10,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X10),X11))))) | s(bool,X14) != s(bool,'const.list.EVERY_2'(s(fun(X10,bool),X12),s('type.list.list'(X10),X11))) | p(s(bool,X14))) ) | $spl138),
  inference(cnf_transformation,[],[f5461_D])).
fof(f5461_D,plain,(
  ( ! [X14,X12,X10,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X10,sK7(X11,X12,X10)),s(fun(X10,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X10),X11))))) | s(bool,X14) != s(bool,'const.list.EVERY_2'(s(fun(X10,bool),X12),s('type.list.list'(X10),X11))) | p(s(bool,X14))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f5445,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X76,sK7(X77,X78,X76)),s(fun(X76,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X76),X77))))) | p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,'const.list.EVERY_2'(s(fun(X76,bool),X78),s('type.list.list'(X76),X77)))) )),
  inference(resolution,[],[f3509,f142])).
fof(f3509,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'const.list.EVERY_2'(s(fun(X54,bool),X55),s('type.list.list'(X54),X56)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X54,sK7(X56,X55,X54)),s(fun(X54,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X54),X56)))))) )),
  inference(resolution,[],[f95,f113])).
fof(f5438,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X52,sK7(X53,X54,X52)),s(fun(X52,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X52),X53))))) | p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,'const.list.EVERY_2'(s(fun(X52,bool),X54),s('type.list.list'(X52),X53))) | s(bool,X56) != s(bool,X57)) )),
  inference(resolution,[],[f3509,f105])).
fof(f5435,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X34,sK7(X35,X36,X34)),s(fun(X34,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X34),X35))))) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(X34,bool),X36),s('type.list.list'(X34),X35)))) )),
  inference(resolution,[],[f3509,f104])).
fof(f5434,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,sK7(X29,X30,X28)),s(fun(X28,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X28),X29))))) | p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.list.EVERY_2'(s(fun(X28,bool),X30),s('type.list.list'(X28),X29)))) )),
  inference(resolution,[],[f3509,f103])).
fof(f5433,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X22,sK7(X23,X24,X22)),s(fun(X22,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X22),X23))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun(X22,bool),X24),s('type.list.list'(X22),X23))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f3509,f101])).
fof(f5431,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X10,sK7(X11,X12,X10)),s(fun(X10,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X10),X11))))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.list.EVERY_2'(s(fun(X10,bool),X12),s('type.list.list'(X10),X11)))) )),
  inference(resolution,[],[f3509,f99])).
fof(f5389,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X71,bool),X72),s('type.list.list'(X71),X73))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(X71,sK7(X73,X72,X71)),s(fun(X71,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X71),X73)))))) )),
  inference(resolution,[],[f3489,f142])).
fof(f5382,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),X49))) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.bool.IN_2'(s(X47,sK7(X49,X48,X47)),s(fun(X47,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X47),X49))))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f3489,f105])).
fof(f5379,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X29,bool),X30),s('type.list.list'(X29),X31))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.bool.IN_2'(s(X29,sK7(X31,X30,X29)),s(fun(X29,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X29),X31)))))) )),
  inference(resolution,[],[f3489,f104])).
fof(f5378,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X23,bool),X24),s('type.list.list'(X23),X25))) | p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(X23,sK7(X25,X24,X23)),s(fun(X23,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X23),X25)))))) )),
  inference(resolution,[],[f3489,f103])).
fof(f5377,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X17,bool),X18),s('type.list.list'(X17),X19))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(X17,sK7(X19,X18,X17)),s(fun(X17,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X17),X19))))) | s(bool,X21) != s(bool,X22)) )),
  inference(resolution,[],[f3489,f101])).
fof(f5375,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X5,bool),X6),s('type.list.list'(X5),X7))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(X5,sK7(X7,X6,X5)),s(fun(X5,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X5),X7)))))) )),
  inference(resolution,[],[f3489,f99])).
fof(f3623,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,X4))) ) | $spl94),
  inference(cnf_transformation,[],[f3623_D])).
fof(f3623_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,X4))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5269,plain,(
  ( ! [X313,X310,X312,X314,X311] : (p(s(bool,X313)) | ~p(s(bool,X314)) | s(bool,X314) != s(bool,happ(s(fun(X310,bool),X311),s(X310,sK7(X312,X311,X310)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X310,bool),X311),s('type.list.list'(X310),X312)))) )),
  inference(subsumption_resolution,[],[f5210,f127])).
fof(f5210,plain,(
  ( ! [X313,X310,X312,X314,X311] : (p(s(bool,'F_0')) | p(s(bool,X313)) | ~p(s(bool,X314)) | s(bool,X314) != s(bool,happ(s(fun(X310,bool),X311),s(X310,sK7(X312,X311,X310)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X310,bool),X311),s('type.list.list'(X310),X312)))) )),
  inference(superposition,[],[f142,f3599])).
fof(f5265,plain,(
  ( ! [X218,X216,X220,X219,X217,X221] : (~p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X220) != s(bool,happ(s(fun(X216,bool),X217),s(X216,sK7(X218,X217,X216)))) | s(bool,X219) != s(bool,X221) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X216,bool),X217),s('type.list.list'(X216),X218)))) )),
  inference(subsumption_resolution,[],[f5186,f140])).
fof(f5186,plain,(
  ( ! [X218,X216,X220,X219,X217,X221] : (p(s(bool,'F_0')) | ~p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X220) != s(bool,happ(s(fun(X216,bool),X217),s(X216,sK7(X218,X217,X216)))) | s(bool,X219) != s(bool,X221) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X216,bool),X217),s('type.list.list'(X216),X218)))) )),
  inference(superposition,[],[f105,f3599])).
fof(f5263,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (p(s(bool,X171)) | ~p(s(bool,X172)) | ~p(s(bool,X173)) | s(bool,X172) != s(bool,happ(s(fun(X168,bool),X169),s(X168,sK7(X170,X169,X168)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X168,bool),X169),s('type.list.list'(X168),X170)))) )),
  inference(subsumption_resolution,[],[f5178,f127])).
fof(f5178,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (p(s(bool,'F_0')) | p(s(bool,X171)) | ~p(s(bool,X172)) | ~p(s(bool,X173)) | s(bool,X172) != s(bool,happ(s(fun(X168,bool),X169),s(X168,sK7(X170,X169,X168)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X168,bool),X169),s('type.list.list'(X168),X170)))) )),
  inference(superposition,[],[f104,f3599])).
fof(f5261,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X141) != s(bool,X143) | s(bool,X142) != s(bool,happ(s(fun(X138,bool),X139),s(X138,sK7(X140,X139,X138)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X138,bool),X139),s('type.list.list'(X138),X140)))) )),
  inference(subsumption_resolution,[],[f5173,f127])).
fof(f5173,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,'F_0')) | p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X141) != s(bool,X143) | s(bool,X142) != s(bool,happ(s(fun(X138,bool),X139),s(X138,sK7(X140,X139,X138)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X138,bool),X139),s('type.list.list'(X138),X140)))) )),
  inference(superposition,[],[f103,f3599])).
fof(f5258,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (~p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X100) != s(bool,happ(s(fun(X96,bool),X97),s(X96,sK7(X98,X97,X96)))) | s(bool,X99) != s(bool,X101) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X96,bool),X97),s('type.list.list'(X96),X98)))) )),
  inference(subsumption_resolution,[],[f5166,f140])).
fof(f5166,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (p(s(bool,'F_0')) | ~p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X100) != s(bool,happ(s(fun(X96,bool),X97),s(X96,sK7(X98,X97,X96)))) | s(bool,X99) != s(bool,X101) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X96,bool),X97),s('type.list.list'(X96),X98)))) )),
  inference(superposition,[],[f101,f3599])).
fof(f5256,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,X59) | s(bool,X58) != s(bool,happ(s(fun(X54,bool),X55),s(X54,sK7(X56,X55,X54)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X54,bool),X55),s('type.list.list'(X54),X56)))) )),
  inference(subsumption_resolution,[],[f5159,f127])).
fof(f5159,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,'F_0')) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,X59) | s(bool,X58) != s(bool,happ(s(fun(X54,bool),X55),s(X54,sK7(X56,X55,X54)))) | s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X54,bool),X55),s('type.list.list'(X54),X56)))) )),
  inference(superposition,[],[f99,f3599])).
fof(f5121,plain,(
  ( ! [X356,X358,X360,X357,X359] : (s(bool,'F_0') = s(bool,X359) | s(bool,X360) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X356),X357))),s('type.bvlSem.v',X358))) | s(bool,'T_0') = s(bool,X360) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X356),X357),s('type.bvlSem.v',X358)))) ) | $spl14),
  inference(subsumption_resolution,[],[f5039,f116])).
fof(f5039,plain,(
  ( ! [X356,X358,X360,X357,X359] : (s(bool,'T_0') != s(bool,X359) | s(bool,'F_0') = s(bool,X359) | s(bool,X360) != s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X356),X357))),s('type.bvlSem.v',X358))) | s(bool,'T_0') = s(bool,X360) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X356),X357),s('type.bvlSem.v',X358)))) ) | $spl14),
  inference(superposition,[],[f1032,f3326])).
fof(f5077,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7))),s('type.list.list'('type.bvlSem.v'),X8)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X7),s('type.bvlSem.v',sK7(X8,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7)),'type.bvlSem.v'))))) ) | $spl132),
  inference(cnf_transformation,[],[f5077_D])).
fof(f5077_D,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7))),s('type.list.list'('type.bvlSem.v'),X8)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X7),s('type.bvlSem.v',sK7(X8,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7)),'type.bvlSem.v'))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f5076,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7))),s('type.list.list'('type.bvlSem.v'),X8)))) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X7),s('type.bvlSem.v',sK7(X8,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7)),'type.bvlSem.v'))))) ) | $spl88),
  inference(subsumption_resolution,[],[f4967,f112])).
fof(f4967,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7))),s('type.list.list'('type.bvlSem.v'),X8)))) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X6),X7),s('type.bvlSem.v',sK7(X8,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num',X6),X7)),'type.bvlSem.v'))))) ) | $spl88),
  inference(superposition,[],[f3315,f3326])).
fof(f4660,plain,(
  ( ! [X6,X10,X8,X7,X5] : (s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) | s(bool,X10) != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK7(X8,X7,X6)))) | ~p(s(bool,X10)) | p(s(bool,X5))) ) | $spl130),
  inference(cnf_transformation,[],[f4660_D])).
fof(f4660_D,plain,(
  ( ! [X6,X10,X8,X7,X5] : (s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) | s(bool,X10) != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK7(X8,X7,X6)))) | ~p(s(bool,X10)) | p(s(bool,X5))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f4642,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,'const.list.EVERY_2'(s(fun(X86,bool),X87),s('type.list.list'(X86),X88))) | p(s(bool,X85)) | p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,happ(s(fun(X86,bool),X87),s(X86,sK7(X88,X87,X86))))) ) | $spl86),
  inference(resolution,[],[f3314,f142])).
fof(f4635,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (s(bool,X54) != s(bool,'const.list.EVERY_2'(s(fun(X55,bool),X56),s('type.list.list'(X55),X57))) | p(s(bool,X54)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,happ(s(fun(X55,bool),X56),s(X55,sK7(X57,X56,X55)))) | s(bool,X58) != s(bool,X60)) ) | $spl86),
  inference(resolution,[],[f3314,f105])).
fof(f4633,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(X41,bool),X42),s('type.list.list'(X41),X43))) | p(s(bool,X40)) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,happ(s(fun(X41,bool),X42),s(X41,sK7(X43,X42,X41))))) ) | $spl86),
  inference(resolution,[],[f3314,f104])).
fof(f4631,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(X27,bool),X28),s('type.list.list'(X27),X29))) | p(s(bool,X26)) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(X27,bool),X28),s(X27,sK7(X29,X28,X27))))) ) | $spl86),
  inference(resolution,[],[f3314,f103])).
fof(f4630,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(X20,bool),X21),s('type.list.list'(X20),X22))) | p(s(bool,X19)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK7(X22,X21,X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl86),
  inference(resolution,[],[f3314,f101])).
fof(f4628,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) | p(s(bool,X5)) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK7(X8,X7,X6))))) ) | $spl86),
  inference(resolution,[],[f3314,f99])).
fof(f4623,plain,(
  ( ! [X6,X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool))) ) | $spl128),
  inference(cnf_transformation,[],[f4623_D])).
fof(f4623_D,plain,(
  ( ! [X6,X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f4610,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X55),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(X56,X55,bool)) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56)))) )),
  inference(resolution,[],[f3307,f142])).
fof(f3307,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X11),s('type.list.list'(bool),X10)))) | ~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(X10,X11,bool))) )),
  inference(superposition,[],[f96,f116])).
fof(f4603,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(X39,X38,bool)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3307,f105])).
fof(f4600,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(X24,X23,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X24)))) )),
  inference(resolution,[],[f3307,f104])).
fof(f4599,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(X19,X18,bool)) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) )),
  inference(resolution,[],[f3307,f103])).
fof(f4598,plain,(
  ( ! [X14,X17,X15,X13,X16] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(X14,X13,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3307,f101])).
fof(f4596,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(X4,X3,bool)) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) )),
  inference(resolution,[],[f3307,f99])).
fof(f4431,plain,(
  ( ! [X6,X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) | $spl126),
  inference(cnf_transformation,[],[f4431_D])).
fof(f4431_D,plain,(
  ( ! [X6,X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f4418,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X55),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(X56,X55,bool)) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56)))) )),
  inference(resolution,[],[f3306,f142])).
fof(f3306,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X9),s('type.list.list'(bool),X8)))) | ~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(X8,X9,bool))) )),
  inference(superposition,[],[f96,f116])).
fof(f4411,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(X39,X38,bool)) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3306,f105])).
fof(f4408,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(X24,X23,bool)) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X24)))) )),
  inference(resolution,[],[f3306,f104])).
fof(f4407,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(X19,X18,bool)) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) )),
  inference(resolution,[],[f3306,f103])).
fof(f4406,plain,(
  ( ! [X14,X17,X15,X13,X16] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(X14,X13,bool)) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3306,f101])).
fof(f4404,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(X4,X3,bool)) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) )),
  inference(resolution,[],[f3306,f99])).
fof(f4399,plain,(
  ( ! [X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK6)))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X3,'type.bvlSem.v'))))) ) | $spl124),
  inference(cnf_transformation,[],[f4399_D])).
fof(f4399_D,plain,(
  ( ! [X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK6)))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X3,'type.bvlSem.v'))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f4386,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X41),s('type.list.list'('type.bvlSem.v'),sK6)))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X41,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f1785,f142])).
fof(f4379,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X31),s('type.list.list'('type.bvlSem.v'),sK6)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X31,'type.bvlSem.v')))) | s(bool,X33) != s(bool,X34)) )),
  inference(resolution,[],[f1785,f105])).
fof(f4376,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X19),s('type.list.list'('type.bvlSem.v'),sK6)))) | p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X19,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f1785,f104])).
fof(f4375,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X15),s('type.list.list'('type.bvlSem.v'),sK6)))) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X15,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f1785,f103])).
fof(f4374,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X11),s('type.list.list'('type.bvlSem.v'),sK6)))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X11,'type.bvlSem.v')))) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1785,f101])).
fof(f4372,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),X3),s('type.list.list'('type.bvlSem.v'),sK6)))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',sK7(sK6,X3,'type.bvlSem.v'))))) )),
  inference(resolution,[],[f1785,f99])).
fof(f2793,plain,(
  ( ! [X5,X3] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X3),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X3)))) ) | $spl82),
  inference(cnf_transformation,[],[f2793_D])).
fof(f2793_D,plain,(
  ( ! [X5,X3] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X3),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X3)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4355,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X41),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X41)))) ) | $spl6),
  inference(resolution,[],[f2269,f142])).
fof(f4348,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X31),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X31))) | s(bool,X33) != s(bool,X34)) ) | $spl6),
  inference(resolution,[],[f2269,f105])).
fof(f4345,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X19),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X19)))) ) | $spl6),
  inference(resolution,[],[f2269,f104])).
fof(f4344,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X15),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X15)))) ) | $spl6),
  inference(resolution,[],[f2269,f103])).
fof(f4343,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X11),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X11))) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f2269,f101])).
fof(f4341,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X3),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X3)))) ) | $spl6),
  inference(resolution,[],[f2269,f99])).
fof(f1810,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | ~p(s(bool,X3)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X1))))) ) | $spl56),
  inference(cnf_transformation,[],[f1810_D])).
fof(f1810_D,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | ~p(s(bool,X3)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X1))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4328,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X39)))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X39),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f142])).
fof(f4321,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X29)))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X29),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f539,f105])).
fof(f4319,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X21)))) | p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X21),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f104])).
fof(f4317,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X13)))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X13),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f103])).
fof(f4316,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X9)))) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X9),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X10) != s(bool,X12)) )),
  inference(resolution,[],[f539,f101])).
fof(f4314,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X1)))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f99])).
fof(f4223,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl122),
  inference(cnf_transformation,[],[f4223_D])).
fof(f4223_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f4184,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X80) != s(bool,X82) | s(bool,X82) != s(bool,X83) | s(bool,'T_0') = s(bool,X83) | p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X81) != s(bool,X85)) ) | $spl50),
  inference(resolution,[],[f1606,f142])).
fof(f4177,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl50),
  inference(resolution,[],[f1606,f105])).
fof(f4175,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,X35) != s(bool,X36) | s(bool,X35) != s(bool,X37) | s(bool,X37) != s(bool,X38) | s(bool,'T_0') = s(bool,X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl50),
  inference(resolution,[],[f1606,f104])).
fof(f4173,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl50),
  inference(resolution,[],[f1606,f103])).
fof(f4172,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl50),
  inference(resolution,[],[f1606,f101])).
fof(f4170,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl50),
  inference(resolution,[],[f1606,f99])).
fof(f4108,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X11)))) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X14))) | ~p(s(bool,X17)) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X10),X14))))),s('type.list.list'('type.bvlSem.v'),X13)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X12),s('type.list.list'('type.bvlSem.v'),X13),s('type.bvlSem.state'(X10),X14))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'const.pair.,_2'(s('type.bvlSem.v',X15),s('type.bvlSem.state'(X10),X11)))))) ) | $spl120),
  inference(cnf_transformation,[],[f4108_D])).
fof(f4108_D,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X11)))) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X14))) | ~p(s(bool,X17)) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X10),X14))))),s('type.list.list'('type.bvlSem.v'),X13)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X12),s('type.list.list'('type.bvlSem.v'),X13),s('type.bvlSem.state'(X10),X14))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'const.pair.,_2'(s('type.bvlSem.v',X15),s('type.bvlSem.state'(X10),X11)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f4073,plain,(
  ( ! [X125,X123,X121,X118,X124,X122,X120,X119] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X118),X119)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X118)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X120),s('type.list.list'('type.bvlSem.v'),X121),s('type.bvlSem.state'(X118),X122))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X118)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X118)),'const.pair.,_2'(s('type.bvlSem.v',X123),s('type.bvlSem.state'(X118),X119))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X118),X122))))),s('type.list.list'('type.bvlSem.v'),X121)))) | p(s(bool,X124)) | ~p(s(bool,X125)) | s(bool,X125) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X118),X122)))) )),
  inference(resolution,[],[f92,f142])).
fof(f4066,plain,(
  ( ! [X80,X78,X76,X74,X81,X79,X77,X75,X73] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X73),X74)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X73)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X75),s('type.list.list'('type.bvlSem.v'),X76),s('type.bvlSem.state'(X73),X77))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X73)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X73)),'const.pair.,_2'(s('type.bvlSem.v',X78),s('type.bvlSem.state'(X73),X74))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X73),X77))))),s('type.list.list'('type.bvlSem.v'),X76)))) | ~p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X73),X77))) | s(bool,X79) != s(bool,X81)) )),
  inference(resolution,[],[f92,f105])).
fof(f4064,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55,X63] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X55),X56)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X55)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X57),s('type.list.list'('type.bvlSem.v'),X58),s('type.bvlSem.state'(X55),X59))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X55)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X55)),'const.pair.,_2'(s('type.bvlSem.v',X60),s('type.bvlSem.state'(X55),X56))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X55),X59))))),s('type.list.list'('type.bvlSem.v'),X58)))) | p(s(bool,X61)) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X55),X59)))) )),
  inference(resolution,[],[f92,f104])).
fof(f4062,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X44,X42,X40] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X37),X38)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X37)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X39),s('type.list.list'('type.bvlSem.v'),X40),s('type.bvlSem.state'(X37),X41))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X37)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X37)),'const.pair.,_2'(s('type.bvlSem.v',X42),s('type.bvlSem.state'(X37),X38))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X37),X41))))),s('type.list.list'('type.bvlSem.v'),X40)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X37),X41)))) )),
  inference(resolution,[],[f92,f103])).
fof(f4061,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X36,X34,X32] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X28),X29)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X28)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X30),s('type.list.list'('type.bvlSem.v'),X31),s('type.bvlSem.state'(X28),X32))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X28)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X28)),'const.pair.,_2'(s('type.bvlSem.v',X33),s('type.bvlSem.state'(X28),X29))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X28),X32))))),s('type.list.list'('type.bvlSem.v'),X31)))) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X28),X32))) | s(bool,X34) != s(bool,X36)) )),
  inference(resolution,[],[f92,f101])).
fof(f4059,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X11)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X12),s('type.list.list'('type.bvlSem.v'),X13),s('type.bvlSem.state'(X10),X14))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'const.pair.,_2'(s('type.bvlSem.v',X15),s('type.bvlSem.state'(X10),X11))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X10),X14))))),s('type.list.list'('type.bvlSem.v'),X13)))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X14)))) )),
  inference(resolution,[],[f92,f99])).
fof(f4107,plain,(
  ( ! [X23,X21,X19,X25,X22,X20,X18] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X18),X19)))) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X18),X22))))),s('type.list.list'('type.bvlSem.v'),X21))) | ~p(s(bool,X25)) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X18),X22)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X20),s('type.list.list'('type.bvlSem.v'),X21),s('type.bvlSem.state'(X18),X22))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'const.pair.,_2'(s('type.bvlSem.v',X23),s('type.bvlSem.state'(X18),X19)))))) ) | $spl118),
  inference(cnf_transformation,[],[f4107_D])).
fof(f4107_D,plain,(
  ( ! [X23,X21,X19,X25,X22,X20,X18] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X18),X19)))) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X18),X22))))),s('type.list.list'('type.bvlSem.v'),X21))) | ~p(s(bool,X25)) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X18),X22)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X20),s('type.list.list'('type.bvlSem.v'),X21),s('type.bvlSem.state'(X18),X22))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'const.pair.,_2'(s('type.bvlSem.v',X23),s('type.bvlSem.state'(X18),X19)))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4053,plain,(
  ( ! [X127,X132,X130,X128,X126,X133,X131,X129] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X126),X127)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X126)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X128),s('type.list.list'('type.bvlSem.v'),X129),s('type.bvlSem.state'(X126),X130))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X126)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X126)),'const.pair.,_2'(s('type.bvlSem.v',X131),s('type.bvlSem.state'(X126),X127))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X126),X130)))) | p(s(bool,X132)) | ~p(s(bool,X133)) | s(bool,X133) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X126),X130))))),s('type.list.list'('type.bvlSem.v'),X129)))) )),
  inference(resolution,[],[f92,f142])).
fof(f4046,plain,(
  ( ! [X88,X87,X85,X83,X81,X89,X86,X84,X82] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X81),X82)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X81)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X83),s('type.list.list'('type.bvlSem.v'),X84),s('type.bvlSem.state'(X81),X85))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X81)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X81)),'const.pair.,_2'(s('type.bvlSem.v',X86),s('type.bvlSem.state'(X81),X82))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X81),X85)))) | ~p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X88) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X81),X85))))),s('type.list.list'('type.bvlSem.v'),X84))) | s(bool,X87) != s(bool,X89)) )),
  inference(resolution,[],[f92,f105])).
fof(f4044,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X63),X64)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X63)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X65),s('type.list.list'('type.bvlSem.v'),X66),s('type.bvlSem.state'(X63),X67))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X63)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X63)),'const.pair.,_2'(s('type.bvlSem.v',X68),s('type.bvlSem.state'(X63),X64))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X63),X67)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X63),X67))))),s('type.list.list'('type.bvlSem.v'),X66)))) )),
  inference(resolution,[],[f92,f104])).
fof(f4042,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X53,X51,X49] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X45),X46)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X45)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X47),s('type.list.list'('type.bvlSem.v'),X48),s('type.bvlSem.state'(X45),X49))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X45)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X45)),'const.pair.,_2'(s('type.bvlSem.v',X50),s('type.bvlSem.state'(X45),X46))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X45),X49)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X45),X49))))),s('type.list.list'('type.bvlSem.v'),X48)))) )),
  inference(resolution,[],[f92,f103])).
fof(f4041,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X36),X37)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X36)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X38),s('type.list.list'('type.bvlSem.v'),X39),s('type.bvlSem.state'(X36),X40))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X36)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X36)),'const.pair.,_2'(s('type.bvlSem.v',X41),s('type.bvlSem.state'(X36),X37))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X36),X40)))) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X43) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X36),X40))))),s('type.list.list'('type.bvlSem.v'),X39))) | s(bool,X42) != s(bool,X44)) )),
  inference(resolution,[],[f92,f101])).
fof(f4039,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X18),X19)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X20),s('type.list.list'('type.bvlSem.v'),X21),s('type.bvlSem.state'(X18),X22))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X18)),'const.pair.,_2'(s('type.bvlSem.v',X23),s('type.bvlSem.state'(X18),X19))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X18),X22)))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X18),X22))))),s('type.list.list'('type.bvlSem.v'),X21)))) )),
  inference(resolution,[],[f92,f99])).
fof(f4104,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X11),s('type.list.list'('type.bvlSem.v'),X12),s('type.bvlSem.state'(X10),X13))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'const.pair.,_2'(s('type.bvlSem.v',X14),s('type.bvlSem.state'(X10),X15))))) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X15))) | p(s(bool,X17)) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X13)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X10),X13))))),s('type.list.list'('type.bvlSem.v'),X12))))) ) | $spl116),
  inference(cnf_transformation,[],[f4104_D])).
fof(f4104_D,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X11),s('type.list.list'('type.bvlSem.v'),X12),s('type.bvlSem.state'(X10),X13))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'const.pair.,_2'(s('type.bvlSem.v',X14),s('type.bvlSem.state'(X10),X15))))) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X15))) | p(s(bool,X17)) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X13)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X10),X13))))),s('type.list.list'('type.bvlSem.v'),X12))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4030,plain,(
  ( ! [X125,X123,X121,X118,X124,X122,X120,X119] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X118)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X119),s('type.list.list'('type.bvlSem.v'),X120),s('type.bvlSem.state'(X118),X121))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X118)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X118)),'const.pair.,_2'(s('type.bvlSem.v',X122),s('type.bvlSem.state'(X118),X123))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X118),X121))))),s('type.list.list'('type.bvlSem.v'),X120)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X118),X121)))) | p(s(bool,X124)) | p(s(bool,X125)) | s(bool,X125) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X118),X123)))) )),
  inference(resolution,[],[f92,f142])).
fof(f4023,plain,(
  ( ! [X80,X78,X76,X74,X81,X79,X77,X75,X73] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X73)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X74),s('type.list.list'('type.bvlSem.v'),X75),s('type.bvlSem.state'(X73),X76))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X73)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X73)),'const.pair.,_2'(s('type.bvlSem.v',X77),s('type.bvlSem.state'(X73),X78))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X73),X76))))),s('type.list.list'('type.bvlSem.v'),X75)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X73),X76)))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X79) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X73),X78))) | s(bool,X80) != s(bool,X81)) )),
  inference(resolution,[],[f92,f105])).
fof(f4020,plain,(
  ( ! [X47,X54,X52,X50,X48,X46,X53,X51,X49] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X46)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X47),s('type.list.list'('type.bvlSem.v'),X48),s('type.bvlSem.state'(X46),X49))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X46)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X46)),'const.pair.,_2'(s('type.bvlSem.v',X50),s('type.bvlSem.state'(X46),X51))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X46),X49))))),s('type.list.list'('type.bvlSem.v'),X48)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X46),X49)))) | p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X54) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X46),X51)))) )),
  inference(resolution,[],[f92,f104])).
fof(f4019,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X44,X42,X40] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X37)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X38),s('type.list.list'('type.bvlSem.v'),X39),s('type.bvlSem.state'(X37),X40))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X37)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X37)),'const.pair.,_2'(s('type.bvlSem.v',X41),s('type.bvlSem.state'(X37),X42))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X37),X40))))),s('type.list.list'('type.bvlSem.v'),X39)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X37),X40)))) | p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X37),X42)))) )),
  inference(resolution,[],[f92,f103])).
fof(f4018,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X36,X34,X32] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X28)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X29),s('type.list.list'('type.bvlSem.v'),X30),s('type.bvlSem.state'(X28),X31))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X28)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X28)),'const.pair.,_2'(s('type.bvlSem.v',X32),s('type.bvlSem.state'(X28),X33))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X28),X31))))),s('type.list.list'('type.bvlSem.v'),X30)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X28),X31)))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X28),X33))) | s(bool,X35) != s(bool,X36)) )),
  inference(resolution,[],[f92,f101])).
fof(f4016,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X11),s('type.list.list'('type.bvlSem.v'),X12),s('type.bvlSem.state'(X10),X13))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X10)),'const.pair.,_2'(s('type.bvlSem.v',X14),s('type.bvlSem.state'(X10),X15))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X10),X13))))),s('type.list.list'('type.bvlSem.v'),X12)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X13)))) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X10),X15)))) )),
  inference(resolution,[],[f92,f99])).
fof(f4011,plain,(
  ( ! [X14,X12,X10,X8,X7,X11,X9] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X8),s('type.list.list'('type.bvlSem.v'),X9),s('type.bvlSem.state'(X7),X10))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'const.pair.,_2'(s('type.bvlSem.v',X11),s('type.bvlSem.state'(X7),X12))))) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X7),X10))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X7),X12))))))) | p(s(bool,X14))) ) | $spl114),
  inference(cnf_transformation,[],[f4011_D])).
fof(f4011_D,plain,(
  ( ! [X14,X12,X10,X8,X7,X11,X9] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X8),s('type.list.list'('type.bvlSem.v'),X9),s('type.bvlSem.state'(X7),X10))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'const.pair.,_2'(s('type.bvlSem.v',X11),s('type.bvlSem.state'(X7),X12))))) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X7),X10))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X7),X12))))))) | p(s(bool,X14))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f3998,plain,(
  ( ! [X121,X118,X116,X122,X120,X119,X117,X115] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X115)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X116),s('type.list.list'('type.bvlSem.v'),X117),s('type.bvlSem.state'(X115),X118))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X115)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X115)),'const.pair.,_2'(s('type.bvlSem.v',X119),s('type.bvlSem.state'(X115),X120))))) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X122) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X115),X118))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X115),X120)))))))) )),
  inference(resolution,[],[f117,f142])).
fof(f3991,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X77,X75,X73] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X70)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X71),s('type.list.list'('type.bvlSem.v'),X72),s('type.bvlSem.state'(X70),X73))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X70)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X70)),'const.pair.,_2'(s('type.bvlSem.v',X74),s('type.bvlSem.state'(X70),X75))))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X76) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X70),X73))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X70),X75))))))) | s(bool,X77) != s(bool,X78)) )),
  inference(resolution,[],[f117,f105])).
fof(f3988,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X51,X49] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X43)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X44),s('type.list.list'('type.bvlSem.v'),X45),s('type.bvlSem.state'(X43),X46))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X43)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X43)),'const.pair.,_2'(s('type.bvlSem.v',X47),s('type.bvlSem.state'(X43),X48))))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X43),X46))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X43),X48)))))))) )),
  inference(resolution,[],[f117,f104])).
fof(f3987,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X34)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X35),s('type.list.list'('type.bvlSem.v'),X36),s('type.bvlSem.state'(X34),X37))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X34)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X34)),'const.pair.,_2'(s('type.bvlSem.v',X38),s('type.bvlSem.state'(X34),X39))))) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X34),X37))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X34),X39)))))))) )),
  inference(resolution,[],[f117,f103])).
fof(f3986,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X25)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X26),s('type.list.list'('type.bvlSem.v'),X27),s('type.bvlSem.state'(X25),X28))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X25)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X25)),'const.pair.,_2'(s('type.bvlSem.v',X29),s('type.bvlSem.state'(X25),X30))))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X25),X28))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X25),X30))))))) | s(bool,X32) != s(bool,X33)) )),
  inference(resolution,[],[f117,f101])).
fof(f3984,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X8),s('type.list.list'('type.bvlSem.v'),X9),s('type.bvlSem.state'(X7),X10))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X7)),'const.pair.,_2'(s('type.bvlSem.v',X11),s('type.bvlSem.state'(X7),X12))))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X7),X10))))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X7),X12)))))))) )),
  inference(resolution,[],[f117,f99])).
fof(f3981,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | s(bool,X3) != s(bool,sK7(X1,X0,bool)) | ~p(s(bool,X3)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1))))) ) | $spl112),
  inference(cnf_transformation,[],[f3981_D])).
fof(f3981_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | s(bool,X3) != s(bool,sK7(X1,X0,bool)) | ~p(s(bool,X3)) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f3970,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X52),s('type.list.list'(bool),X53)))) | ~p(s(bool,happ(s(fun(bool,bool),X52),s(bool,'T_0')))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,sK7(X53,X52,bool))) )),
  inference(resolution,[],[f3304,f142])).
fof(f3963,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X35),s('type.list.list'(bool),X36)))) | ~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,sK7(X36,X35,bool)) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f3304,f105])).
fof(f3961,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X25),s('type.list.list'(bool),X26)))) | ~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0')))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,sK7(X26,X25,bool))) )),
  inference(resolution,[],[f3304,f104])).
fof(f3959,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X15),s('type.list.list'(bool),X16)))) | ~p(s(bool,happ(s(fun(bool,bool),X15),s(bool,'T_0')))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,sK7(X16,X15,bool))) )),
  inference(resolution,[],[f3304,f103])).
fof(f3958,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X10),s('type.list.list'(bool),X11)))) | ~p(s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0')))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,sK7(X11,X10,bool)) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f3304,f101])).
fof(f3956,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X0),s('type.list.list'(bool),X1)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,sK7(X1,X0,bool))) )),
  inference(resolution,[],[f3304,f99])).
fof(f3952,plain,(
  ( ! [X6,X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | p(s(bool,X6)) | p(s(bool,sK7(X4,X3,bool)))) ) | $spl110),
  inference(cnf_transformation,[],[f3952_D])).
fof(f3952_D,plain,(
  ( ! [X6,X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4))) | p(s(bool,X6)) | p(s(bool,sK7(X4,X3,bool)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f3940,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X55),s(bool,'F_0')))) | p(s(bool,sK7(X56,X55,bool))) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X55),s('type.list.list'(bool),X56)))) )),
  inference(resolution,[],[f3302,f142])).
fof(f3933,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,'F_0')))) | p(s(bool,sK7(X39,X38,bool))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X38),s('type.list.list'(bool),X39))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f3302,f105])).
fof(f3930,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,'F_0')))) | p(s(bool,sK7(X24,X23,bool))) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X23),s('type.list.list'(bool),X24)))) )),
  inference(resolution,[],[f3302,f104])).
fof(f3929,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,'F_0')))) | p(s(bool,sK7(X19,X18,bool))) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X18),s('type.list.list'(bool),X19)))) )),
  inference(resolution,[],[f3302,f103])).
fof(f3928,plain,(
  ( ! [X14,X17,X15,X13,X16] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,sK7(X14,X13,bool))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X13),s('type.list.list'(bool),X14))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3302,f101])).
fof(f3926,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | p(s(bool,sK7(X4,X3,bool))) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X3),s('type.list.list'(bool),X4)))) )),
  inference(resolution,[],[f3302,f99])).
fof(f3918,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8)))))) | s(bool,X13) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9))) | ~p(s(bool,X13)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9))))) ) | $spl108),
  inference(cnf_transformation,[],[f3918_D])).
fof(f3918_D,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8)))))) | s(bool,X13) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9))) | ~p(s(bool,X13)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f3891,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X101),X102),s('type.bvlSem.v',X103)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X104),X105))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X101),X102)))))) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X107) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X104),X105),s('type.bvlSem.v',X103)))) )),
  inference(resolution,[],[f118,f142])).
fof(f3884,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X63),X64),s('type.bvlSem.v',X65)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X66),X67))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X63),X64)))))) | ~p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X69) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X66),X67),s('type.bvlSem.v',X65))) | s(bool,X68) != s(bool,X70)) )),
  inference(resolution,[],[f118,f105])).
fof(f3882,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X47),X48),s('type.bvlSem.v',X49)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X50),X51))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X47),X48)))))) | p(s(bool,X52)) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X50),X51),s('type.bvlSem.v',X49)))) )),
  inference(resolution,[],[f118,f104])).
fof(f3880,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X31),X32),s('type.bvlSem.v',X33)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X34),X35))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X31),X32)))))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X34),X35),s('type.bvlSem.v',X33)))) )),
  inference(resolution,[],[f118,f103])).
fof(f3879,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X23),X24),s('type.bvlSem.v',X25)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X26),X27))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X23),X24)))))) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X26),X27),s('type.bvlSem.v',X25))) | s(bool,X28) != s(bool,X30)) )),
  inference(resolution,[],[f118,f101])).
fof(f3877,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X7),X8),s('type.bvlSem.v',X9)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X10),X11))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X7),X8)))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X10),X11),s('type.bvlSem.v',X9)))) )),
  inference(resolution,[],[f118,f99])).
fof(f3917,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X3),X4),s('type.bvlSem.v',X2)))) | s(bool,X6) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X3),X4))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))))) | ~p(s(bool,X6)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',X2))))) ) | $spl106),
  inference(cnf_transformation,[],[f3917_D])).
fof(f3917_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X3),X4),s('type.bvlSem.v',X2)))) | s(bool,X6) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X3),X4))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))))) | ~p(s(bool,X6)) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',X2))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f3871,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X94),X95),s('type.bvlSem.v',X96)))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X97),X98),s('type.bvlSem.v',X96)))) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X100) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X97),X98))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X94),X95)))))) )),
  inference(resolution,[],[f118,f142])).
fof(f3864,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X56),X57),s('type.bvlSem.v',X58)))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X59),X60),s('type.bvlSem.v',X58)))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X59),X60))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X56),X57))))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f118,f105])).
fof(f3862,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X40),X41),s('type.bvlSem.v',X42)))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X43),X44),s('type.bvlSem.v',X42)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X43),X44))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X40),X41)))))) )),
  inference(resolution,[],[f118,f104])).
fof(f3860,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X24),X25),s('type.bvlSem.v',X26)))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X27),X28),s('type.bvlSem.v',X26)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X27),X28))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X24),X25)))))) )),
  inference(resolution,[],[f118,f103])).
fof(f3859,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X16),X17),s('type.bvlSem.v',X18)))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X19),X20),s('type.bvlSem.v',X18)))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X19),X20))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X16),X17))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f118,f101])).
fof(f3857,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',X2)))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X3),X4),s('type.bvlSem.v',X2)))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X3),X4))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X1)))))) )),
  inference(resolution,[],[f118,f99])).
fof(f3916,plain,(
  ( ! [X12,X10,X15,X13,X11,X9] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X11),X12)))))) | s(bool,X15) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X11),X12),s('type.bvlSem.v',X13))) | p(s(bool,X15)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X13))))) ) | $spl104),
  inference(cnf_transformation,[],[f3916_D])).
fof(f3916_D,plain,(
  ( ! [X12,X10,X15,X13,X11,X9] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X11),X12)))))) | s(bool,X15) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X11),X12),s('type.bvlSem.v',X13))) | p(s(bool,X15)) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X13))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3851,plain,(
  ( ! [X103,X109,X107,X105,X108,X106,X104] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X103),X104))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X105),X106)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X103),X104),s('type.bvlSem.v',X107)))) | p(s(bool,X108)) | p(s(bool,X109)) | s(bool,X109) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X105),X106),s('type.bvlSem.v',X107)))) )),
  inference(resolution,[],[f118,f142])).
fof(f3844,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X65),X66))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X67),X68)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X65),X66),s('type.bvlSem.v',X69)))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X67),X68),s('type.bvlSem.v',X69))) | s(bool,X71) != s(bool,X72)) )),
  inference(resolution,[],[f118,f105])).
fof(f3841,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X41),X42))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X43),X44)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X41),X42),s('type.bvlSem.v',X45)))) | p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X43),X44),s('type.bvlSem.v',X45)))) )),
  inference(resolution,[],[f118,f104])).
fof(f3840,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X40] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X33),X34))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X35),X36)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X33),X34),s('type.bvlSem.v',X37)))) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X35),X36),s('type.bvlSem.v',X37)))) )),
  inference(resolution,[],[f118,f103])).
fof(f3839,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25,X32] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X25),X26))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X27),X28)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X25),X26),s('type.bvlSem.v',X29)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X27),X28),s('type.bvlSem.v',X29))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f118,f101])).
fof(f3837,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9,X16] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X9),X10))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X11),X12)))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X9),X10),s('type.bvlSem.v',X13)))) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X11),X12),s('type.bvlSem.v',X13)))) )),
  inference(resolution,[],[f118,f99])).
fof(f3814,plain,(
  ( ! [X6,X4,X8,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,sK7(X6,X5,X4)))) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),X6))) | p(s(bool,X8))) ) | $spl102),
  inference(cnf_transformation,[],[f3814_D])).
fof(f3814_D,plain,(
  ( ! [X6,X4,X8,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,sK7(X6,X5,X4)))) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),X6))) | p(s(bool,X8))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f3799,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(X70,bool),X71),s(X70,sK7(X72,X71,X70)))) | p(s(bool,X73)) | p(s(bool,X74)) | s(bool,X74) != s(bool,'const.list.EVERY_2'(s(fun(X70,bool),X71),s('type.list.list'(X70),X72)))) )),
  inference(resolution,[],[f3282,f142])).
fof(f3792,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,'F_0') = s(bool,happ(s(fun(X46,bool),X47),s(X46,sK7(X48,X47,X46)))) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,'const.list.EVERY_2'(s(fun(X46,bool),X47),s('type.list.list'(X46),X48))) | s(bool,X50) != s(bool,X51)) )),
  inference(resolution,[],[f3282,f105])).
fof(f3789,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X28,bool),X29),s(X28,sK7(X30,X29,X28)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.list.EVERY_2'(s(fun(X28,bool),X29),s('type.list.list'(X28),X30)))) )),
  inference(resolution,[],[f3282,f104])).
fof(f3788,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'F_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,sK7(X24,X23,X22)))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(X22,bool),X23),s('type.list.list'(X22),X24)))) )),
  inference(resolution,[],[f3282,f103])).
fof(f3787,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,sK7(X18,X17,X16)))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(X16,bool),X17),s('type.list.list'(X16),X18))) | s(bool,X20) != s(bool,X21)) )),
  inference(resolution,[],[f3282,f101])).
fof(f3785,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,sK7(X6,X5,X4)))) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),X6)))) )),
  inference(resolution,[],[f3282,f99])).
fof(f3755,plain,(
  ( ! [X10,X8,X13,X11,X9] : (s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(X8,bool),X9),s('type.list.list'(X8),X11))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X13)) | ~p(s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X8),X11))))))) ) | $spl100),
  inference(cnf_transformation,[],[f3755_D])).
fof(f3755_D,plain,(
  ( ! [X10,X8,X13,X11,X9] : (s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(X8,bool),X9),s('type.list.list'(X8),X11))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X13)) | ~p(s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X8),X11))))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f3683,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (p(s(bool,happ(s(fun(X88,bool),X89),s(X88,X90)))) | ~p(s(bool,'const.bool.IN_2'(s(X88,X90),s(fun(X88,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X88),X91)))))) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.list.EVERY_2'(s(fun(X88,bool),X89),s('type.list.list'(X88),X91)))) )),
  inference(resolution,[],[f94,f142])).
fof(f3676,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (p(s(bool,happ(s(fun(X57,bool),X58),s(X57,X59)))) | ~p(s(bool,'const.bool.IN_2'(s(X57,X59),s(fun(X57,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X57),X60)))))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.list.EVERY_2'(s(fun(X57,bool),X58),s('type.list.list'(X57),X60))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f94,f105])).
fof(f3674,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (p(s(bool,happ(s(fun(X43,bool),X44),s(X43,X45)))) | ~p(s(bool,'const.bool.IN_2'(s(X43,X45),s(fun(X43,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X43),X46)))))) | p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,'const.list.EVERY_2'(s(fun(X43,bool),X44),s('type.list.list'(X43),X46)))) )),
  inference(resolution,[],[f94,f104])).
fof(f3672,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X29,bool),X30),s(X29,X31)))) | ~p(s(bool,'const.bool.IN_2'(s(X29,X31),s(fun(X29,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X29),X32)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(X29,bool),X30),s('type.list.list'(X29),X32)))) )),
  inference(resolution,[],[f94,f103])).
fof(f3671,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (p(s(bool,happ(s(fun(X22,bool),X23),s(X22,X24)))) | ~p(s(bool,'const.bool.IN_2'(s(X22,X24),s(fun(X22,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X22),X25)))))) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun(X22,bool),X23),s('type.list.list'(X22),X25))) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f94,f101])).
fof(f3669,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X8),X11)))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(X8,bool),X9),s('type.list.list'(X8),X11)))) )),
  inference(resolution,[],[f94,f99])).
fof(f3754,plain,(
  ( ! [X6,X4,X7,X5,X9] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),X7)))) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(X4,X6),s(fun(X4,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X4),X7))))) | ~p(s(bool,X9)) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl98),
  inference(cnf_transformation,[],[f3754_D])).
fof(f3754_D,plain,(
  ( ! [X6,X4,X7,X5,X9] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),X7)))) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(X4,X6),s(fun(X4,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X4),X7))))) | ~p(s(bool,X9)) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f3662,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (p(s(bool,happ(s(fun(X84,bool),X85),s(X84,X86)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X84,bool),X85),s('type.list.list'(X84),X87)))) | p(s(bool,X88)) | ~p(s(bool,X89)) | s(bool,X89) != s(bool,'const.bool.IN_2'(s(X84,X86),s(fun(X84,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X84),X87)))))) )),
  inference(resolution,[],[f94,f142])).
fof(f3655,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X54),s(X53,X55)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X53,bool),X54),s('type.list.list'(X53),X56)))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X53),X56))))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f94,f105])).
fof(f3653,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X39,bool),X40),s('type.list.list'(X39),X42)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(X39,X41),s(fun(X39,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X39),X42)))))) )),
  inference(resolution,[],[f94,f104])).
fof(f3651,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X25,bool),X26),s('type.list.list'(X25),X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(X25,X27),s(fun(X25,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X25),X28)))))) )),
  inference(resolution,[],[f94,f103])).
fof(f3650,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (p(s(bool,happ(s(fun(X18,bool),X19),s(X18,X20)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X18,bool),X19),s('type.list.list'(X18),X21)))) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(X18,X20),s(fun(X18,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X18),X21))))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f94,f101])).
fof(f3648,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),X7)))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(X4,X6),s(fun(X4,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X4),X7)))))) )),
  inference(resolution,[],[f94,f99])).
fof(f3753,plain,(
  ( ! [X10,X8,X13,X11,X9] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X8,bool),X11),s('type.list.list'(X8),X10)))) | s(bool,X13) != s(bool,happ(s(fun(X8,bool),X11),s(X8,X9))) | p(s(bool,X13)) | ~p(s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X8),X10))))))) ) | $spl96),
  inference(cnf_transformation,[],[f3753_D])).
fof(f3753_D,plain,(
  ( ! [X10,X8,X13,X11,X9] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X8,bool),X11),s('type.list.list'(X8),X10)))) | s(bool,X13) != s(bool,happ(s(fun(X8,bool),X11),s(X8,X9))) | p(s(bool,X13)) | ~p(s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X8),X10))))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f3641,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (~p(s(bool,'const.bool.IN_2'(s(X88,X89),s(fun(X88,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X88),X90)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X88,bool),X91),s('type.list.list'(X88),X90)))) | p(s(bool,X92)) | p(s(bool,X93)) | s(bool,X93) != s(bool,happ(s(fun(X88,bool),X91),s(X88,X89)))) )),
  inference(resolution,[],[f94,f142])).
fof(f3634,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.bool.IN_2'(s(X57,X58),s(fun(X57,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X57),X59)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X57,bool),X60),s('type.list.list'(X57),X59)))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,happ(s(fun(X57,bool),X60),s(X57,X58))) | s(bool,X62) != s(bool,X63)) )),
  inference(resolution,[],[f94,f105])).
fof(f3631,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X36),X38)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X36,bool),X39),s('type.list.list'(X36),X38)))) | p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X39),s(X36,X37)))) )),
  inference(resolution,[],[f94,f104])).
fof(f3630,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X29,X30),s(fun(X29,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X29),X31)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X29,bool),X32),s('type.list.list'(X29),X31)))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,happ(s(fun(X29,bool),X32),s(X29,X30)))) )),
  inference(resolution,[],[f94,f103])).
fof(f3629,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(X22,X23),s(fun(X22,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X22),X24)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X22,bool),X25),s('type.list.list'(X22),X24)))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,happ(s(fun(X22,bool),X25),s(X22,X23))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f94,f101])).
fof(f3627,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X8),X10)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X8,bool),X11),s('type.list.list'(X8),X10)))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(X8,bool),X11),s(X8,X9)))) )),
  inference(resolution,[],[f94,f99])).
fof(f3605,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X66,bool),X67),s('type.list.list'(X66),X68))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,happ(s(fun(X66,bool),X67),s(X66,sK7(X68,X67,X66))))) )),
  inference(resolution,[],[f3264,f142])).
fof(f3598,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X42,bool),X43),s('type.list.list'(X42),X44))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(X42,bool),X43),s(X42,sK7(X44,X43,X42)))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f3264,f105])).
fof(f3596,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X30,bool),X31),s('type.list.list'(X30),X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,happ(s(fun(X30,bool),X31),s(X30,sK7(X32,X31,X30))))) )),
  inference(resolution,[],[f3264,f104])).
fof(f3594,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X18,bool),X19),s('type.list.list'(X18),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(X18,bool),X19),s(X18,sK7(X20,X19,X18))))) )),
  inference(resolution,[],[f3264,f103])).
fof(f3593,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X12,bool),X13),s('type.list.list'(X12),X14))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK7(X14,X13,X12)))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3264,f101])).
fof(f3591,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) )),
  inference(resolution,[],[f3264,f99])).
fof(f2751,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl80),
  inference(cnf_transformation,[],[f2751_D])).
fof(f2751_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3565,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f2232,f142])).
fof(f3558,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f2232,f105])).
fof(f3556,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f2232,f104])).
fof(f3554,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f2232,f103])).
fof(f3553,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f2232,f101])).
fof(f3551,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f2232,f99])).
fof(f3542,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,X4))) ) | $spl92),
  inference(cnf_transformation,[],[f3542_D])).
fof(f3542_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,X4))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3513,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,'const.list.EVERY_2'(s(fun(X66,bool),X67),s('type.list.list'(X66),X68)))) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X70) != s(bool,'const.bool.IN_2'(s(X66,sK7(X68,X67,X66)),s(fun(X66,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X66),X68)))))) )),
  inference(resolution,[],[f95,f142])).
fof(f3506,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.list.EVERY_2'(s(fun(X42,bool),X43),s('type.list.list'(X42),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.bool.IN_2'(s(X42,sK7(X44,X43,X42)),s(fun(X42,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X42),X44))))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f95,f105])).
fof(f3503,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.list.EVERY_2'(s(fun(X24,bool),X25),s('type.list.list'(X24),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK7(X26,X25,X24)),s(fun(X24,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X24),X26)))))) )),
  inference(resolution,[],[f95,f104])).
fof(f3502,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun(X18,bool),X19),s('type.list.list'(X18),X20)))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(X18,sK7(X20,X19,X18)),s(fun(X18,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X18),X20)))))) )),
  inference(resolution,[],[f95,f103])).
fof(f3501,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(X12,bool),X13),s('type.list.list'(X12),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(X12,sK7(X14,X13,X12)),s(fun(X12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X12),X14))))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f95,f101])).
fof(f3499,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) )),
  inference(resolution,[],[f95,f99])).
fof(f3541,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X1,X2,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,X4))) ) | $spl90),
  inference(cnf_transformation,[],[f3541_D])).
fof(f3541_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X1,X2,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,X4))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f3493,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,'const.bool.IN_2'(s(X66,sK7(X67,X68,X66)),s(fun(X66,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X66),X67)))))) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X70) != s(bool,'const.list.EVERY_2'(s(fun(X66,bool),X68),s('type.list.list'(X66),X67)))) )),
  inference(resolution,[],[f95,f142])).
fof(f3486,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(X42,sK7(X43,X44,X42)),s(fun(X42,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X42),X43)))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.list.EVERY_2'(s(fun(X42,bool),X44),s('type.list.list'(X42),X43))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f95,f105])).
fof(f3483,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(X24,sK7(X25,X26,X24)),s(fun(X24,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X24),X25)))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.list.EVERY_2'(s(fun(X24,bool),X26),s('type.list.list'(X24),X25)))) )),
  inference(resolution,[],[f95,f104])).
fof(f3482,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X18,sK7(X19,X20,X18)),s(fun(X18,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X18),X19)))))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(X18,bool),X20),s('type.list.list'(X18),X19)))) )),
  inference(resolution,[],[f95,f103])).
fof(f3481,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(X12,sK7(X13,X14,X12)),s(fun(X12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X12),X13)))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(X12,bool),X14),s('type.list.list'(X12),X13))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f95,f101])).
fof(f3479,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK7(X1,X2,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) )),
  inference(resolution,[],[f95,f99])).
fof(f3315,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,X4))) ) | $spl88),
  inference(cnf_transformation,[],[f3315_D])).
fof(f3315_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,X4))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3288,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,'const.list.EVERY_2'(s(fun(X66,bool),X67),s('type.list.list'(X66),X68)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,happ(s(fun(X66,bool),X67),s(X66,sK7(X68,X67,X66))))) )),
  inference(resolution,[],[f96,f142])).
fof(f3281,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.list.EVERY_2'(s(fun(X42,bool),X43),s('type.list.list'(X42),X44)))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(X42,bool),X43),s(X42,sK7(X44,X43,X42)))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f96,f105])).
fof(f3279,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.list.EVERY_2'(s(fun(X30,bool),X31),s('type.list.list'(X30),X32)))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,happ(s(fun(X30,bool),X31),s(X30,sK7(X32,X31,X30))))) )),
  inference(resolution,[],[f96,f104])).
fof(f3277,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.list.EVERY_2'(s(fun(X18,bool),X19),s('type.list.list'(X18),X20)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(X18,bool),X19),s(X18,sK7(X20,X19,X18))))) )),
  inference(resolution,[],[f96,f103])).
fof(f3276,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.list.EVERY_2'(s(fun(X12,bool),X13),s('type.list.list'(X12),X14)))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK7(X14,X13,X12)))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f96,f101])).
fof(f3274,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) )),
  inference(resolution,[],[f96,f99])).
fof(f3314,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) | s(bool,X4) != s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | p(s(bool,X4))) ) | $spl86),
  inference(cnf_transformation,[],[f3314_D])).
fof(f3314_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) | s(bool,X4) != s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | p(s(bool,X4))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3268,plain,(
  ( ! [X70,X68,X66,X69,X67] : (~p(s(bool,happ(s(fun(X66,bool),X67),s(X66,sK7(X68,X67,X66))))) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X70) != s(bool,'const.list.EVERY_2'(s(fun(X66,bool),X67),s('type.list.list'(X66),X68)))) )),
  inference(resolution,[],[f96,f142])).
fof(f3261,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),X43),s(X42,sK7(X44,X43,X42))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.list.EVERY_2'(s(fun(X42,bool),X43),s('type.list.list'(X42),X44))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f96,f105])).
fof(f3258,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,happ(s(fun(X24,bool),X25),s(X24,sK7(X26,X25,X24))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.list.EVERY_2'(s(fun(X24,bool),X25),s('type.list.list'(X24),X26)))) )),
  inference(resolution,[],[f96,f104])).
fof(f3257,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),X19),s(X18,sK7(X20,X19,X18))))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(X18,bool),X19),s('type.list.list'(X18),X20)))) )),
  inference(resolution,[],[f96,f103])).
fof(f3256,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK7(X14,X13,X12))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.list.EVERY_2'(s(fun(X12,bool),X13),s('type.list.list'(X12),X14))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f96,f101])).
fof(f3254,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(resolution,[],[f96,f99])).
fof(f3187,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f3186,f140])).
fof(f3186,plain,(
  p(s(bool,'F_0')) | $spl54),
  inference(forward_demodulation,[],[f3185,f1751])).
fof(f3185,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) ) | $spl54),
  inference(subsumption_resolution,[],[f2955,f129])).
fof(f129,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7] : ((p(s(bool,X7)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.bvl_to_bviProof.bvl_to_bviProof_3|48|_prove_ltb','thm.bool.AND_CLAUSES|split|3')).
fof(f2955,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) ) | $spl54),
  inference(backward_demodulation,[],[f1751,f95])).
fof(f2941,plain,(
  ( ! [X124] : (s(bool,X124) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X124)) ) | $spl84),
  inference(cnf_transformation,[],[f2941_D])).
fof(f2941_D,plain,(
  ( ! [X124] : (s(bool,X124) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X124)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f2940,plain,(
  ( ! [X123,X124] : (s(bool,'F_0') = s(bool,X123) | s(bool,X124) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f2870,f116])).
fof(f2870,plain,(
  ( ! [X123,X124] : (s(bool,'T_0') != s(bool,X123) | s(bool,'F_0') = s(bool,X123) | s(bool,X124) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl3 | $spl14)),
  inference(superposition,[],[f1032,f2671])).
fof(f2781,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X41),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X41)))) ) | $spl6),
  inference(resolution,[],[f2269,f142])).
fof(f2774,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X31),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X31))) | s(bool,X33) != s(bool,X34)) ) | $spl6),
  inference(resolution,[],[f2269,f105])).
fof(f2771,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X19),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X19)))) ) | $spl6),
  inference(resolution,[],[f2269,f104])).
fof(f2770,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X15),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X15)))) ) | $spl6),
  inference(resolution,[],[f2269,f103])).
fof(f2769,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X11),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X11))) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f2269,f101])).
fof(f2767,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X3),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X3)))) ) | $spl6),
  inference(resolution,[],[f2269,f99])).
fof(f2740,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f2232,f142])).
fof(f2733,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X23) != s(bool,X25)) ) | $spl3),
  inference(resolution,[],[f2232,f105])).
fof(f2731,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f2232,f104])).
fof(f2729,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f2232,f103])).
fof(f2728,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')))) | s(bool,X8) != s(bool,X10)) ) | $spl3),
  inference(resolution,[],[f2232,f101])).
fof(f2726,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v'))))) ) | $spl3),
  inference(resolution,[],[f2232,f99])).
fof(f2687,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl78),
  inference(cnf_transformation,[],[f2687_D])).
fof(f2687_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2675,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl3),
  inference(resolution,[],[f2205,f142])).
fof(f2668,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X23) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f2205,f105])).
fof(f2665,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl3),
  inference(resolution,[],[f2205,f104])).
fof(f2664,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl3),
  inference(resolution,[],[f2205,f103])).
fof(f2663,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5))))) | s(bool,X8) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f2205,f101])).
fof(f2661,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK7(sK5,'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))),'type.bvlSem.v')),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK5)))))) ) | $spl3),
  inference(resolution,[],[f2205,f99])).
fof(f2650,plain,(
  ( ! [X130] : (s(bool,X130) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X130)) ) | $spl76),
  inference(cnf_transformation,[],[f2650_D])).
fof(f2650_D,plain,(
  ( ! [X130] : (s(bool,X130) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X130)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2649,plain,(
  ( ! [X130,X129] : (s(bool,'F_0') = s(bool,X129) | s(bool,X130) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X130)) ) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f2572,f116])).
fof(f2572,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') != s(bool,X129) | s(bool,'F_0') = s(bool,X129) | s(bool,X130) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,'T_0') = s(bool,X130)) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f1032,f2248])).
fof(f2327,plain,(
  $false | ($spl4 | $spl28)),
  inference(subsumption_resolution,[],[f2306,f2275])).
fof(f2275,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | ($spl4 | $spl28)),
  inference(subsumption_resolution,[],[f2274,f115])).
fof(f2274,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3))) ) | ($spl4 | $spl28)),
  inference(forward_demodulation,[],[f331,f2023])).
fof(f2306,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,'const.list.EVERY_2'(s(fun(X84,bool),X85),s('type.list.list'(X84),X86))))) ) | ($spl4 | $spl28)),
  inference(resolution,[],[f2275,f95])).
fof(f2326,plain,(
  $false | ($spl4 | $spl28)),
  inference(subsumption_resolution,[],[f2299,f2275])).
fof(f2299,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,'const.bool.IN_2'(s(X57,sK7(X58,X59,X57)),s(fun(X57,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X57),X58))))))) ) | ($spl4 | $spl28)),
  inference(resolution,[],[f2275,f95])).
fof(f2295,plain,(
  $false | ($spl4 | $spl28)),
  inference(resolution,[],[f2275,f141])).
fof(f2301,plain,(
  $false | ($spl4 | $spl6 | $spl28)),
  inference(resolution,[],[f2275,f153])).
fof(f2270,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl74),
  inference(cnf_transformation,[],[f2270_D])).
fof(f2270_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2252,plain,(
  ( ! [X31,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f153,f142])).
fof(f2245,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X29) != s(bool,X30)) ) | $spl6),
  inference(resolution,[],[f153,f105])).
fof(f2242,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f153,f104])).
fof(f2241,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f153,f103])).
fof(f2240,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X14) != s(bool,X15)) ) | $spl6),
  inference(resolution,[],[f153,f101])).
fof(f2238,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f153,f99])).
fof(f2233,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl72),
  inference(cnf_transformation,[],[f2233_D])).
fof(f2233_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2221,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(resolution,[],[f148,f142])).
fof(f2214,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f148,f105])).
fof(f2212,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(resolution,[],[f148,f104])).
fof(f2210,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(resolution,[],[f148,f103])).
fof(f2209,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f148,f101])).
fof(f2207,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) ) | $spl3),
  inference(resolution,[],[f148,f99])).
fof(f2195,plain,(
  ( ! [X127] : (s(bool,X127) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X127)) ) | $spl70),
  inference(cnf_transformation,[],[f2195_D])).
fof(f2195_D,plain,(
  ( ! [X127] : (s(bool,X127) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X127)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2194,plain,(
  ( ! [X127,X126] : (s(bool,'F_0') = s(bool,X126) | s(bool,X127) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X127)) ) | ($spl4 | $spl14)),
  inference(subsumption_resolution,[],[f2125,f116])).
fof(f2125,plain,(
  ( ! [X127,X126] : (s(bool,'T_0') != s(bool,X126) | s(bool,'F_0') = s(bool,X126) | s(bool,X127) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,'T_0') = s(bool,X127)) ) | ($spl4 | $spl14)),
  inference(superposition,[],[f1032,f2023])).
fof(f2039,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl68),
  inference(cnf_transformation,[],[f2039_D])).
fof(f2039_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2027,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl4),
  inference(resolution,[],[f149,f142])).
fof(f2020,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,X24) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f149,f105])).
fof(f2017,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl4),
  inference(resolution,[],[f149,f104])).
fof(f2016,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl4),
  inference(resolution,[],[f149,f103])).
fof(f2015,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,X9) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f149,f101])).
fof(f2013,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl4),
  inference(resolution,[],[f149,f99])).
fof(f149,plain,(
  p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) | $spl4),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f636,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X6),s('type.list.list'('type.bvlSem.v'),X7),s('type.bvlSem.state'(sK0),X8))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X8))))),s('type.list.list'('type.bvlSem.v'),X7)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X8))))) )),
  inference(forward_demodulation,[],[f633,f97])).
fof(f633,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X6),s('type.list.list'('type.bvlSem.v'),X7),s('type.bvlSem.state'(sK0),X8))) | p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X8))))),s('type.list.list'('type.bvlSem.v'),X7)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X8))))) )),
  inference(superposition,[],[f93,f89])).
fof(f93,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X1))))),s('type.bvlSem.v',X3)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X4),s('type.list.list'('type.bvlSem.v'),X5),s('type.bvlSem.state'(X0),X2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(X0)),'const.pair.,_2'(s('type.bvlSem.v',X3),s('type.bvlSem.state'(X0),X1))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(X0),X2))))),s('type.list.list'('type.bvlSem.v'),X5)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(X0),X2))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f1940,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X0),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))))) )),
  inference(superposition,[],[f538,f89])).
fof(f538,plain,(
  ( ! [X4,X5,X3] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X5),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',X4),s('type.bvlSem.state'(sK0),X3))))) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X3))),s('type.bvlSem.v',X4))))) )),
  inference(subsumption_resolution,[],[f537,f141])).
fof(f537,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X3))),s('type.bvlSem.v',X4)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X5),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',X4),s('type.bvlSem.state'(sK0),X3)))))) )),
  inference(forward_demodulation,[],[f536,f166])).
fof(f536,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X3))),s('type.bvlSem.v',X4)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X5),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',X4),s('type.bvlSem.state'(sK0),X3))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2))))) )),
  inference(forward_demodulation,[],[f511,f97])).
fof(f511,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X3))))),s('type.bvlSem.v',X4)))) | s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X5),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.semanticPrimitives.Rval_1'(s('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'const.pair.,_2'(s('type.bvlSem.v',X4),s('type.bvlSem.state'(sK0),X3))))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2))))) )),
  inference(resolution,[],[f88,f93])).
fof(f2009,plain,(
  $false | $spl5),
  inference(subsumption_resolution,[],[f2008,f141])).
fof(f2008,plain,(
  ~p(s(bool,'T_0')) | $spl5),
  inference(forward_demodulation,[],[f2007,f166])).
fof(f2007,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) | $spl5),
  inference(subsumption_resolution,[],[f2006,f138])).
fof(f2006,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) | $spl5),
  inference(forward_demodulation,[],[f2005,f523])).
fof(f2005,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) | $spl5),
  inference(equality_resolution,[],[f637])).
fof(f637,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',sK4),s('type.list.list'('type.bvlSem.v'),sK6),s('type.bvlSem.state'(sK0),sK2))) != s('type.semanticPrimitives.result'('type.pair.prod'('type.bvlSem.v','type.bvlSem.state'(sK0)),'type.bvlSem.v'),'const.bvlSem.do_app_3'(s('type.closLang.op',X6),s('type.list.list'('type.bvlSem.v'),X7),s('type.bvlSem.state'(sK0),X8))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),X8))))),s('type.list.list'('type.bvlSem.v'),X7)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),X8))))) ) | $spl5),
  inference(subsumption_resolution,[],[f636,f150])).
fof(f1996,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl62),
  inference(cnf_transformation,[],[f1996_D])).
fof(f1996_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1957,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X80) != s(bool,X82) | s(bool,X82) != s(bool,X83) | s(bool,'F_0') = s(bool,X83) | p(s(bool,X84)) | p(s(bool,X85)) | s(bool,X81) != s(bool,X85)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1276,f142])).
fof(f1950,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1276,f105])).
fof(f1947,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1276,f104])).
fof(f1946,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1276,f103])).
fof(f1945,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1276,f101])).
fof(f1943,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | ($spl14 | $spl40)),
  inference(resolution,[],[f1276,f99])).
fof(f1924,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X0) != s(bool,X2)) ) | $spl60),
  inference(cnf_transformation,[],[f1924_D])).
fof(f1924_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1890,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X80) != s(bool,X82) | p(s(bool,X83)) | s(bool,X82) != s(bool,X83) | p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X81) != s(bool,X85)) ) | $spl46),
  inference(resolution,[],[f1307,f142])).
fof(f1883,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | p(s(bool,X52)) | s(bool,X51) != s(bool,X52) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl46),
  inference(resolution,[],[f1307,f105])).
fof(f1881,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,X35) != s(bool,X36) | s(bool,X35) != s(bool,X37) | p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl46),
  inference(resolution,[],[f1307,f104])).
fof(f1879,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | p(s(bool,X24)) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl46),
  inference(resolution,[],[f1307,f103])).
fof(f1878,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | p(s(bool,X17)) | s(bool,X16) != s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl46),
  inference(resolution,[],[f1307,f101])).
fof(f1876,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl46),
  inference(resolution,[],[f1307,f99])).
fof(f1837,plain,(
  ( ! [X4,X7,X5] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X4),X5),s('type.bvlSem.v',sK3)))) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X4),X5))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))))) | ~p(s(bool,X7))) ) | $spl58),
  inference(cnf_transformation,[],[f1837_D])).
fof(f1837_D,plain,(
  ( ! [X4,X7,X5] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X4),X5),s('type.bvlSem.v',sK3)))) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X4),X5))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))))) | ~p(s(bool,X7))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1827,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X56),X57),s('type.bvlSem.v',sK3)))) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X59) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X56),X57))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) ) | $spl5),
  inference(resolution,[],[f308,f142])).
fof(f308,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X0),X1))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) | ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X0),X1),s('type.bvlSem.v',sK3))))) ) | $spl5),
  inference(resolution,[],[f150,f118])).
fof(f1820,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X39),X40),s('type.bvlSem.v',sK3)))) | ~p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X39),X40))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))))) | s(bool,X41) != s(bool,X43)) ) | $spl5),
  inference(resolution,[],[f308,f105])).
fof(f1818,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X29),X30),s('type.bvlSem.v',sK3)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X29),X30))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) ) | $spl5),
  inference(resolution,[],[f308,f104])).
fof(f1816,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X19),X20),s('type.bvlSem.v',sK3)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X19),X20))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) ) | $spl5),
  inference(resolution,[],[f308,f103])).
fof(f1815,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X14),X15),s('type.bvlSem.v',sK3)))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X14),X15))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))))) | s(bool,X16) != s(bool,X18)) ) | $spl5),
  inference(resolution,[],[f308,f101])).
fof(f1813,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num',X4),X5),s('type.bvlSem.v',sK3)))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num',X4),X5))),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1)))))))) ) | $spl5),
  inference(resolution,[],[f308,f99])).
fof(f1800,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X39)))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X39),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f142])).
fof(f1793,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X29)))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X29),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f539,f105])).
fof(f1791,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X21)))) | p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X21),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f104])).
fof(f1789,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X13)))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X13),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f103])).
fof(f1788,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X9)))) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X9),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6))))) | s(bool,X10) != s(bool,X12)) )),
  inference(resolution,[],[f539,f101])).
fof(f1786,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))),s('type.bvlSem.v',X1)))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK6)))))) )),
  inference(resolution,[],[f539,f99])).
fof(f1750,plain,(
  ( ! [X41] : (s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X41)) ) | $spl52),
  inference(cnf_transformation,[],[f1750_D])).
fof(f1750_D,plain,(
  ( ! [X41] : (s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X41)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1749,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') = s(bool,X40) | s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X41)) ) | $spl14),
  inference(subsumption_resolution,[],[f1719,f116])).
fof(f1719,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X40) | s(bool,'F_0') = s(bool,X40) | s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,'T_0') = s(bool,X41)) ) | $spl14),
  inference(superposition,[],[f1032,f523])).
fof(f1606,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl50),
  inference(cnf_transformation,[],[f1606_D])).
fof(f1606_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1572,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X66) != s(bool,X67) | s(bool,'T_0') = s(bool,X68) | s(bool,X66) != s(bool,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X67) != s(bool,X70)) ) | $spl44),
  inference(resolution,[],[f1151,f142])).
fof(f1565,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X44) | s(bool,X42) != s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl44),
  inference(resolution,[],[f1151,f105])).
fof(f1563,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) != s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl44),
  inference(resolution,[],[f1151,f104])).
fof(f1561,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,X20) | s(bool,X18) != s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl44),
  inference(resolution,[],[f1151,f103])).
fof(f1560,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X14) | s(bool,X12) != s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl44),
  inference(resolution,[],[f1151,f101])).
fof(f1558,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl44),
  inference(resolution,[],[f1151,f99])).
fof(f1368,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f1368_D])).
fof(f1368_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1334,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X66) != s(bool,X67) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X68) != s(bool,X70)) ) | $spl14),
  inference(resolution,[],[f836,f142])).
fof(f1327,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,X42) | s(bool,X43) != s(bool,X44) | s(bool,X42) != s(bool,X43) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl14),
  inference(resolution,[],[f836,f105])).
fof(f1324,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,X24) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) ) | $spl14),
  inference(resolution,[],[f836,f104])).
fof(f1323,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,X18) | s(bool,X19) != s(bool,X20) | s(bool,X18) != s(bool,X19) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) ) | $spl14),
  inference(resolution,[],[f836,f103])).
fof(f1322,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'F_0') = s(bool,X12) | s(bool,X13) != s(bool,X14) | s(bool,X12) != s(bool,X13) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl14),
  inference(resolution,[],[f836,f101])).
fof(f1320,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) ) | $spl14),
  inference(resolution,[],[f836,f99])).
fof(f1307,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | s(bool,X0) != s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f1307_D])).
fof(f1307_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1274,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X66) != s(bool,X67) | s(bool,X66) != s(bool,X68) | p(s(bool,X68)) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X67) != s(bool,X70)) ) | $spl40),
  inference(resolution,[],[f807,f142])).
fof(f1267,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl40),
  inference(resolution,[],[f807,f105])).
fof(f1265,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl40),
  inference(resolution,[],[f807,f104])).
fof(f1263,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | p(s(bool,X20)) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl40),
  inference(resolution,[],[f807,f103])).
fof(f1262,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl40),
  inference(resolution,[],[f807,f101])).
fof(f1260,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | p(s(bool,X2)) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl40),
  inference(resolution,[],[f807,f99])).
fof(f1151,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1)) ) | $spl44),
  inference(cnf_transformation,[],[f1151_D])).
fof(f1151_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1118,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') = s(bool,X52) | s(bool,X52) != s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X53) != s(bool,X55)) ) | $spl14),
  inference(resolution,[],[f1102,f142])).
fof(f1111,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,X35) | s(bool,X35) != s(bool,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl14),
  inference(resolution,[],[f1102,f105])).
fof(f1109,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X25) != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | $spl14),
  inference(resolution,[],[f1102,f104])).
fof(f1107,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl14),
  inference(resolution,[],[f1102,f103])).
fof(f1106,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,X10) | s(bool,X10) != s(bool,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl14),
  inference(resolution,[],[f1102,f101])).
fof(f1104,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl14),
  inference(resolution,[],[f1102,f99])).
fof(f867,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) | $spl42),
  inference(cnf_transformation,[],[f867_D])).
fof(f867_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f835,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53) | p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X52) != s(bool,X55)) ) | $spl14),
  inference(resolution,[],[f770,f142])).
fof(f828,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl14),
  inference(resolution,[],[f770,f105])).
fof(f825,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X20) != s(bool,X24)) ) | $spl14),
  inference(resolution,[],[f770,f104])).
fof(f824,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl14),
  inference(resolution,[],[f770,f103])).
fof(f823,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl14),
  inference(resolution,[],[f770,f101])).
fof(f821,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl14),
  inference(resolution,[],[f770,f99])).
fof(f807,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | p(s(bool,X0))) ) | $spl40),
  inference(cnf_transformation,[],[f807_D])).
fof(f807_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | p(s(bool,X0))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f776,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) != s(bool,X53) | p(s(bool,X52)) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X53) != s(bool,X55)) ) | $spl14),
  inference(resolution,[],[f181,f142])).
fof(f769,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | p(s(bool,X35)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl14),
  inference(resolution,[],[f181,f105])).
fof(f767,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X25)) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | $spl14),
  inference(resolution,[],[f181,f104])).
fof(f765,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | p(s(bool,X15)) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl14),
  inference(resolution,[],[f181,f103])).
fof(f764,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | p(s(bool,X10)) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl14),
  inference(resolution,[],[f181,f101])).
fof(f762,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X0)) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl14),
  inference(resolution,[],[f181,f99])).
fof(f181,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | p(s(bool,X11))) ) | $spl14),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | p(s(bool,X11))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f742,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,X110) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109)) ) | $spl0),
  inference(forward_demodulation,[],[f741,f227])).
fof(f741,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109) | s(bool,X110) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(subsumption_resolution,[],[f680,f142])).
fof(f680,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,'T_0')) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109) | s(bool,X110) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f105,f227])).
fof(f738,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,X93) | p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,X95)) ) | $spl0),
  inference(subsumption_resolution,[],[f675,f142])).
fof(f675,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,X93) | p(s(bool,X94)) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,X95)) ) | $spl0),
  inference(superposition,[],[f105,f227])).
fof(f731,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X63) | ~p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X64) != s(bool,X65)) ) | $spl0),
  inference(subsumption_resolution,[],[f665,f105])).
fof(f665,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X63) | p(s(bool,X63)) | ~p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X64) != s(bool,X65)) ) | $spl0),
  inference(superposition,[],[f103,f227])).
fof(f726,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,X50) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49)) ) | $spl0),
  inference(forward_demodulation,[],[f725,f227])).
fof(f725,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(subsumption_resolution,[],[f660,f142])).
fof(f660,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f101,f227])).
fof(f722,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38)) ) | $spl0),
  inference(subsumption_resolution,[],[f656,f142])).
fof(f656,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38)) ) | $spl0),
  inference(superposition,[],[f101,f227])).
fof(f719,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X18) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X20)) ) | $spl0),
  inference(subsumption_resolution,[],[f650,f105])).
fof(f650,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X20)) ) | $spl0),
  inference(superposition,[],[f99,f227])).
fof(f619,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19)) ) | $spl16),
  inference(subsumption_resolution,[],[f607,f182])).
fof(f607,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl16),
  inference(superposition,[],[f182,f116])).
fof(f612,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,'T_0') != s(bool,X17)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f606])).
fof(f606,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X16))) ) | $spl16),
  inference(superposition,[],[f182,f113])).
fof(f613,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f605])).
fof(f605,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl16),
  inference(superposition,[],[f182,f115])).
fof(f618,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13)) ) | $spl16),
  inference(subsumption_resolution,[],[f604,f182])).
fof(f604,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | $spl16),
  inference(superposition,[],[f182,f116])).
fof(f587,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') != s(bool,X63) | ~p(s(bool,X64))) ) | $spl16),
  inference(resolution,[],[f182,f136])).
fof(f586,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X62))) ) | $spl16),
  inference(resolution,[],[f182,f138])).
fof(f561,plain,(
  $false | $spl12),
  inference(trivial_inequality_removal,[],[f551])).
fof(f551,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl12),
  inference(superposition,[],[f180,f166])).
fof(f560,plain,(
  $false | $spl12),
  inference(equality_resolution,[],[f180])).
fof(f540,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | $spl36),
  inference(cnf_transformation,[],[f540_D])).
fof(f540_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f527,plain,(
  ( ! [X31,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f88,f142])).
fof(f520,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,X29) != s(bool,X30)) )),
  inference(resolution,[],[f88,f105])).
fof(f541,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) ) | $spl38),
  inference(cnf_transformation,[],[f541_D])).
fof(f541_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f519,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f88,f105])).
fof(f517,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f88,f104])).
fof(f516,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f88,f103])).
fof(f515,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6))) | s(bool,X14) != s(bool,X15)) )),
  inference(resolution,[],[f88,f101])).
fof(f514,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X11) | s(bool,X12) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f88,f101])).
fof(f513,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK6)))) )),
  inference(resolution,[],[f88,f99])).
fof(f499,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f496])).
fof(f496,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23) | p(s(bool,X22))) ) | $spl10),
  inference(superposition,[],[f179,f109])).
fof(f500,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,X21)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f495])).
fof(f495,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,X21) | p(s(bool,X20))) ) | $spl10),
  inference(superposition,[],[f179,f111])).
fof(f508,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19)) ) | $spl10),
  inference(subsumption_resolution,[],[f494,f179])).
fof(f494,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl10),
  inference(superposition,[],[f179,f116])).
fof(f507,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13)) ) | $spl10),
  inference(subsumption_resolution,[],[f491,f179])).
fof(f491,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | $spl10),
  inference(superposition,[],[f179,f116])).
fof(f474,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | p(s(bool,X60))) ) | $spl10),
  inference(resolution,[],[f179,f127])).
fof(f473,plain,(
  ( ! [X57,X58] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X58))) ) | $spl10),
  inference(resolution,[],[f179,f129])).
fof(f438,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,X110) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109)) )),
  inference(forward_demodulation,[],[f437,f166])).
fof(f437,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109) | s(bool,X110) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(subsumption_resolution,[],[f375,f142])).
fof(f375,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,'T_0')) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,X109) | s(bool,X110) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(superposition,[],[f105,f166])).
fof(f434,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,X93) | p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,X95)) )),
  inference(subsumption_resolution,[],[f370,f142])).
fof(f370,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,X93) | p(s(bool,X94)) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,X95)) )),
  inference(superposition,[],[f105,f166])).
fof(f427,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X63) | ~p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X64) != s(bool,X65)) )),
  inference(subsumption_resolution,[],[f360,f105])).
fof(f360,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X63) | p(s(bool,X63)) | ~p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X64) != s(bool,X65)) )),
  inference(superposition,[],[f103,f166])).
fof(f422,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,X50) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49)) )),
  inference(forward_demodulation,[],[f421,f166])).
fof(f421,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(subsumption_resolution,[],[f355,f142])).
fof(f355,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(superposition,[],[f101,f166])).
fof(f418,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38)) )),
  inference(subsumption_resolution,[],[f351,f142])).
fof(f351,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38)) )),
  inference(superposition,[],[f101,f166])).
fof(f414,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X18) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X20)) )),
  inference(subsumption_resolution,[],[f345,f105])).
fof(f345,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X20)) )),
  inference(superposition,[],[f99,f166])).
fof(f331,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl28),
  inference(cnf_transformation,[],[f331_D])).
fof(f331_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f323,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl5),
  inference(resolution,[],[f150,f142])).
fof(f316,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,X23) != s(bool,X25)) ) | $spl5),
  inference(resolution,[],[f150,f105])).
fof(f314,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl5),
  inference(resolution,[],[f150,f104])).
fof(f332,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl30),
  inference(cnf_transformation,[],[f332_D])).
fof(f332_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f313,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X16) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,X14) != s(bool,X15)) ) | $spl5),
  inference(resolution,[],[f150,f103])).
fof(f312,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl5),
  inference(resolution,[],[f150,f103])).
fof(f311,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,X8) != s(bool,X10)) ) | $spl5),
  inference(resolution,[],[f150,f101])).
fof(f310,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X7) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3))) | s(bool,X5) != s(bool,X6)) ) | $spl5),
  inference(resolution,[],[f150,f99])).
fof(f309,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) ) | $spl5),
  inference(resolution,[],[f150,f99])).
fof(f258,plain,(
  $false | $spl20),
  inference(resolution,[],[f184,f140])).
fof(f239,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl26),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f231,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f145,f142])).
fof(f224,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | s(bool,X30) != s(bool,X31)) ) | $spl0),
  inference(resolution,[],[f145,f105])).
fof(f212,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl24),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f223,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f145,f105])).
fof(f221,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f145,f104])).
fof(f220,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f145,f103])).
fof(f219,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | s(bool,X15) != s(bool,X16)) ) | $spl0),
  inference(resolution,[],[f145,f101])).
fof(f218,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f145,f101])).
fof(f217,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f145,f99])).
fof(f203,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl1),
  inference(resolution,[],[f146,f142])).
fof(f196,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | s(bool,X25) != s(bool,X27)) ) | $spl1),
  inference(resolution,[],[f146,f105])).
fof(f194,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl1),
  inference(resolution,[],[f146,f104])).
fof(f193,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X18) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | s(bool,X16) != s(bool,X17)) ) | $spl1),
  inference(resolution,[],[f146,f103])).
fof(f192,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl1),
  inference(resolution,[],[f146,f103])).
fof(f191,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | s(bool,X10) != s(bool,X12)) ) | $spl1),
  inference(resolution,[],[f146,f101])).
fof(f190,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))) | s(bool,X7) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f146,f99])).
fof(f189,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) ) | $spl1),
  inference(resolution,[],[f146,f99])).
fof(f178,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) ) | $spl8),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f170,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(resolution,[],[f87,f142])).
fof(f163,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f87,f105])).
fof(f180,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) ) | $spl12),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f162,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(resolution,[],[f87,f105])).
fof(f160,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(resolution,[],[f87,f104])).
fof(f159,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(resolution,[],[f87,f103])).
fof(f158,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f87,f101])).
fof(f157,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(resolution,[],[f87,f101])).
fof(f156,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK2)))) )),
  inference(resolution,[],[f87,f99])).
fof(f146,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) | $spl1),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f150,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) | $spl5),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f151,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))))),
  inference(forward_demodulation,[],[f90,f97])).
fof(f90,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK2))))),s('type.list.list'('type.bvlSem.v'),sK5)))) | ~p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))))),
  inference(cnf_transformation,[],[f64])).
fof(f148,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) | $spl3),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f144,plain,(
  ~p(s(bool,'const.bvl_to_bviProof.bv_ok_2'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))),s('type.bvlSem.v',sK3)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))))),
  inference(forward_demodulation,[],[f91,f97])).
fof(f91,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.list.list'('type.bvlSem.v'),sK5)))) | ~p(s(bool,happ(s(fun('type.bvlSem.v',bool),'const.bvl_to_bviProof.bv_ok_1'(s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvlSem.state_refs_1'(s('type.bvlSem.state'(sK0),sK1))))),s('type.bvlSem.v',sK3)))) | ~p(s(bool,'const.bvl_to_bviProof.state_ok_1'(s('type.bvlSem.state'(sK0),sK1))))),
  inference(cnf_transformation,[],[f64])).
% SZS output end Proof for thm.bvl_to_bviProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 320
% Passive clauses: 3149
% Generated clauses: 14411
% Final active clauses: 243
% Final passive clauses: 2052
% Input formulas: 22
% Initial clauses: 55
% 
% Selected by SInE selection: 21
% SInE iterations: 1
% 
% Duplicate literals: 125
% Trivial inequalities: 571
% Fw subsumption resolutions: 1840
% Fw demodulations: 1209
% Bw demodulations: 30
% 
% Simple tautologies: 298
% Forward subsumptions: 6598
% 
% Binary resolution: 3633
% Factoring: 3
% Forward superposition: 3815
% Backward superposition: 2879
% Self superposition: 2
% Equality resolution: 56
% 
% Split clauses: 998
% Split components: 2147
% SAT solver clauses: 1034
% SAT solver unit clauses: 11
% SAT solver binary clauses: 510
% 
% Sat splits: 998
% Sat splitting refutations: 35
% 
% Memory used [KB]: 9083
% Time elapsed: 0.704 s
% ------------------------------
% ------------------------------
% Success in time 12.435 s
