Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_numberProof
% (90250)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 619
% Passive clauses: 8956
% Generated clauses: 16444
% Final active clauses: 556
% Final passive clauses: 7820
% Input formulas: 43
% Initial clauses: 179
% 
% Duplicate literals: 282
% Trivial inequalities: 77
% Fw subsumption resolutions: 1250
% Fw demodulations: 79
% Bw demodulations: 62
% 
% Simple tautologies: 213
% Equational tautologies: 67
% Forward subsumptions: 5357
% 
% Binary resolution: 10280
% Factoring: 1075
% Forward superposition: 1354
% Backward superposition: 1550
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 33
% 
% Split clauses: 286
% Split components: 573
% SAT solver clauses: 612
% SAT solver unit clauses: 18
% SAT solver binary clauses: 146
% 
% Sat splits: 286
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 610
% TWLsolver calls for satisfiability: 355
% 
% Memory used [KB]: 6908
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_numberProof
% (90275)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 467
% Passive clauses: 2671
% Generated clauses: 14913
% Final active clauses: 369
% Final passive clauses: 1371
% Input formulas: 43
% Initial clauses: 179
% 
% Duplicate literals: 437
% Trivial inequalities: 86
% Fw subsumption resolutions: 1128
% Fw demodulations: 576
% Bw demodulations: 725
% Condensations: 521
% Global subsumptions: 13
% 
% Simple tautologies: 307
% Equational tautologies: 23
% Forward subsumptions: 5897
% Bw demodulations to eq. taut.: 181
% 
% Binary resolution: 3368
% Factoring: 94
% Forward superposition: 1986
% Backward superposition: 2393
% Self superposition: 35
% Equality factoring: 36
% Equality resolution: 12
% 
% Split clauses: 157
% Split components: 327
% SAT solver clauses: 3243
% SAT solver unit clauses: 185
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 157
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 3191
% TWLsolver calls for satisfiability: 1920
% 
% Memory used [KB]: 12409
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_numberProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_numberProof
% SZS output start Proof for thm.clos_numberProof
fof(f2248,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f247,f436,f434,f246,f435,f433,f1755,f2247])).
fof(f2247,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2246,f248])).
fof(f248,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_numberProof.lookup_vars_NONE_related_env',aHL_TRUTH)).
fof(f2246,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2007,f1769])).
fof(f1769,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK14))),s('type.num.num',sK16(sK13,sK15,'type.closSem.v')))) | $spl0),
  inference(unit_resulting_resolution,[],[f248,f1763,f308])).
fof(f308,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.clos_numberProof.lookup_vars_NONE_related_env','thm.bool.IMP_ANTISYM_AX')).
fof(f1763,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK14))),s('type.num.num',sK16(sK13,sK15,'type.closSem.v'))))) | $spl0),
  inference(forward_demodulation,[],[f1757,f1249])).
fof(f1249,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK14))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK15)))),
  inference(subsumption_resolution,[],[f1247,f248])).
fof(f1247,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK14))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK15)))),
  inference(superposition,[],[f339,f476])).
fof(f476,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))),
  inference(unit_resulting_resolution,[],[f245,f263])).
fof(f263,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.clos_numberProof.lookup_vars_NONE_related_env','thm.bool.EQ_CLAUSES|split|1')).
fof(f245,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))))),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) & (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK15))) | s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK14)))) & (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK15))) | s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15])],[f162])).
fof(f162,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X1)))) & (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X2))) | s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1)))) & (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X2))) | s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1)))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X1)))) & ((s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X2))) | s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1)))) & (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X2))) | s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1))))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X1)))) & (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X2))) <~> s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1)))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X1)))) => (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X2))) <=> s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X0),s('type.list.list'('type.closSem.v'),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X8),s('type.list.list'('type.closSem.v'),X7)))) => (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X6),s('type.list.list'('type.closSem.v'),X8))) <=> s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X6),s('type.list.list'('type.closSem.v'),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X8),s('type.list.list'('type.closSem.v'),X7)))) => (s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X6),s('type.list.list'('type.closSem.v'),X8))) <=> s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X6),s('type.list.list'('type.closSem.v'),X7)))))),
  file('thm.clos_numberProof.lookup_vars_NONE_related_env','thm.clos_numberProof.lookup_vars_NONE_related_env')).
fof(f339,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9,X17] : ! [X21,X11,X15] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X9,fun(X17,bool)),X21),s('type.list.list'(X9),X11),s('type.list.list'(X17),X15)))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X9),X11))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X17),X15))))),
  file('thm.clos_numberProof.lookup_vars_NONE_related_env','thm.list.LIST_REL_LENGTH')).
fof(f1757,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK15))),s('type.num.num',sK16(sK13,sK15,'type.closSem.v'))))) | $spl0),
  inference(unit_resulting_resolution,[],[f433,f295])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') != s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',sK16(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') != s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1))) | (p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16(X2,X1,X0)),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',sK16(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',X3))))) | s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') = s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') != s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',X3))))) | s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') = s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1)))))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') != s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',X3))))) | s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') = s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1)))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : (s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') = s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',X3))))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') = s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',X3))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9] : ! [X19,X6] : (s('type.option.option'('type.list.list'(X9)),'const.option.NONE_0') = s('type.option.option'('type.list.list'(X9)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X6),s('type.list.list'(X9),X19))) <=> ? [X20] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X20),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X6)))))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X9),X19))),s('type.num.num',X20))))))),
  file('thm.clos_numberProof.lookup_vars_NONE_related_env','thm.closProps.lookup_vars_NONE')).
fof(f2007,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK14))),s('type.num.num',sK16(sK13,sK15,'type.closSem.v'))))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f436,f1756,f296])).
fof(f296,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num',X3)))) | s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') = s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f185])).
fof(f1756,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16(sK13,sK15,'type.closSem.v')),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),sK13)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f433,f294])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.list.list'(X0)),'const.option.NONE_0') != s('type.option.option'('type.list.list'(X0)),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),X2),s('type.list.list'(X0),X1))) | p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16(X2,X1,X0)),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2))))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f1755,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1754,f248])).
fof(f1754,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1753,f1264])).
fof(f1264,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK14))),s('type.num.num',sK16(sK13,sK14,'type.closSem.v')))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f1257,f308])).
fof(f1257,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK14))),s('type.num.num',sK16(sK13,sK14,'type.closSem.v'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f435,f295])).
fof(f1753,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK14))),s('type.num.num',sK16(sK13,sK14,'type.closSem.v'))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1514,f1249])).
fof(f1514,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.closSem.v'),sK15))),s('type.num.num',sK16(sK13,sK14,'type.closSem.v'))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f434,f1500,f296])).
fof(f1500,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16(sK13,sK14,'type.closSem.v')),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),sK13)))))) | $spl2),
  inference(unit_resulting_resolution,[],[f435,f294])).
fof(f433,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK15))) | $spl0),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK15))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f435,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK14))) | $spl2),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK14))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f246,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK14))) | s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') = s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK15)))),
  inference(cnf_transformation,[],[f163])).
fof(f434,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK15))) | $spl1),
  inference(cnf_transformation,[],[f434_D])).
fof(f434_D,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK15))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f436,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK14))) | $spl3),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK14))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f247,plain,(
  s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK14))) | s('type.option.option'('type.list.list'('type.closSem.v')),'const.option.NONE_0') != s('type.option.option'('type.list.list'('type.closSem.v')),'const.closSem.lookup_vars_2'(s('type.list.list'('type.num.num'),sK13),s('type.list.list'('type.closSem.v'),sK15)))),
  inference(cnf_transformation,[],[f163])).
% SZS output end Proof for thm.clos_numberProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 248
% Passive clauses: 649
% Generated clauses: 2002
% Final active clauses: 236
% Final passive clauses: 222
% Input formulas: 43
% Initial clauses: 179
% 
% Duplicate literals: 6
% Trivial inequalities: 6
% Fw subsumption resolutions: 34
% Fw demodulations: 176
% Bw demodulations: 5
% Condensations: 5
% 
% Simple tautologies: 42
% Equational tautologies: 23
% Forward subsumptions: 1057
% 
% Unit resulting resolution: 1402
% Factoring: 1
% Forward superposition: 84
% Backward superposition: 95
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 657
% SAT solver unit clauses: 122
% SAT solver binary clauses: 437
% 
% Sat splits: 2
% Sat splitting refutations: 2
% 
% Memory used [KB]: 2046
% Time elapsed: 0.090 s
% ------------------------------
% ------------------------------
% Success in time 1.422 s
