Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closProps
% (88783)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 142
% Passive clauses: 3756
% Generated clauses: 17692
% Final active clauses: 131
% Final passive clauses: 3490
% Input formulas: 105
% Initial clauses: 117
% 
% Duplicate literals: 657
% Trivial inequalities: 177
% Fw subsumption resolutions: 1343
% Fw demodulations: 1977
% Bw demodulations: 2
% 
% Simple tautologies: 267
% Equational tautologies: 65
% Forward subsumptions: 9057
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 373
% Factoring: 4
% Forward superposition: 7029
% Backward superposition: 5550
% Self superposition: 230
% Equality factoring: 126
% Equality resolution: 16
% 
% Split clauses: 218
% Split components: 445
% SAT solver clauses: 465
% SAT solver unit clauses: 10
% SAT solver binary clauses: 95
% 
% Sat splits: 218
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 465
% TWLsolver calls for satisfiability: 102
% 
% Memory used [KB]: 13944
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% (88805)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 121
% Passive clauses: 1131
% Generated clauses: 13266
% Final active clauses: 111
% Final passive clauses: 996
% Input formulas: 105
% Initial clauses: 117
% 
% Duplicate literals: 450
% Trivial inequalities: 89
% Fw subsumption resolutions: 566
% Fw demodulations: 1201
% Bw demodulations: 14
% Condensations: 680
% 
% Simple tautologies: 142
% Equational tautologies: 25
% Forward subsumptions: 7217
% Fw demodulations to eq. taut.: 15
% 
% Binary resolution: 91
% Factoring: 1
% Forward superposition: 4619
% Backward superposition: 4149
% Self superposition: 172
% Equality factoring: 47
% Equality resolution: 9
% 
% Split clauses: 160
% Split components: 327
% SAT solver clauses: 1860
% SAT solver unit clauses: 196
% SAT solver binary clauses: 40
% 
% Sat splits: 160
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1860
% TWLsolver calls for satisfiability: 3553
% 
% Memory used [KB]: 18933
% Time elapsed: 0.711 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.closProps
% (88852)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 289
% Passive clauses: 4019
% Generated clauses: 94398
% Final active clauses: 87
% Final passive clauses: 307
% Input formulas: 105
% Initial clauses: 117
% Discarded non-redundant clauses: 35328
% 
% Duplicate literals: 1369
% Trivial inequalities: 262
% Fw subsumption resolutions: 1816
% Fw demodulations: 5407
% Bw demodulations: 1169
% Condensations: 4336
% Global subsumptions: 101
% 
% Simple tautologies: 338
% Equational tautologies: 189
% Forward subsumptions: 70842
% Fw demodulations to eq. taut.: 84
% Bw demodulations to eq. taut.: 156
% 
% Unit resulting resolution: 967
% Factoring: 4
% Forward superposition: 47909
% Backward superposition: 27885
% Self superposition: 822
% Equality factoring: 130
% Equality resolution: 26
% 
% Split clauses: 97
% Split components: 201
% SAT solver clauses: 4472
% SAT solver unit clauses: 260
% SAT solver binary clauses: 540
% 
% Sat splits: 97
% Sat splitting refutations: 156
% 
% Memory used [KB]: 166308
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f25826,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f130,f144,f142,f129,f145,f140,f128,f127,f143,f126,f147,f125,f141,f139,f159,f162,f163,f164,f165,f551,f558,f1206,f1284,f1285,f553,f1613,f1610,f1608,f1600,f1636,f1619,f1587,f1654,f1685,f1688,f1885,f1888,f1889,f1892,f1897,f1900,f1901,f1904,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1914,f1915,f1916,f1918,f1919,f1921,f1922,f1923,f1924,f1925,f1926,f1927,f1928,f1929,f1930,f1931,f1932,f1933,f2005,f2007,f2205,f2207,f2208,f2206,f2209,f2210,f2211,f2213,f2222,f2272,f2279,f2544,f2545,f2556,f2560,f2561,f2656,f2657,f2840,f2853,f2865,f3124,f3125,f3340,f3341,f3342,f3343,f3344,f3345,f3346,f3347,f3348,f3349,f3350,f3351,f3352,f3353,f3354,f3355,f3356,f3357,f3358,f1902,f3359,f3360,f1898,f3361,f3362,f1891,f3363,f3364,f1887,f1886,f3379,f3380,f1890,f3382,f3383,f3863,f4101,f4389,f4590,f4605,f5227,f5390,f5396,f5392,f2212,f5618,f5621,f5631,f5634,f5648,f5651,f5891,f6317,f6329,f6332,f6336,f6338,f6330,f6372,f6375,f1687,f6474,f6475,f6484,f6485,f7148,f7478,f7482,f7480,f7479,f7737,f7740,f7750,f7753,f7768,f7772,f7775,f8047,f8050,f8057,f8060,f8082,f8110,f8136,f8141,f8184,f8204,f8205,f8206,f8207,f8510,f8513,f8520,f8522,f8531,f8534,f8539,f8541,f8587,f8929,f8931,f8938,f8940,f8949,f8953,f8955,f8511,f9074,f9045,f9044,f9043,f9042,f8788,f8787,f8802,f8800,f9608,f9911,f9924,f9925,f9941,f9943,f9946,f9949,f9952,f9959,f9965,f9966,f9974,f9976,f10006,f10007,f10009,f9919,f9078,f149,f10022,f9912,f10011,f10055,f9927,f9926,f10160,f10161,f10185,f10243,f10244,f10271,f10313,f10832,f10833,f11191,f11295,f11369,f11531,f11532,f12168,f12171,f12175,f12177,f12183,f12185,f12188,f12176,f12189,f12934,f12936,f12935,f12945,f13331,f13593,f13595,f13614,f13615,f14172,f14174,f14173,f14157,f12690,f11941,f11818,f8536,f11037,f8535,f14209,f14276,f14277,f14289,f14405,f14406,f14494,f14499,f14455,f14503,f14466,f14469,f14473,f14475,f14476,f14505,f14477,f14507,f14478,f14508,f14782,f14791,f14792,f14799,f14801,f14812,f15050,f15051,f15106,f15239,f15242,f15272,f15274,f15240,f15540,f15541,f12169,f15552,f15553,f12178,f15596,f15597,f15876,f15877,f16299,f16301,f16569,f16962,f16963,f16986,f17109,f16987,f14180,f17108,f17312,f17314,f17539,f17541,f17549,f16833,f16832,f16300,f12946,f16866,f16865,f17554,f16806,f17556,f16802,f17557,f16792,f16791,f17560,f17561,f15814,f16853,f17564,f17339,f17567,f17568,f17046,f17045,f17132,f17043,f17131,f17577,f17021,f17571,f17020,f17573,f17124,f17018,f17123,f17575,f17578,f16957,f17579,f17580,f16944,f17581,f17582,f16869,f17583,f16863,f16861,f16845,f16842,f15827,f15825,f15819,f15817,f11296,f17612,f17576,f17668,f17669,f14504,f17671,f17707,f17713,f17714,f17715,f17716,f17719,f17720,f17565,f17764,f17570,f17933,f17935,f17991,f17992,f17993,f17994,f14783,f17763,f17998,f18007,f17589,f17588,f12947,f16839,f17585,f16846,f16843,f18286,f18288,f18293,f18295,f18707,f18708,f17313,f19158,f14176,f19233,f19235,f19234,f19456,f19457,f17540,f19465,f17550,f19876,f19877,f12172,f19885,f13332,f20314,f20315,f20407,f20817,f20818,f20837,f20838,f13594,f20984,f20986,f20985,f20997,f20998,f21058,f21060,f21913,f21916,f21920,f21923,f21931,f22026,f22032,f22036,f22043,f22045,f22047,f22048,f21991,f10329,f21952,f21950,f21921,f21936,f21934,f21914,f22086,f21327,f21326,f22091,f21325,f21324,f22092,f20153,f20152,f22097,f20151,f20150,f22098,f18608,f18606,f22109,f22111,f15826,f15828,f22116,f22118,f22120,f22123,f22124,f22125,f22127,f22129,f22121,f22134,f22136,f17598,f22180,f22182,f10015,f10016,f22185,f22186,f22187,f22188,f22220,f22222,f9084,f9085,f22283,f9086,f9087,f22284,f15607,f15608,f22285,f15628,f15629,f22286,f15630,f15631,f22287,f15642,f15643,f22117,f6373,f17012,f17037,f17572,f22289,f18551,f18552,f22107,f21047,f17566,f21046,f21045,f21031,f21030,f17574,f22291,f19751,f19749,f22293,f15605,f15606,f22295,f15640,f15641,f22296,f6374,f19739,f19738,f22297,f19737,f19736,f22298,f22299,f22301,f22300,f19564,f19563,f22302,f19562,f19561,f22303,f22304,f22305,f19536,f19535,f22307,f22306,f18582,f18581,f17597,f18580,f18579,f22308,f18550,f18549,f22294,f18329,f18328,f17587,f18327,f18326,f17586,f18318,f18317,f22290,f18316,f18315,f17601,f18032,f18031,f22310,f18030,f18029,f22311,f14514,f14513,f14335,f14334,f1686,f14333,f14332,f22309,f22312,f22313,f22181,f10071,f10070,f10069,f10068,f10013,f10014,f21017,f21016,f8218,f21059,f21015,f21014,f22316,f14512,f14511,f8217,f10479,f10480,f8216,f10481,f10482,f10000,f10001,f172,f10002,f10003,f10234,f10235,f10236,f10237,f19537,f19538,f17590,f19710,f19709,f16844,f22318,f19708,f19707,f22319,f22292,f10020,f10018,f171,f22429,f22430,f9939,f9938,f22431,f22432,f9975,f10058,f22530,f22532,f22536,f22537,f22538,f22773,f22353,f22792,f11192,f21088,f21087,f22800,f19165,f19164,f22803,f17049,f17047,f17024,f22808,f17022,f22809,f22799,f22816,f22815,f22880,f22917,f22918,f173,f22919,f22920,f174,f22928,f22929,f6483,f22930,f22931,f9076,f22926,f23115,f23123,f23122,f23121,f23138,f23352,f23355,f23731,f23742,f23772,f23781,f23807,f23893,f23900,f23902,f23782,f24141,f10008,f24149,f23733,f1903,f24180,f23732,f1899,f24252,f24259,f24364,f24365,f24367,f24368,f24370,f24374,f24378,f24379,f24381,f24383,f24622,f24623,f9079,f24702,f24703,f23773,f24931,f24932,f24701,f24686,f24934,f24936,f24935,f22927,f24251,f24236,f23693,f23776,f24938,f23774,f10024,f23771,f23686,f22044,f10023,f23767,f23765,f23683,f23682,f23741,f23646,f24943,f23645,f24944,f23644,f1917,f23641,f10322,f10320,f23640,f23639,f23725,f23723,f23630,f23629,f24939,f8215,f23710,f23708,f22317,f25051,f25055,f25198,f24933,f552,f25244,f25504,f25505,f17760,f2006,f25506,f25507,f17762,f17761,f25604,f25824,f25825])).
fof(f25825,plain,(
  $false | ($spl0 | $spl3 | $spl84 | $spl102)),
  inference(subsumption_resolution,[],[f25818,f8216])).
fof(f25818,plain,(
  ~p(s(bool,sK6(sK1))) | ($spl0 | $spl3 | $spl102)),
  inference(trivial_inequality_removal,[],[f25813])).
fof(f25813,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ~p(s(bool,sK6(sK1))) | ($spl0 | $spl3 | $spl102)),
  inference(superposition,[],[f9975,f22027])).
fof(f22027,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f21886,f142])).
fof(f21886,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | $spl0),
  inference(superposition,[],[f139,f1302])).
fof(f1302,plain,(
  ( ! [X6,X4,X5,X3] : (s(X4,happ(s(fun(bool,X4),X5),s(bool,sK6(X3)))) = s(X4,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X3),s(fun(bool,X4),X5),s(X4,X6))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X3)) )),
  inference(superposition,[],[f133,f132])).
fof(f132,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X0))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X0)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f111])).
fof(f111,plain,(
  ! [X0] : (? [X1] : s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X84] : (? [X11] : s('type.semanticPrimitives.eq_result',X84) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X11))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X84))),
  file('thm.closProps.eqs|split|6','thm.semanticPrimitives.eq_result_nchotomy')).
fof(f133,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))),s(fun(bool,X0),X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))),s(fun(bool,X0),X2),s(X0,X3)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))),s(fun(bool,X0),X2),s(X0,X3)))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X6] : ! [X14,X18,X9] : s(X6,happ(s(fun(bool,X6),X18),s(bool,X14))) = s(X6,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X14))),s(fun(bool,X6),X18),s(X6,X9)))),
  file('thm.closProps.eqs|split|6','thm.semanticPrimitives.eq_result_case_def|split|0')).
fof(f25824,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f25819,f8084])).
fof(f8084,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | $spl3),
  inference(trivial_inequality_removal,[],[f8083])).
fof(f8083,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | $spl3),
  inference(superposition,[],[f142,f132])).
fof(f25819,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | ($spl0 | $spl3 | $spl6)),
  inference(trivial_inequality_removal,[],[f25812])).
fof(f25812,plain,(
  s(sK0,sK2) != s(sK0,sK2) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | ($spl0 | $spl3 | $spl6)),
  inference(superposition,[],[f145,f22027])).
fof(f25604,plain,(
  $false | ($spl13 | $spl77 | $spl224)),
  inference(resolution,[],[f25595,f171])).
fof(f25595,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | ($spl13 | $spl77 | $spl224)),
  inference(subsumption_resolution,[],[f17601,f25251])).
fof(f25251,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5))))) ) | ($spl13 | $spl77)),
  inference(resolution,[],[f25052,f22321])).
fof(f22321,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,sK5) = s(bool,X0)) ) | $spl13),
  inference(resolution,[],[f171,f135])).
fof(f135,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X15,X16] : ((p(s(bool,X15)) => p(s(bool,X16))) => ((p(s(bool,X16)) => p(s(bool,X15))) => s(bool,X15) = s(bool,X16)))),
  file('thm.closProps.eqs|split|6','thm.bool.IMP_ANTISYM_AX')).
fof(f25052,plain,(
  ( ! [X0] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) | p(s(bool,X0))) ) | ($spl13 | $spl77)),
  inference(subsumption_resolution,[],[f25040,f171])).
fof(f25040,plain,(
  ( ! [X0] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) | p(s(bool,X0)) | p(s(bool,sK5))) ) | $spl77),
  inference(superposition,[],[f6373,f135])).
fof(f17761,plain,(
  ( ! [X8] : (~p(s(bool,sK6(X8))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X8)) ) | $spl230),
  inference(cnf_transformation,[],[f17761_D])).
fof(f17761_D,plain,(
  ( ! [X8] : (~p(s(bool,sK6(X8))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X8)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f17762,plain,(
  ( ! [X9] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X9))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))))) | ~p(s(bool,X9))) ) | $spl232),
  inference(cnf_transformation,[],[f17762_D])).
fof(f17762_D,plain,(
  ( ! [X9] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X9))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))))) | ~p(s(bool,X9))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f25507,plain,(
  ( ! [X8,X7] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X8)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X7) | ~p(s(bool,X8)) | ~p(s(bool,sK6(X7)))) )),
  inference(superposition,[],[f1259,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f119])).
fof(f1259,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1))))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1)) )),
  inference(duplicate_literal_removal,[],[f1242])).
fof(f1242,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))))))) )),
  inference(superposition,[],[f132,f132])).
fof(f25506,plain,(
  ( ! [X6,X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X6)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X5) | ~p(s(bool,sK6(X5))) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f1259,f138])).
fof(f2006,plain,(
  ( ! [X0] : (p(s(bool,sK6(X0))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | $spl56),
  inference(cnf_transformation,[],[f2006_D])).
fof(f2006_D,plain,(
  ( ! [X0] : (p(s(bool,sK6(X0))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f17760,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))))) | p(s(bool,X5))) ) | $spl228),
  inference(cnf_transformation,[],[f17760_D])).
fof(f17760_D,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))))) | p(s(bool,X5))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f25505,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X4)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X3) | p(s(bool,sK6(X3))) | p(s(bool,X4))) )),
  inference(superposition,[],[f1259,f135])).
fof(f25504,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) | p(s(bool,X2)) | p(s(bool,sK6(X1)))) )),
  inference(superposition,[],[f1259,f135])).
fof(f25244,plain,(
  $false | ($spl13 | $spl56 | $spl77 | $spl346)),
  inference(subsumption_resolution,[],[f25243,f24935])).
fof(f25243,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl13 | $spl56 | $spl77)),
  inference(subsumption_resolution,[],[f25235,f171])).
fof(f25235,plain,(
  p(s(bool,sK5)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl13 | $spl56 | $spl77)),
  inference(superposition,[],[f2006,f25038])).
fof(f25038,plain,(
  s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ($spl13 | $spl77)),
  inference(resolution,[],[f6373,f22321])).
fof(f552,plain,(
  ( ! [X41] : (s(bool,sK5) = s(bool,X41)) ) | $spl18),
  inference(cnf_transformation,[],[f552_D])).
fof(f552_D,plain,(
  ( ! [X41] : (s(bool,sK5) = s(bool,X41)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f24933,plain,(
  ( ! [X26,X25] : (s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) ) | $spl344),
  inference(cnf_transformation,[],[f24933_D])).
fof(f24933_D,plain,(
  ( ! [X26,X25] : (s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f25198,plain,(
  ( ! [X21,X19,X20] : (s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X21))) | s(bool,sK5) = s(bool,X19)) ) | ($spl13 | $spl30 | $spl84)),
  inference(subsumption_resolution,[],[f25189,f22321])).
fof(f25189,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X21))) | s(bool,sK5) = s(bool,X19)) ) | ($spl13 | $spl30 | $spl84)),
  inference(superposition,[],[f4524,f22677])).
fof(f22677,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X4)) ) | ($spl13 | $spl84)),
  inference(resolution,[],[f22332,f22321])).
fof(f22332,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,sK6(sK1))) ) | $spl84),
  inference(resolution,[],[f8216,f138])).
fof(f4524,plain,(
  ( ! [X43,X46,X44] : (~p(s(bool,X46)) | s(bool,X46) = s(bool,happ(s(fun(bool,bool),X44),s(bool,X43)))) ) | $spl30),
  inference(forward_demodulation,[],[f4523,f133])).
fof(f4523,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X46) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X43))),s(fun(bool,bool),X44),s(bool,X45))) | ~p(s(bool,X46))) ) | $spl30),
  inference(subsumption_resolution,[],[f4501,f1886])).
fof(f4501,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X44),s(bool,X43)))) | s(bool,X46) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X43))),s(fun(bool,bool),X44),s(bool,X45))) | ~p(s(bool,X46))) )),
  inference(superposition,[],[f138,f133])).
fof(f25055,plain,(
  $false | ($spl13 | $spl26 | $spl77 | $spl84 | $spl216 | $spl268)),
  inference(subsumption_resolution,[],[f25054,f22850])).
fof(f22850,plain,(
  ( ! [X4] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X4))))) ) | ($spl26 | $spl84 | $spl216 | $spl268)),
  inference(resolution,[],[f22845,f1686])).
fof(f22845,plain,(
  ( ! [X2] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2)))))) ) | ($spl84 | $spl216 | $spl268)),
  inference(subsumption_resolution,[],[f22844,f8216])).
fof(f22844,plain,(
  ( ! [X2] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | ~p(s(bool,sK6(sK1)))) ) | ($spl216 | $spl268)),
  inference(subsumption_resolution,[],[f22836,f22283])).
fof(f22836,plain,(
  ( ! [X2] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X2))) ) | $spl216),
  inference(superposition,[],[f17586,f138])).
fof(f25054,plain,(
  s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ($spl13 | $spl77 | $spl84)),
  inference(subsumption_resolution,[],[f25048,f8216])).
fof(f25048,plain,(
  ~p(s(bool,sK6(sK1))) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ($spl13 | $spl77 | $spl84)),
  inference(superposition,[],[f6373,f22677])).
fof(f25051,plain,(
  $false | ($spl13 | $spl26 | $spl77 | $spl84 | $spl216 | $spl268)),
  inference(subsumption_resolution,[],[f25038,f22850])).
fof(f22317,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,sK6(sK1)) | ~p(s(bool,X4))) ) | $spl318),
  inference(cnf_transformation,[],[f22317_D])).
fof(f22317_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,sK6(sK1)) | ~p(s(bool,X4))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f23708,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X4) | ~p(s(bool,X5))) ) | ($spl13 | $spl84)),
  inference(subsumption_resolution,[],[f23618,f22321])).
fof(f23618,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f22677,f138])).
fof(f23710,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X6) | ~p(s(bool,X7))) ) | ($spl13 | $spl84)),
  inference(subsumption_resolution,[],[f23619,f22321])).
fof(f23619,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f22677,f138])).
fof(f8215,plain,(
  ~p(s(bool,sK6(sK1))) | $spl85),
  inference(cnf_transformation,[],[f8215_D])).
fof(f8215_D,plain,(
  ~p(s(bool,sK6(sK1))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f24939,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,X19) | s(bool,sK5) = s(bool,X18) | ~p(s(bool,X19))) ) | $spl350),
  inference(cnf_transformation,[],[f24939_D])).
fof(f24939_D,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,X19) | s(bool,sK5) = s(bool,X18) | ~p(s(bool,X19))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f23629,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | s(bool,sK5) = s(bool,X8) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X7))) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f22677,f138])).
fof(f23630,plain,(
  ( ! [X10,X9] : (s(bool,X9) = s(bool,X10) | s(bool,sK5) = s(bool,X10) | ~p(s(bool,X9)) | ~p(s(bool,sK6(sK1)))) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f22677,f138])).
fof(f23723,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,sK6(sK1)) | ~p(s(bool,X13)) | s(bool,sK5) = s(bool,X12)) ) | ($spl13 | $spl84)),
  inference(subsumption_resolution,[],[f23637,f22321])).
fof(f23637,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,sK6(sK1)) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,sK5) = s(bool,X12)) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f138,f22677])).
fof(f23725,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,sK6(sK1)) | ~p(s(bool,X15)) | s(bool,sK5) = s(bool,X14)) ) | ($spl13 | $spl84)),
  inference(subsumption_resolution,[],[f23638,f22321])).
fof(f23638,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,sK6(sK1)) | ~p(s(bool,X15)) | ~p(s(bool,X14)) | s(bool,sK5) = s(bool,X14)) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f138,f22677])).
fof(f23639,plain,(
  ( ! [X17,X16] : (~p(s(bool,sK6(sK1))) | s(bool,X16) = s(bool,X17) | ~p(s(bool,X17)) | s(bool,sK5) = s(bool,X16)) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f138,f22677])).
fof(f23640,plain,(
  ( ! [X19,X18] : (~p(s(bool,sK6(sK1))) | s(bool,X18) = s(bool,X19) | ~p(s(bool,X19)) | s(bool,sK5) = s(bool,X18)) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f138,f22677])).
fof(f10320,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))) | s(bool,sK5) = s(bool,X16)) ) | $spl112),
  inference(cnf_transformation,[],[f10320_D])).
fof(f10320_D,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))) | s(bool,sK5) = s(bool,X16)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f10322,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | $spl115),
  inference(cnf_transformation,[],[f10322_D])).
fof(f10322_D,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f23641,plain,(
  ( ! [X20] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X20)) ) | ($spl6 | $spl13 | $spl84)),
  inference(superposition,[],[f145,f22677])).
fof(f1917,plain,(
  ( ! [X57,X54,X55] : (s(bool,X57) = s(bool,happ(s(fun(bool,bool),X54),s(bool,X55))) | s(bool,sK5) = s(bool,X57)) ) | $spl48),
  inference(cnf_transformation,[],[f1917_D])).
fof(f1917_D,plain,(
  ( ! [X57,X54,X55] : (s(bool,X57) = s(bool,happ(s(fun(bool,bool),X54),s(bool,X55))) | s(bool,sK5) = s(bool,X57)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f23644,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,sK6(sK1))) | s(bool,X24) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,sK5) = s(bool,X24)) ) | ($spl13 | $spl42 | $spl84)),
  inference(superposition,[],[f1902,f22677])).
fof(f24944,plain,(
  ( ! [X27] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) | s(bool,sK5) = s(bool,X27)) ) | $spl358),
  inference(cnf_transformation,[],[f24944_D])).
fof(f24944_D,plain,(
  ( ! [X27] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) | s(bool,sK5) = s(bool,X27)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f23645,plain,(
  ( ! [X27] : (~p(s(bool,sK6(sK1))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) | s(bool,sK5) = s(bool,X27)) ) | ($spl13 | $spl84 | $spl88)),
  inference(superposition,[],[f8218,f22677])).
fof(f24943,plain,(
  ( ! [X28] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X28))) | s(bool,sK5) = s(bool,X28)) ) | $spl356),
  inference(cnf_transformation,[],[f24943_D])).
fof(f24943_D,plain,(
  ( ! [X28] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X28))) | s(bool,sK5) = s(bool,X28)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f23646,plain,(
  ( ! [X28] : (~p(s(bool,sK6(sK1))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X28))) | s(bool,sK5) = s(bool,X28)) ) | ($spl13 | $spl84 | $spl94)),
  inference(superposition,[],[f9076,f22677])).
fof(f23741,plain,(
  ( ! [X30] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X30)) ) | ($spl13 | $spl84 | $spl102)),
  inference(subsumption_resolution,[],[f23648,f22321])).
fof(f23648,plain,(
  ( ! [X30] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | ~p(s(bool,X30)) | s(bool,sK5) = s(bool,X30)) ) | ($spl13 | $spl84 | $spl102)),
  inference(superposition,[],[f9975,f22677])).
fof(f23682,plain,(
  ( ! [X17,X16] : (s(bool,X16) = s(bool,X17) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X17)) | s(bool,sK5) = s(bool,X16)) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f138,f22677])).
fof(f23683,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,sK6(sK1))) | s(bool,sK5) = s(bool,X18)) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f138,f22677])).
fof(f23765,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,sK6(sK1)) | ~p(s(bool,X21)) | s(bool,sK5) = s(bool,X20)) ) | ($spl13 | $spl84)),
  inference(subsumption_resolution,[],[f23684,f22321])).
fof(f23684,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,X21) = s(bool,sK6(sK1)) | ~p(s(bool,X21)) | s(bool,sK5) = s(bool,X20)) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f138,f22677])).
fof(f23767,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,sK6(sK1)) | ~p(s(bool,X23)) | s(bool,sK5) = s(bool,X22)) ) | ($spl13 | $spl84)),
  inference(subsumption_resolution,[],[f23685,f22321])).
fof(f23685,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) = s(bool,sK6(sK1)) | ~p(s(bool,X23)) | s(bool,sK5) = s(bool,X22)) ) | ($spl13 | $spl84)),
  inference(superposition,[],[f138,f22677])).
fof(f10023,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | $spl109),
  inference(cnf_transformation,[],[f10023_D])).
fof(f10023_D,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f22044,plain,(
  ( ! [X57] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X57))) | s(bool,sK5) = s(bool,X57)) ) | $spl250),
  inference(cnf_transformation,[],[f22044_D])).
fof(f22044_D,plain,(
  ( ! [X57] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X57))) | s(bool,sK5) = s(bool,X57)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f23686,plain,(
  ( ! [X24] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X24))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X24)) ) | ($spl6 | $spl13 | $spl84)),
  inference(superposition,[],[f145,f22677])).
fof(f23771,plain,(
  ( ! [X30,X28,X29] : (s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,sK5) = s(bool,X28)) ) | ($spl13 | $spl42 | $spl84)),
  inference(subsumption_resolution,[],[f23689,f22321])).
fof(f23689,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,sK5) = s(bool,X28)) ) | ($spl13 | $spl42 | $spl84)),
  inference(superposition,[],[f1902,f22677])).
fof(f10024,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | $spl108),
  inference(cnf_transformation,[],[f10024_D])).
fof(f10024_D,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f23774,plain,(
  ( ! [X31] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X31)) ) | ($spl13 | $spl84 | $spl88)),
  inference(subsumption_resolution,[],[f23690,f22321])).
fof(f23690,plain,(
  ( ! [X31] : (~p(s(bool,X31)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X31)) ) | ($spl13 | $spl84 | $spl88)),
  inference(superposition,[],[f8218,f22677])).
fof(f24938,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | $spl349),
  inference(cnf_transformation,[],[f24938_D])).
fof(f24938_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f23776,plain,(
  ( ! [X32] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X32)) ) | ($spl13 | $spl84 | $spl94)),
  inference(subsumption_resolution,[],[f23691,f22321])).
fof(f23691,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X32)) ) | ($spl13 | $spl84 | $spl94)),
  inference(superposition,[],[f9076,f22677])).
fof(f23693,plain,(
  ( ! [X34] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X34))) | ~p(s(bool,sK6(sK1))) | s(bool,sK5) = s(bool,X34)) ) | ($spl13 | $spl84 | $spl102)),
  inference(superposition,[],[f9975,f22677])).
fof(f24236,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,sK6(sK1))) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | s(bool,sK5) = s(bool,X16)) ) | ($spl13 | $spl42 | $spl84)),
  inference(superposition,[],[f1902,f22677])).
fof(f24251,plain,(
  ( ! [X24,X23,X22] : (s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X24))) | s(bool,sK5) = s(bool,X22)) ) | ($spl13 | $spl42 | $spl84)),
  inference(subsumption_resolution,[],[f24238,f22321])).
fof(f24238,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X24))) | s(bool,sK5) = s(bool,X22)) ) | ($spl13 | $spl42 | $spl84)),
  inference(superposition,[],[f1902,f22677])).
fof(f22927,plain,(
  p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | $spl336),
  inference(cnf_transformation,[],[f22927_D])).
fof(f22927_D,plain,(
  p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f24935,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | $spl346),
  inference(cnf_transformation,[],[f24935_D])).
fof(f24935_D,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f24936,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) ) | ($spl94 | $spl327)),
  inference(subsumption_resolution,[],[f24482,f9076])).
fof(f24482,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,X1)) | p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) ) | $spl327),
  inference(superposition,[],[f22798,f135])).
fof(f22798,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | $spl327),
  inference(cnf_transformation,[],[f22798_D])).
fof(f22798_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f24934,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) ) | ($spl94 | $spl327)),
  inference(subsumption_resolution,[],[f24483,f9076])).
fof(f24483,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | p(s(bool,X2))) ) | $spl327),
  inference(superposition,[],[f22798,f135])).
fof(f24686,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,sK6(sK1))) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | s(bool,sK5) = s(bool,X16)) ) | ($spl13 | $spl42 | $spl84)),
  inference(superposition,[],[f1902,f22677])).
fof(f24701,plain,(
  ( ! [X26,X24,X25] : (s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,sK5) = s(bool,X24)) ) | ($spl13 | $spl42 | $spl84)),
  inference(subsumption_resolution,[],[f24689,f22321])).
fof(f24689,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,sK5) = s(bool,X24)) ) | ($spl13 | $spl42 | $spl84)),
  inference(superposition,[],[f1902,f22677])).
fof(f24932,plain,(
  $false | ($spl0 | $spl3 | $spl84 | $spl102 | $spl336)),
  inference(subsumption_resolution,[],[f24925,f22927])).
fof(f24925,plain,(
  ~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | ($spl0 | $spl3 | $spl84 | $spl102 | $spl336)),
  inference(trivial_inequality_removal,[],[f24920])).
fof(f24920,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | ($spl0 | $spl3 | $spl84 | $spl102 | $spl336)),
  inference(superposition,[],[f9975,f24909])).
fof(f24909,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | ($spl0 | $spl3 | $spl84 | $spl336)),
  inference(forward_demodulation,[],[f22027,f22939])).
fof(f22939,plain,(
  s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,sK6(sK1)) | ($spl84 | $spl336)),
  inference(resolution,[],[f22927,f22332])).
fof(f24931,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl84 | $spl336)),
  inference(subsumption_resolution,[],[f24926,f22950])).
fof(f22950,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | ($spl3 | $spl84 | $spl336)),
  inference(backward_demodulation,[],[f22939,f8084])).
fof(f24926,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | ($spl0 | $spl3 | $spl6 | $spl84 | $spl336)),
  inference(trivial_inequality_removal,[],[f24919])).
fof(f24919,plain,(
  s(sK0,sK2) != s(sK0,sK2) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | ($spl0 | $spl3 | $spl6 | $spl84 | $spl336)),
  inference(superposition,[],[f145,f24909])).
fof(f23773,plain,(
  ( ! [X30,X29] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) ) | $spl340),
  inference(cnf_transformation,[],[f23773_D])).
fof(f23773_D,plain,(
  ( ! [X30,X29] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f24703,plain,(
  ( ! [X28,X29,X27] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X28),s(bool,X29))) | s(bool,sK5) = s(bool,X27)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(subsumption_resolution,[],[f24690,f22321])).
fof(f24690,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X28),s(bool,X29))) | s(bool,sK5) = s(bool,X27)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(superposition,[],[f1902,f23035])).
fof(f23035,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) | s(bool,sK5) = s(bool,X5)) ) | ($spl13 | $spl84 | $spl336)),
  inference(resolution,[],[f22953,f22321])).
fof(f22953,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) ) | ($spl84 | $spl336)),
  inference(backward_demodulation,[],[f22939,f22332])).
fof(f24702,plain,(
  ( ! [X26,X24,X25] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,sK5) = s(bool,X24)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(forward_demodulation,[],[f24701,f22939])).
fof(f9079,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) ) | $spl96),
  inference(cnf_transformation,[],[f9079_D])).
fof(f9079_D,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f24623,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | p(s(bool,sK6(X4)))) ) | ($spl3 | $spl84 | $spl88 | $spl216 | $spl268 | $spl336)),
  inference(subsumption_resolution,[],[f24604,f8218])).
fof(f24604,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | p(s(bool,sK6(X4))) | p(s(bool,X5))) ) | ($spl3 | $spl84 | $spl216 | $spl268 | $spl336)),
  inference(superposition,[],[f23904,f135])).
fof(f23904,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1)) ) | ($spl3 | $spl84 | $spl216 | $spl268 | $spl336)),
  inference(forward_demodulation,[],[f23903,f22950])).
fof(f23903,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1)) ) | ($spl84 | $spl216 | $spl268 | $spl336)),
  inference(forward_demodulation,[],[f1259,f23038])).
fof(f23038,plain,(
  ( ! [X8] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X8))))) ) | ($spl84 | $spl216 | $spl268 | $spl336)),
  inference(resolution,[],[f22953,f22845])).
fof(f24622,plain,(
  ( ! [X2,X3] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X2) | p(s(bool,sK6(X2)))) ) | ($spl3 | $spl84 | $spl88 | $spl216 | $spl268 | $spl336)),
  inference(subsumption_resolution,[],[f24603,f8218])).
fof(f24603,plain,(
  ( ! [X2,X3] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X2) | p(s(bool,X3)) | p(s(bool,sK6(X2)))) ) | ($spl3 | $spl84 | $spl216 | $spl268 | $spl336)),
  inference(superposition,[],[f23904,f135])).
fof(f24383,plain,(
  $false | ($spl26 | $spl84 | $spl216 | $spl220 | $spl268 | $spl322 | $spl336)),
  inference(subsumption_resolution,[],[f24382,f24221])).
fof(f24221,plain,(
  ( ! [X9] : (s(bool,sK5) = s(bool,X9)) ) | ($spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(subsumption_resolution,[],[f24220,f22927])).
fof(f24220,plain,(
  ( ! [X9] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | s(bool,sK5) = s(bool,X9)) ) | ($spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(forward_demodulation,[],[f17597,f23038])).
fof(f24382,plain,(
  ( ! [X3] : (s(bool,sK5) != s(bool,X3)) ) | ($spl26 | $spl84 | $spl216 | $spl220 | $spl268 | $spl322 | $spl336)),
  inference(subsumption_resolution,[],[f24275,f1686])).
fof(f24275,plain,(
  ( ! [X3] : (s(bool,sK5) != s(bool,X3) | p(s(bool,X3))) ) | ($spl84 | $spl216 | $spl220 | $spl268 | $spl322 | $spl336)),
  inference(backward_demodulation,[],[f24221,f23217])).
fof(f23217,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) ) | ($spl84 | $spl322 | $spl336)),
  inference(forward_demodulation,[],[f22319,f22939])).
fof(f24381,plain,(
  $false | ($spl3 | $spl11 | $spl76 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(subsumption_resolution,[],[f24380,f142])).
fof(f24380,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl11 | $spl76 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(forward_demodulation,[],[f24273,f22786])).
fof(f22786,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl3 | $spl11 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f22785,f148])).
fof(f148,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | $spl11),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f22785,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl3 | $spl76 | $spl84)),
  inference(forward_demodulation,[],[f22744,f8084])).
fof(f22744,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl76 | $spl84)),
  inference(superposition,[],[f132,f22680])).
fof(f22680,plain,(
  s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ($spl76 | $spl84)),
  inference(resolution,[],[f22332,f6374])).
fof(f24273,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl3 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(backward_demodulation,[],[f24221,f22950])).
fof(f24379,plain,(
  $false | ($spl13 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(subsumption_resolution,[],[f24271,f171])).
fof(f24271,plain,(
  p(s(bool,sK5)) | ($spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(backward_demodulation,[],[f24221,f22927])).
fof(f24378,plain,(
  $false | ($spl3 | $spl11 | $spl76 | $spl84 | $spl216 | $spl220 | $spl268 | $spl327 | $spl336)),
  inference(subsumption_resolution,[],[f24270,f22786])).
fof(f24270,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl84 | $spl216 | $spl220 | $spl268 | $spl327 | $spl336)),
  inference(backward_demodulation,[],[f24221,f22798])).
fof(f24374,plain,(
  $false | ($spl13 | $spl30 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(subsumption_resolution,[],[f24264,f171])).
fof(f24264,plain,(
  p(s(bool,sK5)) | ($spl30 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(backward_demodulation,[],[f24221,f1886])).
fof(f24370,plain,(
  $false | ($spl26 | $spl30 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(trivial_inequality_removal,[],[f24344])).
fof(f24344,plain,(
  s(bool,sK5) != s(bool,sK5) | ($spl26 | $spl30 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(superposition,[],[f22658,f24221])).
fof(f22658,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,happ(s(fun(bool,bool),X6),s(bool,X7)))) ) | ($spl26 | $spl30)),
  inference(resolution,[],[f1886,f1686])).
fof(f24368,plain,(
  $false | ($spl26 | $spl76 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(trivial_inequality_removal,[],[f24350])).
fof(f24350,plain,(
  s(bool,sK5) != s(bool,sK5) | ($spl26 | $spl76 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(superposition,[],[f22414,f24221])).
fof(f22414,plain,(
  s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ($spl26 | $spl76)),
  inference(resolution,[],[f1686,f6374])).
fof(f24367,plain,(
  $false | ($spl26 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(trivial_inequality_removal,[],[f24351])).
fof(f24351,plain,(
  s(bool,sK5) != s(bool,sK5) | ($spl26 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(superposition,[],[f22850,f24221])).
fof(f24365,plain,(
  $false | ($spl26 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(trivial_inequality_removal,[],[f24353])).
fof(f24353,plain,(
  s(bool,sK5) != s(bool,sK5) | ($spl26 | $spl84 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(superposition,[],[f22937,f24221])).
fof(f22937,plain,(
  s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) | ($spl26 | $spl336)),
  inference(resolution,[],[f22927,f1686])).
fof(f24364,plain,(
  $false | ($spl84 | $spl93 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(trivial_inequality_removal,[],[f24354])).
fof(f24354,plain,(
  s(bool,sK5) != s(bool,sK5) | ($spl84 | $spl93 | $spl216 | $spl220 | $spl268 | $spl336)),
  inference(superposition,[],[f8535,f24221])).
fof(f24259,plain,(
  ( ! [X33,X34,X32] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X33),s(bool,X34))) | s(bool,sK5) = s(bool,X32)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(subsumption_resolution,[],[f24242,f22321])).
fof(f24242,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X33),s(bool,X34))) | s(bool,sK5) = s(bool,X32)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(superposition,[],[f1902,f23035])).
fof(f24252,plain,(
  ( ! [X24,X23,X22] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X24))) | s(bool,sK5) = s(bool,X22)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(forward_demodulation,[],[f24251,f22939])).
fof(f1899,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) | $spl40),
  inference(cnf_transformation,[],[f1899_D])).
fof(f1899_D,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f23732,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | $spl338),
  inference(cnf_transformation,[],[f23732_D])).
fof(f23732_D,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f24180,plain,(
  ( ! [X35] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | p(s(bool,X35))) ) | ($spl8 | $spl13 | $spl84 | $spl142 | $spl336)),
  inference(forward_demodulation,[],[f24116,f22939])).
fof(f24116,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1))))) ) | ($spl8 | $spl13 | $spl84 | $spl142)),
  inference(superposition,[],[f22351,f24026])).
fof(f24026,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,sK6(sK1))) ) | ($spl13 | $spl84 | $spl142)),
  inference(subsumption_resolution,[],[f22324,f22677])).
fof(f22324,plain,(
  ( ! [X3] : (s(bool,sK5) != s(bool,X3) | s(bool,X3) = s(bool,sK6(sK1))) ) | ($spl13 | $spl142)),
  inference(resolution,[],[f171,f14173])).
fof(f22351,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X0)))) ) | ($spl8 | $spl13)),
  inference(subsumption_resolution,[],[f22341,f171])).
fof(f22341,plain,(
  ( ! [X0] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK5))) ) | $spl8),
  inference(superposition,[],[f147,f135])).
fof(f1903,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) | $spl44),
  inference(cnf_transformation,[],[f1903_D])).
fof(f1903_D,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f23733,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | $spl339),
  inference(cnf_transformation,[],[f23733_D])).
fof(f23733_D,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f24149,plain,(
  ( ! [X25] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | ~p(s(bool,X25))) ) | ($spl13 | $spl84 | $spl102 | $spl142 | $spl336)),
  inference(forward_demodulation,[],[f24063,f22939])).
fof(f24063,plain,(
  ( ! [X25] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | ~p(s(bool,X25))) ) | ($spl13 | $spl84 | $spl102 | $spl142)),
  inference(superposition,[],[f9975,f24026])).
fof(f10008,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4)))) ) | $spl104),
  inference(cnf_transformation,[],[f10008_D])).
fof(f10008_D,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4)))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f24141,plain,(
  ( ! [X20] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20)))) ) | ($spl6 | $spl13 | $spl84 | $spl142 | $spl336)),
  inference(forward_demodulation,[],[f24058,f22939])).
fof(f24058,plain,(
  ( ! [X20] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20)))) ) | ($spl6 | $spl13 | $spl84 | $spl142)),
  inference(superposition,[],[f145,f24026])).
fof(f23782,plain,(
  ( ! [X43] : (s('type.semanticPrimitives.eq_result',X43) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X43)) ) | $spl342),
  inference(cnf_transformation,[],[f23782_D])).
fof(f23782_D,plain,(
  ( ! [X43] : (s('type.semanticPrimitives.eq_result',X43) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X43)) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f23902,plain,(
  ( ! [X45,X46] : (s('type.semanticPrimitives.eq_result',X46) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X46) | s(bool,sK5) = s(bool,X45)) ) | ($spl3 | $spl11 | $spl13 | $spl76 | $spl84 | $spl336)),
  inference(subsumption_resolution,[],[f23872,f22321])).
fof(f23872,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | s('type.semanticPrimitives.eq_result',X46) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X46) | s(bool,sK5) = s(bool,X45)) ) | ($spl3 | $spl11 | $spl13 | $spl76 | $spl84 | $spl336)),
  inference(superposition,[],[f23325,f23035])).
fof(f23325,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4)) ) | ($spl3 | $spl11 | $spl13 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f1236,f23307])).
fof(f23307,plain,(
  ( ! [X16] : (p(s(bool,sK6(X16))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16)) ) | ($spl3 | $spl11 | $spl13 | $spl76 | $spl84)),
  inference(duplicate_literal_removal,[],[f23306])).
fof(f23306,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16) | p(s(bool,sK6(X16))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16)) ) | ($spl3 | $spl11 | $spl13 | $spl76 | $spl84)),
  inference(forward_demodulation,[],[f23277,f22786])).
fof(f23277,plain,(
  ( ! [X16] : (p(s(bool,sK6(X16))) | s('type.semanticPrimitives.eq_result',X16) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16)) ) | $spl13),
  inference(resolution,[],[f1234,f171])).
fof(f1234,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,sK6(X0))) | s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) )),
  inference(superposition,[],[f132,f135])).
fof(f1236,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK6(X4))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | ~p(s(bool,X5))) )),
  inference(superposition,[],[f132,f138])).
fof(f23900,plain,(
  ( ! [X30,X31,X32] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X31),s(bool,X32))) | s(bool,sK5) = s(bool,X30)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(subsumption_resolution,[],[f23860,f22321])).
fof(f23860,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X31),s(bool,X32))) | s(bool,sK5) = s(bool,X30)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(superposition,[],[f1902,f23035])).
fof(f23893,plain,(
  ( ! [X30] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s(bool,sK5) = s(bool,X30)) ) | ($spl13 | $spl84 | $spl102 | $spl336)),
  inference(subsumption_resolution,[],[f23814,f22321])).
fof(f23814,plain,(
  ( ! [X30] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | ~p(s(bool,X30)) | s(bool,sK5) = s(bool,X30)) ) | ($spl13 | $spl84 | $spl102 | $spl336)),
  inference(superposition,[],[f9975,f23035])).
fof(f23807,plain,(
  ( ! [X20] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X20)) ) | ($spl6 | $spl13 | $spl84 | $spl336)),
  inference(superposition,[],[f145,f23035])).
fof(f23781,plain,(
  ( ! [X43,X42] : (s('type.semanticPrimitives.eq_result',X43) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X43) | s(bool,sK5) = s(bool,X42)) ) | ($spl3 | $spl11 | $spl13 | $spl76 | $spl84 | $spl336)),
  inference(forward_demodulation,[],[f23780,f22939])).
fof(f23780,plain,(
  ( ! [X43,X42] : (s('type.semanticPrimitives.eq_result',X43) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X43) | s(bool,sK5) = s(bool,X42)) ) | ($spl3 | $spl11 | $spl13 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f23701,f22321])).
fof(f23701,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | s('type.semanticPrimitives.eq_result',X43) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X43) | s(bool,sK5) = s(bool,X42)) ) | ($spl3 | $spl11 | $spl13 | $spl76 | $spl84)),
  inference(superposition,[],[f23325,f22677])).
fof(f23772,plain,(
  ( ! [X30,X28,X29] : (s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,sK5) = s(bool,X28)) ) | ($spl13 | $spl42 | $spl84 | $spl336)),
  inference(forward_demodulation,[],[f23771,f22939])).
fof(f23742,plain,(
  ( ! [X30] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s(bool,sK5) = s(bool,X30)) ) | ($spl13 | $spl84 | $spl102 | $spl336)),
  inference(forward_demodulation,[],[f23741,f22939])).
fof(f23731,plain,(
  ( ! [X20] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X20)) ) | ($spl6 | $spl13 | $spl84 | $spl336)),
  inference(forward_demodulation,[],[f23641,f22939])).
fof(f23355,plain,(
  $false | ($spl3 | $spl11 | $spl76 | $spl84 | $spl88 | $spl244)),
  inference(subsumption_resolution,[],[f23345,f142])).
fof(f23345,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl11 | $spl76 | $spl84 | $spl88 | $spl244)),
  inference(superposition,[],[f22786,f23326])).
fof(f23326,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | ($spl3 | $spl88 | $spl244)),
  inference(subsumption_resolution,[],[f22315,f8218])).
fof(f22315,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0))) ) | ($spl3 | $spl244)),
  inference(subsumption_resolution,[],[f22314,f142])).
fof(f22314,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0))) ) | $spl244),
  inference(forward_demodulation,[],[f21112,f131])).
fof(f131,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X1),s(X0,X2)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X1),s(X0,X2)))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6] : ! [X18,X9] : s(X6,X9) = s(X6,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X6),X18),s(X6,X9)))),
  file('thm.closProps.eqs|split|6','thm.semanticPrimitives.eq_result_case_def|split|1')).
fof(f21112,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X1),s('type.semanticPrimitives.eq_result',sK1)))) ) | $spl244),
  inference(forward_demodulation,[],[f21070,f131])).
fof(f21070,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X1),s('type.semanticPrimitives.eq_result',sK1))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X1),s('type.semanticPrimitives.eq_result',sK1)))) ) | $spl244),
  inference(resolution,[],[f21059,f1234])).
fof(f23352,plain,(
  $false | ($spl3 | $spl11 | $spl76 | $spl84 | $spl88 | $spl244)),
  inference(subsumption_resolution,[],[f23341,f142])).
fof(f23341,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl11 | $spl76 | $spl84 | $spl88 | $spl244)),
  inference(superposition,[],[f23326,f22786])).
fof(f23138,plain,(
  $false | ($spl84 | $spl216 | $spl268 | $spl312)),
  inference(subsumption_resolution,[],[f23124,f23116])).
fof(f23116,plain,(
  ( ! [X4] : (~p(s(bool,X4))) ) | ($spl84 | $spl216 | $spl268 | $spl312)),
  inference(subsumption_resolution,[],[f22310,f22845])).
fof(f23124,plain,(
  ( ! [X6] : (p(s(bool,X6))) ) | ($spl84 | $spl216 | $spl268 | $spl312)),
  inference(resolution,[],[f23116,f22283])).
fof(f23121,plain,(
  $false | ($spl30 | $spl84 | $spl216 | $spl268 | $spl312)),
  inference(resolution,[],[f23116,f1886])).
fof(f23122,plain,(
  $false | ($spl84 | $spl216 | $spl268 | $spl312 | $spl336)),
  inference(resolution,[],[f23116,f22927])).
fof(f23123,plain,(
  $false | ($spl84 | $spl216 | $spl268 | $spl312)),
  inference(resolution,[],[f23116,f22845])).
fof(f23115,plain,(
  $false | ($spl26 | $spl84 | $spl136 | $spl216 | $spl268)),
  inference(subsumption_resolution,[],[f12946,f22850])).
fof(f22926,plain,(
  ~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | $spl337),
  inference(cnf_transformation,[],[f22926_D])).
fof(f22926_D,plain,(
  ~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) <=> ~$spl337),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl337])])).
fof(f9076,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) ) | $spl94),
  inference(cnf_transformation,[],[f9076_D])).
fof(f9076_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f22931,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) ) | $spl327),
  inference(superposition,[],[f22798,f138])).
fof(f22930,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | ~p(s(bool,X2))) ) | $spl327),
  inference(superposition,[],[f22798,f138])).
fof(f6483,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | $spl78),
  inference(cnf_transformation,[],[f6483_D])).
fof(f6483_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f22929,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | p(s(bool,X1))) ) | $spl327),
  inference(superposition,[],[f22798,f135])).
fof(f22928,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) ) | $spl327),
  inference(superposition,[],[f22798,f135])).
fof(f174,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | $spl16),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f22920,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) ) | $spl326),
  inference(superposition,[],[f22799,f138])).
fof(f22919,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | ~p(s(bool,X2))) ) | $spl326),
  inference(superposition,[],[f22799,f138])).
fof(f173,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | $spl14),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f22918,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))) | p(s(bool,X1))) ) | $spl326),
  inference(superposition,[],[f22799,f135])).
fof(f22917,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) ) | $spl326),
  inference(superposition,[],[f22799,f135])).
fof(f22880,plain,(
  $false | ($spl3 | $spl98 | $spl110)),
  inference(trivial_inequality_removal,[],[f22878])).
fof(f22878,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl98 | $spl110)),
  inference(superposition,[],[f22810,f8084])).
fof(f22810,plain,(
  ( ! [X6] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6)))) ) | ($spl98 | $spl110)),
  inference(subsumption_resolution,[],[f10058,f9938])).
fof(f22815,plain,(
  $false | ($spl3 | $spl13 | $spl98 | $spl217)),
  inference(resolution,[],[f22790,f171])).
fof(f22790,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl98 | $spl217)),
  inference(subsumption_resolution,[],[f22789,f142])).
fof(f22789,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | p(s(bool,X0))) ) | ($spl3 | $spl98 | $spl217)),
  inference(forward_demodulation,[],[f22788,f8084])).
fof(f22788,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl3 | $spl98 | $spl217)),
  inference(subsumption_resolution,[],[f22787,f9938])).
fof(f22787,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl3 | $spl217)),
  inference(forward_demodulation,[],[f22369,f8084])).
fof(f22369,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | $spl217),
  inference(resolution,[],[f17587,f1234])).
fof(f22816,plain,(
  $false | ($spl3 | $spl98 | $spl217)),
  inference(resolution,[],[f22790,f17587])).
fof(f22799,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | $spl326),
  inference(cnf_transformation,[],[f22799_D])).
fof(f22799_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f22809,plain,(
  ( ! [X29] : (s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X29),s(fun(bool,sK0),sK4),s(sK0,sK3))) | s('type.semanticPrimitives.eq_result',X29) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X29))))) ) | $spl334),
  inference(cnf_transformation,[],[f22809_D])).
fof(f22809_D,plain,(
  ( ! [X29] : (s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X29),s(fun(bool,sK0),sK4),s(sK0,sK3))) | s('type.semanticPrimitives.eq_result',X29) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X29))))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f17022,plain,(
  ( ! [X29] : (s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X29),s(fun(bool,sK0),sK4),s(sK0,sK3))) | s('type.semanticPrimitives.eq_result',X29) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X29)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))))) ) | $spl107),
  inference(superposition,[],[f10010,f1229])).
fof(f1229,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result',X1) | s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) | s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X0))))) )),
  inference(superposition,[],[f132,f132])).
fof(f10010,plain,(
  s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK3))) | $spl107),
  inference(cnf_transformation,[],[f10010_D])).
fof(f10010_D,plain,(
  s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK3))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f22808,plain,(
  ( ! [X39,X37,X38,X36] : (s(X37,X38) = s(X37,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X36),s(fun(bool,X37),X39),s(X37,X38))) | s('type.semanticPrimitives.eq_result',X36) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X36))))) ) | $spl332),
  inference(cnf_transformation,[],[f22808_D])).
fof(f22808_D,plain,(
  ( ! [X39,X37,X38,X36] : (s(X37,X38) = s(X37,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X36),s(fun(bool,X37),X39),s(X37,X38))) | s('type.semanticPrimitives.eq_result',X36) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X36))))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f17024,plain,(
  ( ! [X39,X37,X38,X36] : (s(X37,X38) = s(X37,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X36),s(fun(bool,X37),X39),s(X37,X38))) | s('type.semanticPrimitives.eq_result',X36) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X36)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))))) )),
  inference(superposition,[],[f131,f1229])).
fof(f17047,plain,(
  ( ! [X29] : (s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X29),s(fun(bool,sK0),sK4),s(sK0,sK3))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result',X29) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X29))))) ) | $spl107),
  inference(superposition,[],[f10010,f1229])).
fof(f17049,plain,(
  ( ! [X39,X37,X38,X36] : (s(X37,X38) = s(X37,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X36),s(fun(bool,X37),X39),s(X37,X38))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result',X36) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X36))))) )),
  inference(superposition,[],[f131,f1229])).
fof(f22803,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.eq_result',X7) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X6),s(fun(bool,'type.semanticPrimitives.eq_result'),X8),s('type.semanticPrimitives.eq_result',X7)))))) | s('type.semanticPrimitives.eq_result',X6) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X6)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X7)) ) | $spl330),
  inference(cnf_transformation,[],[f22803_D])).
fof(f22803_D,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.eq_result',X7) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X6),s(fun(bool,'type.semanticPrimitives.eq_result'),X8),s('type.semanticPrimitives.eq_result',X7)))))) | s('type.semanticPrimitives.eq_result',X6) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X6)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X7)) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f19164,plain,(
  ( ! [X4,X5,X3] : (s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X3),s(fun(bool,'type.semanticPrimitives.eq_result'),X5),s('type.semanticPrimitives.eq_result',X4)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result',X3) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X3))))) )),
  inference(superposition,[],[f1261,f1229])).
fof(f1261,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',X4)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4)) )),
  inference(forward_demodulation,[],[f1232,f131])).
fof(f1232,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',X4))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',X4))))))) )),
  inference(superposition,[],[f132,f131])).
fof(f19165,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.eq_result',X7) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X6),s(fun(bool,'type.semanticPrimitives.eq_result'),X8),s('type.semanticPrimitives.eq_result',X7)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X7) | s('type.semanticPrimitives.eq_result',X6) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X6)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))))) )),
  inference(superposition,[],[f1261,f1229])).
fof(f22800,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X4),s(fun(bool,'type.semanticPrimitives.eq_result'),X5),s('type.semanticPrimitives.eq_result',sK1))))) | s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X4))))) ) | $spl328),
  inference(cnf_transformation,[],[f22800_D])).
fof(f22800_D,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X4),s(fun(bool,'type.semanticPrimitives.eq_result'),X5),s('type.semanticPrimitives.eq_result',sK1))))) | s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X4))))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f21087,plain,(
  ( ! [X2,X3] : (~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X2),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',sK1))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0')))) | s('type.semanticPrimitives.eq_result',X2) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X2))))) ) | $spl244),
  inference(superposition,[],[f21059,f1229])).
fof(f21088,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X4),s(fun(bool,'type.semanticPrimitives.eq_result'),X5),s('type.semanticPrimitives.eq_result',sK1))))) | s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X4)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_type_error_0'))))) ) | $spl244),
  inference(superposition,[],[f21059,f1229])).
fof(f11192,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | ~p(s(bool,X1))) ) | $spl120),
  inference(cnf_transformation,[],[f11192_D])).
fof(f11192_D,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | ~p(s(bool,X1))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f22792,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | $spl324),
  inference(cnf_transformation,[],[f22792_D])).
fof(f22792_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f22353,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | ~p(s(bool,X0))) ) | $spl76),
  inference(resolution,[],[f6374,f1236])).
fof(f22773,plain,(
  $false | ($spl3 | $spl11 | $spl76 | $spl78 | $spl84 | $spl94)),
  inference(subsumption_resolution,[],[f22772,f148])).
fof(f22772,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl3 | $spl76 | $spl78 | $spl84 | $spl94)),
  inference(forward_demodulation,[],[f22771,f8084])).
fof(f22771,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | ($spl76 | $spl78 | $spl84 | $spl94)),
  inference(subsumption_resolution,[],[f22744,f9367])).
fof(f9367,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) ) | ($spl78 | $spl94)),
  inference(subsumption_resolution,[],[f9076,f6483])).
fof(f22538,plain,(
  $false | ($spl13 | $spl84 | $spl260)),
  inference(subsumption_resolution,[],[f22464,f171])).
fof(f22464,plain,(
  p(s(bool,sK5)) | ($spl13 | $spl84 | $spl260)),
  inference(backward_demodulation,[],[f22453,f8216])).
fof(f22453,plain,(
  ( ! [X10] : (s(bool,sK5) = s(bool,X10)) ) | ($spl13 | $spl260)),
  inference(subsumption_resolution,[],[f22107,f22321])).
fof(f22537,plain,(
  $false | ($spl3 | $spl11 | $spl13 | $spl260)),
  inference(subsumption_resolution,[],[f22463,f148])).
fof(f22463,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl3 | $spl13 | $spl260)),
  inference(backward_demodulation,[],[f22453,f8084])).
fof(f22536,plain,(
  $false | ($spl13 | $spl76 | $spl260)),
  inference(subsumption_resolution,[],[f22462,f171])).
fof(f22462,plain,(
  p(s(bool,sK5)) | ($spl13 | $spl76 | $spl260)),
  inference(backward_demodulation,[],[f22453,f6374])).
fof(f22532,plain,(
  $false | ($spl13 | $spl26 | $spl76 | $spl260)),
  inference(trivial_inequality_removal,[],[f22520])).
fof(f22520,plain,(
  s(bool,sK5) != s(bool,sK5) | ($spl13 | $spl26 | $spl76 | $spl260)),
  inference(superposition,[],[f22414,f22453])).
fof(f22530,plain,(
  $false | ($spl13 | $spl93 | $spl260)),
  inference(trivial_inequality_removal,[],[f22522])).
fof(f22522,plain,(
  s(bool,sK5) != s(bool,sK5) | ($spl13 | $spl93 | $spl260)),
  inference(superposition,[],[f8535,f22453])).
fof(f10058,plain,(
  ( ! [X6] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X6))) ) | $spl110),
  inference(cnf_transformation,[],[f10058_D])).
fof(f10058_D,plain,(
  ( ! [X6] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X6))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f9975,plain,(
  ( ! [X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7))) | ~p(s(bool,X7))) ) | $spl102),
  inference(cnf_transformation,[],[f9975_D])).
fof(f9975_D,plain,(
  ( ! [X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7))) | ~p(s(bool,X7))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f22432,plain,(
  ( ! [X6,X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl6),
  inference(superposition,[],[f145,f138])).
fof(f22431,plain,(
  ( ! [X4,X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl6),
  inference(superposition,[],[f145,f138])).
fof(f9938,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | $spl98),
  inference(cnf_transformation,[],[f9938_D])).
fof(f9938_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f9939,plain,(
  ( ! [X1] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X1))) | p(s(bool,X1))) ) | $spl100),
  inference(cnf_transformation,[],[f9939_D])).
fof(f9939_D,plain,(
  ( ! [X1] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X1))) | p(s(bool,X1))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f22430,plain,(
  ( ! [X2,X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f145,f135])).
fof(f22429,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X1))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl6),
  inference(superposition,[],[f145,f135])).
fof(f171,plain,(
  ~p(s(bool,sK5)) | $spl13),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  ~p(s(bool,sK5)) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f10018,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK5))) ) | ($spl11 | $spl98)),
  inference(subsumption_resolution,[],[f10015,f9938])).
fof(f10020,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,sK5))) ) | ($spl11 | $spl98)),
  inference(subsumption_resolution,[],[f10016,f9938])).
fof(f22292,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X2))) ) | $spl284),
  inference(cnf_transformation,[],[f22292_D])).
fof(f22292_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X2))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f22319,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,sK6(sK1))) ) | $spl322),
  inference(cnf_transformation,[],[f22319_D])).
fof(f22319_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,sK6(sK1))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f19707,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,sK6(sK1)) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl184),
  inference(superposition,[],[f17550,f135])).
fof(f19708,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,sK6(sK1)) | s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl184),
  inference(superposition,[],[f17550,f135])).
fof(f22318,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X6))) ) | $spl320),
  inference(cnf_transformation,[],[f22318_D])).
fof(f22318_D,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X6))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f16844,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,sK6(sK1))) ) | $spl174),
  inference(cnf_transformation,[],[f16844_D])).
fof(f16844_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,sK6(sK1))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f19709,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,sK6(sK1)) | s(bool,X4) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl184),
  inference(superposition,[],[f17550,f138])).
fof(f19710,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6(sK1)) | s(bool,X6) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl184),
  inference(superposition,[],[f17550,f138])).
fof(f17590,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X7))) ) | $spl218),
  inference(cnf_transformation,[],[f17590_D])).
fof(f17590_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X7))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f19538,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X6) = s(bool,sK6(sK1)) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f19537,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X4) = s(bool,sK6(sK1)) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f10237,plain,(
  ( ! [X6,X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl6),
  inference(superposition,[],[f145,f138])).
fof(f10236,plain,(
  ( ! [X4,X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl6),
  inference(superposition,[],[f145,f138])).
fof(f10235,plain,(
  ( ! [X2,X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f145,f135])).
fof(f10234,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X1))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl6),
  inference(superposition,[],[f145,f135])).
fof(f10003,plain,(
  ( ! [X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK5))) ) | $spl9),
  inference(superposition,[],[f146,f138])).
fof(f146,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | $spl9),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f10002,plain,(
  ( ! [X2] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X2))) | ~p(s(bool,sK5)) | ~p(s(bool,X2))) ) | $spl9),
  inference(superposition,[],[f146,f138])).
fof(f172,plain,(
  p(s(bool,sK5)) | $spl12),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  p(s(bool,sK5)) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f10001,plain,(
  ( ! [X1] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X1))) | p(s(bool,sK5)) | p(s(bool,X1))) ) | $spl9),
  inference(superposition,[],[f146,f135])).
fof(f10000,plain,(
  ( ! [X0] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK5))) ) | $spl9),
  inference(superposition,[],[f146,f135])).
fof(f10482,plain,(
  ( ! [X4] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X4))) | ~p(s(bool,X4)) | ~p(s(bool,sK6(sK1)))) ) | $spl115),
  inference(superposition,[],[f10322,f138])).
fof(f10481,plain,(
  ( ! [X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X3))) ) | $spl115),
  inference(superposition,[],[f10322,f138])).
fof(f8216,plain,(
  p(s(bool,sK6(sK1))) | $spl84),
  inference(cnf_transformation,[],[f8216_D])).
fof(f8216_D,plain,(
  p(s(bool,sK6(sK1))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f10480,plain,(
  ( ! [X2] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X2))) | p(s(bool,sK6(sK1))) | p(s(bool,X2))) ) | $spl115),
  inference(superposition,[],[f10322,f135])).
fof(f10479,plain,(
  ( ! [X1] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X1))) | p(s(bool,X1)) | p(s(bool,sK6(sK1)))) ) | $spl115),
  inference(superposition,[],[f10322,f135])).
fof(f8217,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | $spl86),
  inference(cnf_transformation,[],[f8217_D])).
fof(f8217_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f14511,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,X1)) | p(s(bool,sK6(sK1)))) ) | $spl3),
  inference(superposition,[],[f8084,f135])).
fof(f14512,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | p(s(bool,sK6(sK1))) | p(s(bool,X2))) ) | $spl3),
  inference(superposition,[],[f8084,f135])).
fof(f22316,plain,(
  ( ! [X3] : (p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',sK1)))))) ) | $spl316),
  inference(cnf_transformation,[],[f22316_D])).
fof(f22316_D,plain,(
  ( ! [X3] : (p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',sK1)))))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f21014,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | p(s(bool,X2)) | p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X1),s('type.semanticPrimitives.eq_result',sK1)))))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f19231,f135])).
fof(f19231,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X16),s('type.semanticPrimitives.eq_result',sK1))))))) ) | ($spl3 | $spl78 | $spl94)),
  inference(subsumption_resolution,[],[f19177,f9367])).
fof(f19177,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X16),s('type.semanticPrimitives.eq_result',sK1)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | $spl3),
  inference(superposition,[],[f1261,f8084])).
fof(f21015,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',sK1))))) | p(s(bool,X4))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f19231,f135])).
fof(f21059,plain,(
  ( ! [X5] : (~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X5),s('type.semanticPrimitives.eq_result',sK1)))))) ) | $spl244),
  inference(cnf_transformation,[],[f21059_D])).
fof(f21059_D,plain,(
  ( ! [X5] : (~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X5),s('type.semanticPrimitives.eq_result',sK1)))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f8218,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | $spl88),
  inference(cnf_transformation,[],[f8218_D])).
fof(f8218_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f21016,plain,(
  ( ! [X6,X5] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X5),s('type.semanticPrimitives.eq_result',sK1))))) | ~p(s(bool,X6))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f19231,f138])).
fof(f21017,plain,(
  ( ! [X8,X7] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8))) | ~p(s(bool,X8)) | ~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X7),s('type.semanticPrimitives.eq_result',sK1)))))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f19231,f138])).
fof(f10014,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK5)) | p(s(bool,X1))) ) | $spl11),
  inference(superposition,[],[f148,f135])).
fof(f10013,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK5))) ) | $spl11),
  inference(superposition,[],[f148,f135])).
fof(f10068,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK6(sK1)))) ) | $spl3),
  inference(superposition,[],[f8084,f135])).
fof(f10069,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6(sK1))) | p(s(bool,X1))) ) | $spl3),
  inference(superposition,[],[f8084,f135])).
fof(f10070,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X2))) ) | $spl3),
  inference(superposition,[],[f8084,f138])).
fof(f10071,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK6(sK1)))) ) | $spl3),
  inference(superposition,[],[f8084,f138])).
fof(f22181,plain,(
  ( ! [X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7)))) ) | $spl266),
  inference(cnf_transformation,[],[f22181_D])).
fof(f22181_D,plain,(
  ( ! [X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7)))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f22313,plain,(
  ( ! [X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | ~p(s(bool,sK6(sK1)))) ) | ($spl100 | $spl115)),
  inference(subsumption_resolution,[],[f10481,f9939])).
fof(f22312,plain,(
  ( ! [X4] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X4))) | ~p(s(bool,sK6(sK1)))) ) | ($spl100 | $spl115)),
  inference(subsumption_resolution,[],[f10482,f9939])).
fof(f22309,plain,(
  ( ! [X7] : (s(bool,sK5) != s(bool,X7) | p(s(bool,X7))) ) | $spl310),
  inference(cnf_transformation,[],[f22309_D])).
fof(f22309_D,plain,(
  ( ! [X7] : (s(bool,sK5) != s(bool,X7) | p(s(bool,X7))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f14332,plain,(
  ( ! [X1] : (s(bool,sK5) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,sK6(sK1)))) ) | $spl93),
  inference(superposition,[],[f8535,f135])).
fof(f14333,plain,(
  ( ! [X2] : (s(bool,sK5) != s(bool,X2) | p(s(bool,sK6(sK1))) | p(s(bool,X2))) ) | $spl93),
  inference(superposition,[],[f8535,f135])).
fof(f1686,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,sK5) != s(bool,X4)) ) | $spl26),
  inference(cnf_transformation,[],[f1686_D])).
fof(f1686_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,sK5) != s(bool,X4)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f14334,plain,(
  ( ! [X3] : (s(bool,sK5) != s(bool,X3) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X3))) ) | $spl93),
  inference(superposition,[],[f8535,f138])).
fof(f14335,plain,(
  ( ! [X4] : (s(bool,sK5) != s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,sK6(sK1)))) ) | $spl93),
  inference(superposition,[],[f8535,f138])).
fof(f14513,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X3))) ) | $spl3),
  inference(superposition,[],[f8084,f138])).
fof(f14514,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | ~p(s(bool,X4)) | ~p(s(bool,sK6(sK1)))) ) | $spl3),
  inference(superposition,[],[f8084,f138])).
fof(f22311,plain,(
  ( ! [X2] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | p(s(bool,X2))) ) | $spl314),
  inference(cnf_transformation,[],[f22311_D])).
fof(f22311_D,plain,(
  ( ! [X2] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | p(s(bool,X2))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f18029,plain,(
  ( ! [X1] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1))))) | p(s(bool,X1)) | p(s(bool,sK6(sK1)))) ) | $spl217),
  inference(superposition,[],[f17587,f135])).
fof(f18030,plain,(
  ( ! [X2] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | p(s(bool,sK6(sK1))) | p(s(bool,X2))) ) | $spl217),
  inference(superposition,[],[f17587,f135])).
fof(f22310,plain,(
  ( ! [X4] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X4))))) | ~p(s(bool,X4))) ) | $spl312),
  inference(cnf_transformation,[],[f22310_D])).
fof(f22310_D,plain,(
  ( ! [X4] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X4))))) | ~p(s(bool,X4))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f18031,plain,(
  ( ! [X3] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3))))) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X3))) ) | $spl217),
  inference(superposition,[],[f17587,f138])).
fof(f18032,plain,(
  ( ! [X4] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X4))))) | ~p(s(bool,X4)) | ~p(s(bool,sK6(sK1)))) ) | $spl217),
  inference(superposition,[],[f17587,f138])).
fof(f17601,plain,(
  ( ! [X3] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) | p(s(bool,X3))) ) | $spl224),
  inference(cnf_transformation,[],[f17601_D])).
fof(f17601_D,plain,(
  ( ! [X3] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) | p(s(bool,X3))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f18315,plain,(
  ( ! [X1] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) | p(s(bool,X1)) | p(s(bool,sK6(sK1)))) ) | $spl137),
  inference(superposition,[],[f12947,f135])).
fof(f18316,plain,(
  ( ! [X2] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) | p(s(bool,sK6(sK1))) | p(s(bool,X2))) ) | $spl137),
  inference(superposition,[],[f12947,f135])).
fof(f22290,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5))))) ) | $spl280),
  inference(cnf_transformation,[],[f22290_D])).
fof(f22290_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f18317,plain,(
  ( ! [X3] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X3))) ) | $spl137),
  inference(superposition,[],[f12947,f138])).
fof(f18318,plain,(
  ( ! [X4] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X4)))) | ~p(s(bool,X4)) | ~p(s(bool,sK6(sK1)))) ) | $spl137),
  inference(superposition,[],[f12947,f138])).
fof(f17586,plain,(
  p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) | $spl216),
  inference(cnf_transformation,[],[f17586_D])).
fof(f17586_D,plain,(
  p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f18326,plain,(
  ( ! [X0] : (s(bool,sK5) != s(bool,X0) | p(s(bool,X0)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) ) | $spl137),
  inference(superposition,[],[f12947,f135])).
fof(f18327,plain,(
  ( ! [X1] : (s(bool,sK5) != s(bool,X1) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) | p(s(bool,X1))) ) | $spl137),
  inference(superposition,[],[f12947,f135])).
fof(f17587,plain,(
  ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) | $spl217),
  inference(cnf_transformation,[],[f17587_D])).
fof(f17587_D,plain,(
  ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f18328,plain,(
  ( ! [X2] : (s(bool,sK5) != s(bool,X2) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) | ~p(s(bool,X2))) ) | $spl137),
  inference(superposition,[],[f12947,f138])).
fof(f18329,plain,(
  ( ! [X3] : (s(bool,sK5) != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) ) | $spl137),
  inference(superposition,[],[f12947,f138])).
fof(f22294,plain,(
  ( ! [X0] : (s(bool,sK5) = s(bool,X0) | p(s(bool,X0))) ) | $spl288),
  inference(cnf_transformation,[],[f22294_D])).
fof(f22294_D,plain,(
  ( ! [X0] : (s(bool,sK5) = s(bool,X0) | p(s(bool,X0))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f18549,plain,(
  ( ! [X0,X1] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) | s(bool,sK5) = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl134),
  inference(superposition,[],[f12935,f135])).
fof(f18550,plain,(
  ( ! [X2,X3] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) | s(bool,sK5) = s(bool,X2) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl134),
  inference(superposition,[],[f12935,f135])).
fof(f22308,plain,(
  ( ! [X6] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X6))))) | s(bool,sK5) = s(bool,X6)) ) | $spl308),
  inference(cnf_transformation,[],[f22308_D])).
fof(f22308_D,plain,(
  ( ! [X6] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X6))))) | s(bool,sK5) = s(bool,X6)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f18579,plain,(
  ( ! [X4,X5] : (s(bool,sK5) != s(bool,X5) | s(bool,sK5) = s(bool,X4) | p(s(bool,X5)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X4)))))) ) | $spl134),
  inference(superposition,[],[f12935,f135])).
fof(f18580,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,X7) | s(bool,sK5) = s(bool,X6) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X6))))) | p(s(bool,X7))) ) | $spl134),
  inference(superposition,[],[f12935,f135])).
fof(f17597,plain,(
  ( ! [X9] : (s(bool,sK5) = s(bool,X9) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))))) ) | $spl220),
  inference(cnf_transformation,[],[f17597_D])).
fof(f17597_D,plain,(
  ( ! [X9] : (s(bool,sK5) = s(bool,X9) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f18581,plain,(
  ( ! [X8,X9] : (s(bool,sK5) != s(bool,X9) | s(bool,sK5) = s(bool,X8) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X8))))) | ~p(s(bool,X9))) ) | $spl134),
  inference(superposition,[],[f12935,f138])).
fof(f18582,plain,(
  ( ! [X10,X11] : (s(bool,sK5) != s(bool,X11) | s(bool,sK5) = s(bool,X10) | ~p(s(bool,X11)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X10)))))) ) | $spl134),
  inference(superposition,[],[f12935,f138])).
fof(f22306,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,sK6(sK1)) | p(s(bool,X2))) ) | $spl304),
  inference(cnf_transformation,[],[f22306_D])).
fof(f22306_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,sK6(sK1)) | p(s(bool,X2))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f22307,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X3))) ) | $spl306),
  inference(cnf_transformation,[],[f22307_D])).
fof(f22307_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X3))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f19535,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X0) = s(bool,sK6(sK1)) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl144),
  inference(superposition,[],[f14176,f135])).
fof(f19536,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X2) = s(bool,sK6(sK1)) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl144),
  inference(superposition,[],[f14176,f135])).
fof(f22305,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl144 | $spl174)),
  inference(subsumption_resolution,[],[f19537,f16844])).
fof(f22304,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl144 | $spl174)),
  inference(subsumption_resolution,[],[f19538,f16844])).
fof(f22303,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))) | p(s(bool,X5)) | s(bool,X6) = s(bool,sK6(sK1))) ) | $spl302),
  inference(cnf_transformation,[],[f22303_D])).
fof(f22303_D,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))) | p(s(bool,X5)) | s(bool,X6) = s(bool,sK6(sK1))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f19561,plain,(
  ( ! [X4,X3] : (s(bool,X4) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) | s(bool,X4) = s(bool,sK6(sK1)) | p(s(bool,X3)) | p(s(bool,sK6(sK1)))) ) | $spl144),
  inference(superposition,[],[f14176,f135])).
fof(f19562,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))) | s(bool,X6) = s(bool,sK6(sK1)) | p(s(bool,sK6(sK1))) | p(s(bool,X5))) ) | $spl144),
  inference(superposition,[],[f14176,f135])).
fof(f22302,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))) | ~p(s(bool,X9)) | s(bool,X10) = s(bool,sK6(sK1))) ) | $spl300),
  inference(cnf_transformation,[],[f22302_D])).
fof(f22302_D,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))) | ~p(s(bool,X9)) | s(bool,X10) = s(bool,sK6(sK1))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f19563,plain,(
  ( ! [X8,X7] : (s(bool,X8) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X7)))) | s(bool,X8) = s(bool,sK6(sK1)) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X7))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f19564,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))) | s(bool,X10) = s(bool,sK6(sK1)) | ~p(s(bool,X9)) | ~p(s(bool,sK6(sK1)))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f22300,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,X13) = s(bool,sK6(sK1))) ) | $spl298),
  inference(cnf_transformation,[],[f22300_D])).
fof(f22300_D,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,X13) = s(bool,sK6(sK1))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f22301,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,X11) = s(bool,sK6(sK1)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) ) | ($spl142 | $spl144)),
  inference(subsumption_resolution,[],[f19577,f14173])).
fof(f19577,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,X11) = s(bool,sK6(sK1)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) | ~p(s(bool,X10))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f22299,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,X13) = s(bool,sK6(sK1)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) ) | ($spl142 | $spl144)),
  inference(subsumption_resolution,[],[f19578,f14173])).
fof(f19578,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,X13) = s(bool,sK6(sK1)) | ~p(s(bool,X12)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f22298,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X5)) | s(bool,X6) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) | $spl296),
  inference(cnf_transformation,[],[f22298_D])).
fof(f22298_D,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X5)) | s(bool,X6) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f19736,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X3)) | p(s(bool,sK6(sK1)))) ) | $spl184),
  inference(superposition,[],[f17550,f135])).
fof(f19737,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | s(bool,X6) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,sK6(sK1))) | p(s(bool,X5))) ) | $spl184),
  inference(superposition,[],[f17550,f135])).
fof(f22297,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9)) | s(bool,X10) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) | $spl294),
  inference(cnf_transformation,[],[f22297_D])).
fof(f22297_D,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9)) | s(bool,X10) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f19738,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,X8) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X7))) ) | $spl184),
  inference(superposition,[],[f17550,f138])).
fof(f19739,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X10) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ~p(s(bool,X9)) | ~p(s(bool,sK6(sK1)))) ) | $spl184),
  inference(superposition,[],[f17550,f138])).
fof(f6374,plain,(
  p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | $spl76),
  inference(cnf_transformation,[],[f6374_D])).
fof(f6374_D,plain,(
  p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f22296,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X5)) | s(bool,sK5) = s(bool,X6)) ) | $spl292),
  inference(cnf_transformation,[],[f22296_D])).
fof(f22296_D,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X5)) | s(bool,sK5) = s(bool,X6)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f15641,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,sK5) = s(bool,X8) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | p(s(bool,X7))) ) | $spl168),
  inference(superposition,[],[f15597,f135])).
fof(f15640,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | s(bool,sK5) = s(bool,X6) | p(s(bool,X5)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | $spl168),
  inference(superposition,[],[f15597,f135])).
fof(f22295,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | p(s(bool,X1))) ) | $spl290),
  inference(cnf_transformation,[],[f22295_D])).
fof(f22295_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | p(s(bool,X1))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f15606,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X2) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl168),
  inference(superposition,[],[f15597,f135])).
fof(f15605,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl168),
  inference(superposition,[],[f15597,f135])).
fof(f22293,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,sK6(sK1))) ) | $spl286),
  inference(cnf_transformation,[],[f22293_D])).
fof(f22293_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,sK6(sK1))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f19749,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,sK6(sK1)) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X0))) ) | ($spl174 | $spl184)),
  inference(subsumption_resolution,[],[f19707,f16844])).
fof(f19751,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,sK6(sK1)) | s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | p(s(bool,X2))) ) | ($spl174 | $spl184)),
  inference(subsumption_resolution,[],[f19708,f16844])).
fof(f22291,plain,(
  ( ! [X31] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X31),s('type.semanticPrimitives.eq_result',sK1)))))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X31),s('type.semanticPrimitives.eq_result',sK1)))))))))) ) | $spl282),
  inference(cnf_transformation,[],[f22291_D])).
fof(f22291_D,plain,(
  ( ! [X31] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X31),s('type.semanticPrimitives.eq_result',sK1)))))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X31),s('type.semanticPrimitives.eq_result',sK1)))))))))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f17574,plain,(
  ( ! [X24] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X24) | s('type.semanticPrimitives.eq_result',X24) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X24))))) ) | $spl208),
  inference(cnf_transformation,[],[f17574_D])).
fof(f17574_D,plain,(
  ( ! [X24] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X24) | s('type.semanticPrimitives.eq_result',X24) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X24))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f21030,plain,(
  ( ! [X10,X11] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X11) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X10),s('type.semanticPrimitives.eq_result',sK1)))))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X10),s('type.semanticPrimitives.eq_result',sK1))))))))) | s('type.semanticPrimitives.eq_result',X11) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X11))))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f19231,f1229])).
fof(f21031,plain,(
  ( ! [X12,X13] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X13) | s('type.semanticPrimitives.eq_result',X13) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X13)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X12),s('type.semanticPrimitives.eq_result',sK1)))))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X12),s('type.semanticPrimitives.eq_result',sK1)))))))))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f19231,f1229])).
fof(f21045,plain,(
  ( ! [X28,X27] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X28) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X27),s('type.semanticPrimitives.eq_result',sK1)))))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X27),s('type.semanticPrimitives.eq_result',sK1))))))))) | s('type.semanticPrimitives.eq_result',X28) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X28))))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f1229,f19231])).
fof(f21046,plain,(
  ( ! [X30,X29] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X30) | s('type.semanticPrimitives.eq_result',X30) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X30)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X29),s('type.semanticPrimitives.eq_result',sK1)))))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X29),s('type.semanticPrimitives.eq_result',sK1)))))))))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f1229,f19231])).
fof(f17566,plain,(
  ( ! [X41,X42,X40] : (s(X40,X41) = s(X40,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X40),X42),s(X40,X41)))) ) | $spl198),
  inference(cnf_transformation,[],[f17566_D])).
fof(f17566_D,plain,(
  ( ! [X41,X42,X40] : (s(X40,X41) = s(X40,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X40),X42),s(X40,X41)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f21047,plain,(
  ( ! [X33,X31,X34,X32] : (s(X32,X33) = s(X32,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X32),X34),s(X32,X33))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X31),s('type.semanticPrimitives.eq_result',sK1)))))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X31),s('type.semanticPrimitives.eq_result',sK1)))))))))) ) | ($spl3 | $spl78 | $spl94)),
  inference(superposition,[],[f1244,f19231])).
fof(f1244,plain,(
  ( ! [X4,X2,X3,X1] : (s(X2,X3) = s(X2,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X2),X4),s(X2,X3))) | s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1))))) )),
  inference(superposition,[],[f131,f132])).
fof(f22107,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,sK5) = s(bool,X10)) ) | $spl260),
  inference(cnf_transformation,[],[f22107_D])).
fof(f22107_D,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,sK5) = s(bool,X10)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f18552,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X7)))) | s(bool,sK5) = s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl134),
  inference(superposition,[],[f12935,f138])).
fof(f18551,plain,(
  ( ! [X4,X5] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))) | s(bool,sK5) = s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl134),
  inference(superposition,[],[f12935,f138])).
fof(f22289,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | $spl278),
  inference(cnf_transformation,[],[f22289_D])).
fof(f22289_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f17572,plain,(
  ( ! [X28] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',X28) | s('type.semanticPrimitives.eq_result',X28) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X28))))) ) | $spl204),
  inference(cnf_transformation,[],[f17572_D])).
fof(f17572_D,plain,(
  ( ! [X28] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',X28) | s('type.semanticPrimitives.eq_result',X28) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X28))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f17037,plain,(
  ( ! [X9] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',X9) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | s('type.semanticPrimitives.eq_result',X9) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X9))))) ) | $spl11),
  inference(superposition,[],[f148,f1229])).
fof(f17012,plain,(
  ( ! [X9] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',X9) | s('type.semanticPrimitives.eq_result',X9) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X9)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))))) ) | $spl11),
  inference(superposition,[],[f148,f1229])).
fof(f6373,plain,(
  ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | $spl77),
  inference(cnf_transformation,[],[f6373_D])).
fof(f6373_D,plain,(
  ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f22117,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | ~p(s(bool,X11)) | s(bool,sK5) = s(bool,X12)) ) | $spl262),
  inference(cnf_transformation,[],[f22117_D])).
fof(f22117_D,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | ~p(s(bool,X11)) | s(bool,sK5) = s(bool,X12)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f15643,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | s(bool,sK5) = s(bool,X12) | ~p(s(bool,X11)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | $spl168),
  inference(superposition,[],[f15597,f138])).
fof(f15642,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,sK5) = s(bool,X10) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | ~p(s(bool,X9))) ) | $spl168),
  inference(superposition,[],[f15597,f138])).
fof(f22287,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X10)))) | ~p(s(bool,X10)) | s(bool,sK5) = s(bool,X11)) ) | $spl276),
  inference(cnf_transformation,[],[f22287_D])).
fof(f22287_D,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X10)))) | ~p(s(bool,X10)) | s(bool,sK5) = s(bool,X11)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f15631,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X12)))) | s(bool,sK5) = s(bool,X13) | ~p(s(bool,X12)) | ~p(s(bool,sK5))) ) | $spl168),
  inference(superposition,[],[f15597,f138])).
fof(f15630,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X10)))) | s(bool,sK5) = s(bool,X11) | ~p(s(bool,sK5)) | ~p(s(bool,X10))) ) | $spl168),
  inference(superposition,[],[f15597,f138])).
fof(f22286,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X6)))) | p(s(bool,X6)) | s(bool,sK5) = s(bool,X7)) ) | $spl274),
  inference(cnf_transformation,[],[f22286_D])).
fof(f22286_D,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X6)))) | p(s(bool,X6)) | s(bool,sK5) = s(bool,X7)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f15629,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X8)))) | s(bool,sK5) = s(bool,X9) | p(s(bool,sK5)) | p(s(bool,X8))) ) | $spl168),
  inference(superposition,[],[f15597,f135])).
fof(f15628,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X6)))) | s(bool,sK5) = s(bool,X7) | p(s(bool,X6)) | p(s(bool,sK5))) ) | $spl168),
  inference(superposition,[],[f15597,f135])).
fof(f22285,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ~p(s(bool,X5))) ) | $spl272),
  inference(cnf_transformation,[],[f22285_D])).
fof(f22285_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ~p(s(bool,X5))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f15608,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl168),
  inference(superposition,[],[f15597,f138])).
fof(f15607,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl168),
  inference(superposition,[],[f15597,f138])).
fof(f22284,plain,(
  ( ! [X2] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | ~p(s(bool,X2))) ) | $spl270),
  inference(cnf_transformation,[],[f22284_D])).
fof(f22284_D,plain,(
  ( ! [X2] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | ~p(s(bool,X2))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f9087,plain,(
  ( ! [X3] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3))))) | ~p(s(bool,X3)) | ~p(s(bool,sK5))) ) | $spl76),
  inference(superposition,[],[f6374,f138])).
fof(f9086,plain,(
  ( ! [X2] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | ~p(s(bool,sK5)) | ~p(s(bool,X2))) ) | $spl76),
  inference(superposition,[],[f6374,f138])).
fof(f22283,plain,(
  ( ! [X0] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) | p(s(bool,X0))) ) | $spl268),
  inference(cnf_transformation,[],[f22283_D])).
fof(f22283_D,plain,(
  ( ! [X0] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) | p(s(bool,X0))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f9085,plain,(
  ( ! [X1] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1))))) | p(s(bool,sK5)) | p(s(bool,X1))) ) | $spl76),
  inference(superposition,[],[f6374,f135])).
fof(f9084,plain,(
  ( ! [X0] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) | p(s(bool,X0)) | p(s(bool,sK5))) ) | $spl76),
  inference(superposition,[],[f6374,f135])).
fof(f22222,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,sK5))) ) | ($spl76 | $spl80)),
  inference(duplicate_literal_removal,[],[f22221])).
fof(f22221,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,sK5))) ) | ($spl76 | $spl80)),
  inference(forward_demodulation,[],[f9084,f7479])).
fof(f22220,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,sK5))) ) | ($spl76 | $spl80)),
  inference(duplicate_literal_removal,[],[f22219])).
fof(f22219,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,sK5)) | p(s(bool,X1))) ) | ($spl76 | $spl80)),
  inference(forward_demodulation,[],[f9085,f7479])).
fof(f22188,plain,(
  ( ! [X2] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X2))) | ~p(s(bool,sK5))) ) | ($spl9 | $spl100)),
  inference(subsumption_resolution,[],[f10002,f9939])).
fof(f22187,plain,(
  ( ! [X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | ~p(s(bool,sK5))) ) | ($spl9 | $spl100)),
  inference(subsumption_resolution,[],[f10003,f9939])).
fof(f22186,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,sK5))) ) | ($spl11 | $spl86)),
  inference(subsumption_resolution,[],[f10013,f8217])).
fof(f22185,plain,(
  ( ! [X1] : (p(s(bool,sK5)) | p(s(bool,X1))) ) | ($spl11 | $spl86)),
  inference(subsumption_resolution,[],[f10014,f8217])).
fof(f10016,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK5))) ) | $spl11),
  inference(superposition,[],[f148,f138])).
fof(f10015,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK5)) | ~p(s(bool,X2))) ) | $spl11),
  inference(superposition,[],[f148,f138])).
fof(f22182,plain,(
  ( ! [X4,X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | ~p(s(bool,X4))) ) | ($spl6 | $spl100)),
  inference(subsumption_resolution,[],[f10236,f9939])).
fof(f22180,plain,(
  ( ! [X6,X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X6))) ) | ($spl6 | $spl100)),
  inference(subsumption_resolution,[],[f10237,f9939])).
fof(f17598,plain,(
  ( ! [X10] : (s(bool,sK5) != s(bool,X10)) ) | $spl222),
  inference(cnf_transformation,[],[f17598_D])).
fof(f17598_D,plain,(
  ( ! [X10] : (s(bool,sK5) != s(bool,X10)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f22136,plain,(
  ( ! [X4,X5] : (s(bool,sK5) != s(bool,X5) | s(bool,sK5) = s(bool,X4) | ~p(s(bool,X4))) ) | ($spl80 | $spl85 | $spl93 | $spl168)),
  inference(subsumption_resolution,[],[f22135,f14343])).
fof(f14343,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK5) != s(bool,X1)) ) | ($spl85 | $spl93)),
  inference(subsumption_resolution,[],[f14332,f8215])).
fof(f22135,plain,(
  ( ! [X4,X5] : (s(bool,sK5) != s(bool,X5) | s(bool,sK5) = s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl80 | $spl168)),
  inference(forward_demodulation,[],[f15607,f7479])).
fof(f22134,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,X7) | s(bool,sK5) = s(bool,X6) | ~p(s(bool,X6))) ) | ($spl80 | $spl85 | $spl93 | $spl168)),
  inference(subsumption_resolution,[],[f22133,f14343])).
fof(f22133,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,X7) | s(bool,sK5) = s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl80 | $spl168)),
  inference(forward_demodulation,[],[f15608,f7479])).
fof(f22121,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | p(s(bool,X7))) ) | $spl264),
  inference(cnf_transformation,[],[f22121_D])).
fof(f22121_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | p(s(bool,X7))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f22129,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X6)) | p(s(bool,sK5))) ) | ($spl80 | $spl85 | $spl93 | $spl168 | $spl174)),
  inference(subsumption_resolution,[],[f22128,f18473])).
fof(f18473,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,sK5) != s(bool,X4) | s(bool,X3) != s(bool,X4)) ) | ($spl85 | $spl93 | $spl174)),
  inference(subsumption_resolution,[],[f18452,f8215])).
fof(f18452,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | s(bool,sK5) != s(bool,X4) | p(s(bool,X3)) | p(s(bool,sK6(sK1)))) ) | ($spl85 | $spl93 | $spl174)),
  inference(superposition,[],[f18359,f135])).
fof(f18359,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,sK6(sK1)) | s(bool,sK5) != s(bool,X12)) ) | ($spl85 | $spl93 | $spl174)),
  inference(resolution,[],[f16844,f14343])).
fof(f22128,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,sK5) = s(bool,X7) | p(s(bool,X6)) | p(s(bool,sK5))) ) | ($spl80 | $spl168)),
  inference(forward_demodulation,[],[f15628,f7479])).
fof(f22127,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,sK5)) | p(s(bool,X8))) ) | ($spl80 | $spl85 | $spl93 | $spl168 | $spl174)),
  inference(subsumption_resolution,[],[f22126,f18473])).
fof(f22126,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,sK5) = s(bool,X9) | p(s(bool,sK5)) | p(s(bool,X8))) ) | ($spl80 | $spl168)),
  inference(forward_demodulation,[],[f15629,f7479])).
fof(f22125,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,sK5) = s(bool,X11) | ~p(s(bool,sK5)) | ~p(s(bool,X10))) ) | ($spl80 | $spl168)),
  inference(forward_demodulation,[],[f15630,f7479])).
fof(f22124,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,sK5) = s(bool,X13) | ~p(s(bool,X12)) | ~p(s(bool,sK5))) ) | ($spl80 | $spl168)),
  inference(forward_demodulation,[],[f15631,f7479])).
fof(f22123,plain,(
  ( ! [X6,X5] : (p(s(bool,sK5)) | s(bool,X5) != s(bool,X6) | p(s(bool,X5))) ) | ($spl80 | $spl85 | $spl93 | $spl168 | $spl174)),
  inference(forward_demodulation,[],[f22122,f7479])).
fof(f22122,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X5)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | ($spl85 | $spl93 | $spl168 | $spl174)),
  inference(subsumption_resolution,[],[f15640,f18473])).
fof(f22120,plain,(
  ( ! [X8,X7] : (p(s(bool,sK5)) | s(bool,X7) != s(bool,X8) | p(s(bool,X7))) ) | ($spl80 | $spl85 | $spl93 | $spl168 | $spl174)),
  inference(forward_demodulation,[],[f22119,f7479])).
fof(f22119,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | p(s(bool,X7))) ) | ($spl85 | $spl93 | $spl168 | $spl174)),
  inference(subsumption_resolution,[],[f15641,f18473])).
fof(f22118,plain,(
  ( ! [X10,X9] : (~p(s(bool,sK5)) | s(bool,X9) != s(bool,X10) | s(bool,sK5) = s(bool,X10) | ~p(s(bool,X9))) ) | ($spl80 | $spl168)),
  inference(forward_demodulation,[],[f15642,f7479])).
fof(f22116,plain,(
  ( ! [X12,X11] : (~p(s(bool,sK5)) | s(bool,X11) != s(bool,X12) | s(bool,sK5) = s(bool,X12) | ~p(s(bool,X11))) ) | ($spl80 | $spl168)),
  inference(forward_demodulation,[],[f15643,f7479])).
fof(f15828,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,sK5) != s(bool,X10) | s(bool,sK5) = s(bool,X9)) ) | ($spl80 | $spl85 | $spl93 | $spl134)),
  inference(forward_demodulation,[],[f15827,f7479])).
fof(f15826,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | s(bool,sK5) != s(bool,X8) | s(bool,sK5) = s(bool,X7)) ) | ($spl80 | $spl85 | $spl93 | $spl134)),
  inference(forward_demodulation,[],[f15825,f7479])).
fof(f22111,plain,(
  ( ! [X4,X5] : (s(bool,sK5) != s(bool,X5) | s(bool,sK5) = s(bool,X4) | ~p(s(bool,X4))) ) | ($spl80 | $spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f22110,f14343])).
fof(f22110,plain,(
  ( ! [X4,X5] : (s(bool,sK5) != s(bool,X5) | s(bool,sK5) = s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl80 | $spl134)),
  inference(forward_demodulation,[],[f18551,f7479])).
fof(f22109,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,X7) | s(bool,sK5) = s(bool,X6) | ~p(s(bool,X6))) ) | ($spl80 | $spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f22108,f14343])).
fof(f22108,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,X7) | s(bool,sK5) = s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl80 | $spl134)),
  inference(forward_demodulation,[],[f18552,f7479])).
fof(f18606,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,sK5) != s(bool,X9) | s(bool,sK5) = s(bool,X8)) ) | ($spl80 | $spl85 | $spl93 | $spl134)),
  inference(forward_demodulation,[],[f18605,f7479])).
fof(f18605,plain,(
  ( ! [X8,X9] : (s(bool,sK5) != s(bool,X9) | s(bool,sK5) = s(bool,X8) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X8)))))) ) | ($spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f18581,f14343])).
fof(f18608,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s(bool,sK5) != s(bool,X11) | s(bool,sK5) = s(bool,X10)) ) | ($spl80 | $spl85 | $spl93 | $spl134)),
  inference(forward_demodulation,[],[f18607,f7479])).
fof(f18607,plain,(
  ( ! [X10,X11] : (s(bool,sK5) != s(bool,X11) | s(bool,sK5) = s(bool,X10) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X10)))))) ) | ($spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f18582,f14343])).
fof(f22098,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | p(s(bool,X11)) | s(bool,X12) = s(bool,X13) | s(bool,sK5) != s(bool,X13)) ) | $spl258),
  inference(cnf_transformation,[],[f22098_D])).
fof(f22098_D,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | p(s(bool,X11)) | s(bool,X12) = s(bool,X13) | s(bool,sK5) != s(bool,X13)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f20150,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,sK5) != s(bool,X10) | s(bool,X9) = s(bool,X10) | p(s(bool,X8)) | p(s(bool,sK5))) ) | ($spl85 | $spl93)),
  inference(superposition,[],[f15013,f135])).
fof(f15013,plain,(
  ( ! [X26,X25] : (s(bool,sK5) != s(bool,X26) | s(bool,sK5) != s(bool,X25) | s(bool,X25) = s(bool,X26)) ) | ($spl85 | $spl93)),
  inference(resolution,[],[f14531,f14343])).
fof(f14531,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,sK5) != s(bool,X0)) ) | ($spl85 | $spl93)),
  inference(resolution,[],[f14343,f138])).
fof(f20151,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,sK5) != s(bool,X13) | s(bool,X12) = s(bool,X13) | p(s(bool,sK5)) | p(s(bool,X11))) ) | ($spl85 | $spl93)),
  inference(superposition,[],[f15013,f135])).
fof(f22097,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | ~p(s(bool,X17)) | s(bool,X18) = s(bool,X19) | s(bool,sK5) != s(bool,X19)) ) | $spl256),
  inference(cnf_transformation,[],[f22097_D])).
fof(f22097_D,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | ~p(s(bool,X17)) | s(bool,X18) = s(bool,X19) | s(bool,sK5) != s(bool,X19)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f20152,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) != s(bool,X15) | s(bool,sK5) != s(bool,X16) | s(bool,X15) = s(bool,X16) | ~p(s(bool,sK5)) | ~p(s(bool,X14))) ) | ($spl85 | $spl93)),
  inference(superposition,[],[f15013,f138])).
fof(f20153,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,sK5) != s(bool,X19) | s(bool,X18) = s(bool,X19) | ~p(s(bool,X17)) | ~p(s(bool,sK5))) ) | ($spl85 | $spl93)),
  inference(superposition,[],[f15013,f138])).
fof(f22092,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | p(s(bool,X11)) | s(bool,X12) = s(bool,X13) | s(bool,X13) = s(bool,sK6(sK1))) ) | $spl254),
  inference(cnf_transformation,[],[f22092_D])).
fof(f22092_D,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | p(s(bool,X11)) | s(bool,X12) = s(bool,X13) | s(bool,X13) = s(bool,sK6(sK1))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f21324,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,X10) = s(bool,sK6(sK1)) | s(bool,X9) = s(bool,X10) | p(s(bool,X8)) | p(s(bool,sK5))) ) | ($spl85 | $spl93)),
  inference(superposition,[],[f15839,f135])).
fof(f15839,plain,(
  ( ! [X17,X16] : (s(bool,sK5) != s(bool,X17) | s(bool,X16) = s(bool,sK6(sK1)) | s(bool,X16) = s(bool,X17)) ) | ($spl85 | $spl93)),
  inference(resolution,[],[f10111,f14343])).
fof(f10111,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,sK6(sK1))) ) | $spl85),
  inference(resolution,[],[f10059,f138])).
fof(f10059,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,sK6(sK1))) ) | $spl85),
  inference(resolution,[],[f8215,f135])).
fof(f21325,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,X13) = s(bool,sK6(sK1)) | s(bool,X12) = s(bool,X13) | p(s(bool,sK5)) | p(s(bool,X11))) ) | ($spl85 | $spl93)),
  inference(superposition,[],[f15839,f135])).
fof(f22091,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | ~p(s(bool,X17)) | s(bool,X18) = s(bool,X19) | s(bool,X19) = s(bool,sK6(sK1))) ) | $spl252),
  inference(cnf_transformation,[],[f22091_D])).
fof(f22091_D,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | ~p(s(bool,X17)) | s(bool,X18) = s(bool,X19) | s(bool,X19) = s(bool,sK6(sK1))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f21326,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X16) = s(bool,sK6(sK1)) | s(bool,X15) = s(bool,X16) | ~p(s(bool,sK5)) | ~p(s(bool,X14))) ) | ($spl85 | $spl93)),
  inference(superposition,[],[f15839,f138])).
fof(f21327,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,X19) = s(bool,sK6(sK1)) | s(bool,X18) = s(bool,X19) | ~p(s(bool,X17)) | ~p(s(bool,sK5))) ) | ($spl85 | $spl93)),
  inference(superposition,[],[f15839,f138])).
fof(f22086,plain,(
  $false | ($spl0 | $spl3 | $spl9 | $spl12 | $spl100)),
  inference(subsumption_resolution,[],[f22027,f22051])).
fof(f22051,plain,(
  ( ! [X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3)))) ) | ($spl9 | $spl12 | $spl100)),
  inference(subsumption_resolution,[],[f10005,f9939])).
fof(f10005,plain,(
  ( ! [X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | ~p(s(bool,X3))) ) | ($spl9 | $spl12)),
  inference(subsumption_resolution,[],[f10003,f172])).
fof(f21914,plain,(
  ( ! [X28,X26,X27] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X26),s(fun(bool,bool),X27),s(bool,X28))))) ) | $spl246),
  inference(cnf_transformation,[],[f21914_D])).
fof(f21914_D,plain,(
  ( ! [X28,X26,X27] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X26),s(fun(bool,bool),X27),s(bool,X28))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f21934,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,sK5) = s(bool,X56) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X53),s(fun(bool,bool),X54),s(bool,X55)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X53)) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f21933,f2548])).
fof(f2548,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK5) = s(bool,X1)) ) | $spl12),
  inference(resolution,[],[f172,f138])).
fof(f21933,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,sK5) = s(bool,X56) | p(s(bool,X56)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X53),s(fun(bool,bool),X54),s(bool,X55)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X53)) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f21795,f4419])).
fof(f4419,plain,(
  ( ! [X0,X1] : (s(bool,sK5) = s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) ) | ($spl12 | $spl30)),
  inference(resolution,[],[f1886,f2548])).
fof(f21795,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X56) = s(bool,happ(s(fun(bool,bool),X54),s(bool,sK6(X53)))) | p(s(bool,X56)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X53),s(fun(bool,bool),X54),s(bool,X55)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X53)) )),
  inference(superposition,[],[f135,f1302])).
fof(f21936,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,sK5) = s(bool,X60) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X57),s(fun(bool,bool),X58),s(bool,X59)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X57)) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f21935,f2548])).
fof(f21935,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,sK5) = s(bool,X60) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X57),s(fun(bool,bool),X58),s(bool,X59)))) | p(s(bool,X60)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X57)) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f21796,f4419])).
fof(f21796,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X60) = s(bool,happ(s(fun(bool,bool),X58),s(bool,sK6(X57)))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X57),s(fun(bool,bool),X58),s(bool,X59)))) | p(s(bool,X60)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X57)) )),
  inference(superposition,[],[f135,f1302])).
fof(f21921,plain,(
  ( ! [X52,X50,X51] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X50) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X50),s(fun(bool,bool),X51),s(bool,X52)))) ) | $spl248),
  inference(cnf_transformation,[],[f21921_D])).
fof(f21921_D,plain,(
  ( ! [X52,X50,X51] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X50) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X50),s(fun(bool,bool),X51),s(bool,X52)))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f21950,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,sK5) = s(bool,X121) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X118),s(fun(bool,bool),X119),s(bool,X120))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X118)) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(forward_demodulation,[],[f21949,f4419])).
fof(f21949,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,X121) = s(bool,happ(s(fun(bool,bool),X119),s(bool,sK6(X118)))) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X118),s(fun(bool,bool),X119),s(bool,X120))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X118)) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f21811,f15049])).
fof(f15049,plain,(
  ( ! [X37,X35,X36] : (s(bool,sK5) != s(bool,X37) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X36)))) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f15029,f172])).
fof(f15029,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,sK5)) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X36))) | s(bool,sK5) != s(bool,X37)) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(superposition,[],[f14531,f4419])).
fof(f21811,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,X121) = s(bool,happ(s(fun(bool,bool),X119),s(bool,sK6(X118)))) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X118),s(fun(bool,bool),X119),s(bool,X120))) | s(bool,sK5) = s(bool,X121) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X118)) ) | $spl12),
  inference(superposition,[],[f7351,f1302])).
fof(f7351,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,X7) | s(bool,sK5) = s(bool,X6) | s(bool,sK5) = s(bool,X7)) ) | $spl12),
  inference(resolution,[],[f4406,f2548])).
fof(f4406,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,sK5) = s(bool,X0)) ) | $spl12),
  inference(resolution,[],[f2548,f135])).
fof(f21952,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,sK5) = s(bool,X125) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X122),s(fun(bool,bool),X123),s(bool,X124))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X122)) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(forward_demodulation,[],[f21951,f4419])).
fof(f21951,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,X125) = s(bool,happ(s(fun(bool,bool),X123),s(bool,sK6(X122)))) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X122),s(fun(bool,bool),X123),s(bool,X124))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X122)) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f21812,f15049])).
fof(f21812,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,X125) = s(bool,happ(s(fun(bool,bool),X123),s(bool,sK6(X122)))) | s(bool,sK5) = s(bool,X125) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X122),s(fun(bool,bool),X123),s(bool,X124))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X122)) ) | $spl12),
  inference(superposition,[],[f7351,f1302])).
fof(f10329,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6(sK1))) ) | $spl116),
  inference(cnf_transformation,[],[f10329_D])).
fof(f10329_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6(sK1))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f21991,plain,(
  ( ! [X229,X227,X228,X226] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X226),s(fun(bool,bool),X227),s(bool,X228))) | s(bool,X229) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X226)) ) | ($spl12 | $spl30 | $spl76 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f21990,f10061])).
fof(f10061,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X2)) ) | ($spl12 | $spl85)),
  inference(resolution,[],[f8215,f4406])).
fof(f21990,plain,(
  ( ! [X229,X227,X228,X226] : (s(bool,sK5) != s(bool,X229) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X226),s(fun(bool,bool),X227),s(bool,X228))) | s(bool,X229) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X226)) ) | ($spl12 | $spl30 | $spl76 | $spl85 | $spl93)),
  inference(forward_demodulation,[],[f21843,f4419])).
fof(f21843,plain,(
  ( ! [X229,X227,X228,X226] : (s(bool,X229) != s(bool,happ(s(fun(bool,bool),X227),s(bool,sK6(X226)))) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X226),s(fun(bool,bool),X227),s(bool,X228))) | s(bool,X229) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X226)) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(superposition,[],[f15561,f1302])).
fof(f15561,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | s(bool,sK5) = s(bool,X15) | s(bool,X14) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(resolution,[],[f15467,f10059])).
fof(f15467,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s(bool,X10) != s(bool,X11) | s(bool,sK5) = s(bool,X11)) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(forward_demodulation,[],[f15466,f9082])).
fof(f9082,plain,(
  s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ($spl12 | $spl76)),
  inference(resolution,[],[f6374,f2548])).
fof(f15466,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,X11) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ~p(s(bool,X10))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f15428,f172])).
fof(f15428,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,X11) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ~p(s(bool,sK5)) | ~p(s(bool,X10))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(superposition,[],[f15052,f138])).
fof(f15052,plain,(
  ( ! [X46] : (s(bool,sK5) != s(bool,X46) | s(bool,X46) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f15033,f172])).
fof(f15033,plain,(
  ( ! [X46] : (~p(s(bool,sK5)) | s(bool,X46) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) != s(bool,X46)) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(superposition,[],[f14531,f9082])).
fof(f22048,plain,(
  ( ! [X30,X29] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X30))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X29))) | s(bool,sK5) = s(bool,X30)) ) | ($spl6 | $spl11 | $spl12 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f11643,f14534])).
fof(f14534,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(bool,sK5) != s(bool,X5)) ) | ($spl11 | $spl12 | $spl85 | $spl93)),
  inference(resolution,[],[f14343,f10056])).
fof(f10056,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) ) | ($spl11 | $spl12)),
  inference(subsumption_resolution,[],[f10016,f172])).
fof(f11643,plain,(
  ( ! [X30,X29] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X30))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X29))) | s(bool,sK5) = s(bool,X29) | s(bool,sK5) = s(bool,X30)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f145,f7351])).
fof(f22047,plain,(
  ( ! [X57,X56] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X57))) | s(bool,sK5) = s(bool,X56) | s(bool,sK5) = s(bool,X57)) ) | ($spl12 | $spl100)),
  inference(subsumption_resolution,[],[f11653,f2548])).
fof(f11653,plain,(
  ( ! [X57,X56] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X57))) | p(s(bool,X56)) | s(bool,sK5) = s(bool,X56) | s(bool,sK5) = s(bool,X57)) ) | ($spl12 | $spl100)),
  inference(superposition,[],[f9939,f7351])).
fof(f22045,plain,(
  ( ! [X30,X29] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X30))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X29))) | s(bool,sK5) = s(bool,X30)) ) | ($spl6 | $spl11 | $spl12 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f11718,f14534])).
fof(f11718,plain,(
  ( ! [X30,X29] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X30))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X29))) | s(bool,sK5) = s(bool,X30) | s(bool,sK5) = s(bool,X29)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f145,f7351])).
fof(f22043,plain,(
  ( ! [X57,X56] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X57))) | s(bool,sK5) = s(bool,X57) | s(bool,sK5) = s(bool,X56)) ) | ($spl12 | $spl100)),
  inference(subsumption_resolution,[],[f11728,f2548])).
fof(f11728,plain,(
  ( ! [X57,X56] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X57))) | p(s(bool,X56)) | s(bool,sK5) = s(bool,X57) | s(bool,sK5) = s(bool,X56)) ) | ($spl12 | $spl100)),
  inference(superposition,[],[f9939,f7351])).
fof(f22036,plain,(
  ( ! [X38] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X38)) ) | ($spl12 | $spl80 | $spl100 | $spl137)),
  inference(forward_demodulation,[],[f22035,f7479])).
fof(f22035,plain,(
  ( ! [X38] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s(bool,sK5) = s(bool,X38)) ) | ($spl12 | $spl100 | $spl137)),
  inference(subsumption_resolution,[],[f18938,f2548])).
fof(f18938,plain,(
  ( ! [X38] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | p(s(bool,X38)) | s(bool,sK5) = s(bool,X38)) ) | ($spl12 | $spl100 | $spl137)),
  inference(superposition,[],[f9939,f18337])).
fof(f18337,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X5)) ) | ($spl12 | $spl137)),
  inference(trivial_inequality_removal,[],[f18331])).
fof(f18331,plain,(
  ( ! [X5] : (s(bool,sK5) != s(bool,sK5) | s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X5)) ) | ($spl12 | $spl137)),
  inference(superposition,[],[f12947,f7351])).
fof(f22032,plain,(
  ( ! [X56] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X56)) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl100 | $spl134)),
  inference(subsumption_resolution,[],[f22031,f2548])).
fof(f22031,plain,(
  ( ! [X56] : (p(s(bool,X56)) | s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X56)) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl100 | $spl134)),
  inference(forward_demodulation,[],[f19326,f7479])).
fof(f19326,plain,(
  ( ! [X56] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X56))))) | s(bool,sK5) = s(bool,X56)) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl100 | $spl134)),
  inference(superposition,[],[f9939,f18609])).
fof(f18609,plain,(
  ( ! [X21] : (s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X21)))) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl134)),
  inference(subsumption_resolution,[],[f18588,f7479])).
fof(f18588,plain,(
  ( ! [X21] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X21) | s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X21))))) ) | ($spl12 | $spl76 | $spl85 | $spl134)),
  inference(superposition,[],[f12935,f10781])).
fof(f10781,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X17) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85)),
  inference(subsumption_resolution,[],[f10771,f172])).
fof(f10771,plain,(
  ( ! [X17] : (~p(s(bool,sK5)) | s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X17) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85)),
  inference(superposition,[],[f10111,f9082])).
fof(f22026,plain,(
  $false | ($spl0 | $spl3 | $spl100 | $spl102)),
  inference(subsumption_resolution,[],[f22025,f142])).
fof(f22025,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl0 | $spl100 | $spl102)),
  inference(subsumption_resolution,[],[f21886,f10640])).
fof(f10640,plain,(
  ( ! [X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7)))) ) | ($spl100 | $spl102)),
  inference(subsumption_resolution,[],[f9975,f9939])).
fof(f21931,plain,(
  $false | ($spl0 | $spl3 | $spl100 | $spl102)),
  inference(subsumption_resolution,[],[f21930,f142])).
fof(f21930,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl0 | $spl100 | $spl102)),
  inference(subsumption_resolution,[],[f21783,f10640])).
fof(f21783,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | $spl0),
  inference(superposition,[],[f1302,f139])).
fof(f21923,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,sK5) = s(bool,X57) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X54) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X54),s(fun(bool,bool),X55),s(bool,X56)))) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(forward_demodulation,[],[f21922,f4419])).
fof(f21922,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X57) = s(bool,happ(s(fun(bool,bool),X55),s(bool,sK6(X54)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X54) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X54),s(fun(bool,bool),X55),s(bool,X56)))) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f21775,f15049])).
fof(f21775,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X57) = s(bool,happ(s(fun(bool,bool),X55),s(bool,sK6(X54)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X54) | s(bool,sK5) = s(bool,X57) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X54),s(fun(bool,bool),X55),s(bool,X56)))) ) | $spl12),
  inference(superposition,[],[f1302,f7351])).
fof(f21920,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,sK5) = s(bool,X53) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X50) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X50),s(fun(bool,bool),X51),s(bool,X52)))) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(forward_demodulation,[],[f21919,f4419])).
fof(f21919,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X53) = s(bool,happ(s(fun(bool,bool),X51),s(bool,sK6(X50)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X50) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X50),s(fun(bool,bool),X51),s(bool,X52)))) ) | ($spl12 | $spl30 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f21774,f15049])).
fof(f21774,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X53) = s(bool,happ(s(fun(bool,bool),X51),s(bool,sK6(X50)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X50) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X50),s(fun(bool,bool),X51),s(bool,X52))) | s(bool,sK5) = s(bool,X53)) ) | $spl12),
  inference(superposition,[],[f1302,f7351])).
fof(f21916,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,sK5) = s(bool,X33) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X30) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X30),s(fun(bool,bool),X31),s(bool,X32))))) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f21915,f2548])).
fof(f21915,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,sK5) = s(bool,X33) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X30) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X30),s(fun(bool,bool),X31),s(bool,X32)))) | p(s(bool,X33))) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f21769,f4419])).
fof(f21769,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X33) = s(bool,happ(s(fun(bool,bool),X31),s(bool,sK6(X30)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X30) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X30),s(fun(bool,bool),X31),s(bool,X32)))) | p(s(bool,X33))) )),
  inference(superposition,[],[f1302,f135])).
fof(f21913,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,sK5) = s(bool,X29) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X26),s(fun(bool,bool),X27),s(bool,X28))))) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f21912,f2548])).
fof(f21912,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,sK5) = s(bool,X29) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X26) | p(s(bool,X29)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X26),s(fun(bool,bool),X27),s(bool,X28))))) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f21768,f4419])).
fof(f21768,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,sK6(X26)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X26) | p(s(bool,X29)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X26),s(fun(bool,bool),X27),s(bool,X28))))) )),
  inference(superposition,[],[f1302,f135])).
fof(f21060,plain,(
  ( ! [X8,X7] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8))) | ~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X7),s('type.semanticPrimitives.eq_result',sK1)))))) ) | ($spl3 | $spl78 | $spl86 | $spl94)),
  inference(subsumption_resolution,[],[f21017,f8217])).
fof(f21058,plain,(
  ( ! [X6,X5] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X5),s('type.semanticPrimitives.eq_result',sK1)))))) ) | ($spl3 | $spl78 | $spl86 | $spl94)),
  inference(subsumption_resolution,[],[f21016,f8217])).
fof(f20998,plain,(
  ( ! [X62] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X62))) | s(bool,X62) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) | $spl242),
  inference(cnf_transformation,[],[f20998_D])).
fof(f20998_D,plain,(
  ( ! [X62] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X62))) | s(bool,X62) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f20997,plain,(
  ( ! [X61,X62] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X62))) | s(bool,X62) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X61)) ) | ($spl12 | $spl85 | $spl86 | $spl137 | $spl174)),
  inference(subsumption_resolution,[],[f20975,f2548])).
fof(f20975,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X62))) | s(bool,X62) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X61)) ) | ($spl12 | $spl85 | $spl86 | $spl137 | $spl174)),
  inference(superposition,[],[f18661,f18337])).
fof(f18661,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | s(bool,X3) != s(bool,X4)) ) | ($spl85 | $spl86 | $spl174)),
  inference(subsumption_resolution,[],[f18637,f8215])).
fof(f18637,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | p(s(bool,X3)) | p(s(bool,sK6(sK1)))) ) | ($spl86 | $spl174)),
  inference(superposition,[],[f18356,f135])).
fof(f18356,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8)))) ) | ($spl86 | $spl174)),
  inference(resolution,[],[f16844,f8217])).
fof(f20985,plain,(
  ( ! [X20,X18] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,X18) != s(bool,X20)) ) | $spl240),
  inference(cnf_transformation,[],[f20985_D])).
fof(f20985_D,plain,(
  ( ! [X20,X18] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,X18) != s(bool,X20)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f20986,plain,(
  ( ! [X23,X21,X22] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,X21) != s(bool,X23) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl85 | $spl86 | $spl174)),
  inference(subsumption_resolution,[],[f20958,f2548])).
fof(f20958,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,X21) != s(bool,X23) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl85 | $spl86 | $spl174)),
  inference(superposition,[],[f18661,f7351])).
fof(f20984,plain,(
  ( ! [X19,X20,X18] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,X18) != s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl85 | $spl86 | $spl174)),
  inference(subsumption_resolution,[],[f20957,f2548])).
fof(f20957,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,X18) != s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl85 | $spl86 | $spl174)),
  inference(superposition,[],[f18661,f7351])).
fof(f13594,plain,(
  ( ! [X20,X18] : (p(s(bool,X20)) | s(bool,sK5) = s(bool,X18) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20)))) ) | $spl140),
  inference(cnf_transformation,[],[f13594_D])).
fof(f13594_D,plain,(
  ( ! [X20,X18] : (p(s(bool,X20)) | s(bool,sK5) = s(bool,X18) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f20838,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X23)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f20787,f2548])).
fof(f20787,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(superposition,[],[f18161,f7351])).
fof(f18161,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X7)))) ) | ($spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f18107,f9367])).
fof(f18107,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X7))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6)))) ) | $spl80),
  inference(superposition,[],[f1234,f7479])).
fof(f20837,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f20786,f2548])).
fof(f20786,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | p(s(bool,X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(superposition,[],[f18161,f7351])).
fof(f20818,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X23)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f20756,f2548])).
fof(f20756,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(superposition,[],[f18161,f7351])).
fof(f20817,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f20755,f2548])).
fof(f20755,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | p(s(bool,X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(superposition,[],[f18161,f7351])).
fof(f20407,plain,(
  $false | ($spl12 | $spl73 | $spl76 | $spl78 | $spl94 | $spl176)),
  inference(subsumption_resolution,[],[f20401,f9217])).
fof(f9217,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl12 | $spl73 | $spl76)),
  inference(backward_demodulation,[],[f9082,f5391])).
fof(f5391,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | $spl73),
  inference(cnf_transformation,[],[f5391_D])).
fof(f5391_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f20401,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl78 | $spl94 | $spl176)),
  inference(equality_resolution,[],[f20346])).
fof(f20346,plain,(
  ( ! [X12] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',X12) | s('type.semanticPrimitives.eq_result',X12) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) ) | ($spl78 | $spl94 | $spl176)),
  inference(superposition,[],[f9367,f17108])).
fof(f20315,plain,(
  ( ! [X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X20),s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))))))) ) | $spl238),
  inference(cnf_transformation,[],[f20315_D])).
fof(f20315_D,plain,(
  ( ! [X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X20),s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f20314,plain,(
  ( ! [X19,X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X20),s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))))) | s(bool,X19) = s(bool,sK6(sK1))) ) | ($spl12 | $spl73 | $spl76 | $spl78 | $spl85 | $spl94)),
  inference(subsumption_resolution,[],[f20286,f9367])).
fof(f20286,plain,(
  ( ! [X19,X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X20),s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) | s(bool,X19) = s(bool,sK6(sK1))) ) | ($spl12 | $spl73 | $spl76 | $spl85)),
  inference(superposition,[],[f1261,f11023])).
fof(f11023,plain,(
  ( ! [X7] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X7))) | s(bool,X7) = s(bool,sK6(sK1))) ) | ($spl12 | $spl73 | $spl76 | $spl85)),
  inference(subsumption_resolution,[],[f10932,f9217])).
fof(f10932,plain,(
  ( ! [X7] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X7))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s(bool,X7) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85)),
  inference(superposition,[],[f132,f10781])).
fof(f13332,plain,(
  ( ! [X61] : (s(bool,X61) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X61)))) ) | $spl138),
  inference(cnf_transformation,[],[f13332_D])).
fof(f13332_D,plain,(
  ( ! [X61] : (s(bool,X61) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X61)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f19885,plain,(
  ( ! [X56,X55] : (s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X56))) | s(bool,sK5) = s(bool,X55)) ) | ($spl11 | $spl12 | $spl137)),
  inference(subsumption_resolution,[],[f19867,f2548])).
fof(f19867,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X56))) | s(bool,sK5) = s(bool,X55)) ) | ($spl11 | $spl12 | $spl137)),
  inference(superposition,[],[f10084,f18337])).
fof(f10084,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | ($spl11 | $spl12)),
  inference(resolution,[],[f10056,f135])).
fof(f12172,plain,(
  ( ! [X66,X64] : (s(bool,X64) = s(bool,X66) | s(bool,sK5) = s(bool,X64) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X66)))) ) | $spl128),
  inference(cnf_transformation,[],[f12172_D])).
fof(f12172_D,plain,(
  ( ! [X66,X64] : (s(bool,X64) = s(bool,X66) | s(bool,sK5) = s(bool,X64) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X66)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f19877,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) = s(bool,X23) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl11 | $spl12)),
  inference(subsumption_resolution,[],[f19852,f2548])).
fof(f19852,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s(bool,X21) = s(bool,X23) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl11 | $spl12)),
  inference(superposition,[],[f10084,f7351])).
fof(f19876,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) = s(bool,X20) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl11 | $spl12)),
  inference(subsumption_resolution,[],[f19851,f2548])).
fof(f19851,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s(bool,X18) = s(bool,X20) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl11 | $spl12)),
  inference(superposition,[],[f10084,f7351])).
fof(f17550,plain,(
  ( ! [X56] : (s(bool,X56) != s(bool,sK6(sK1)) | s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) | $spl184),
  inference(cnf_transformation,[],[f17550_D])).
fof(f17550_D,plain,(
  ( ! [X56] : (s(bool,X56) != s(bool,sK6(sK1)) | s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f19465,plain,(
  ( ! [X56,X55] : (s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X56) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X55)) ) | ($spl12 | $spl137 | $spl174)),
  inference(subsumption_resolution,[],[f19447,f2548])).
fof(f19447,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X56) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X55)) ) | ($spl12 | $spl137 | $spl174)),
  inference(superposition,[],[f18352,f18337])).
fof(f18352,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,X0) != s(bool,sK6(sK1))) ) | $spl174),
  inference(resolution,[],[f16844,f135])).
fof(f17540,plain,(
  ( ! [X20,X18] : (s(bool,X18) = s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,X20) != s(bool,sK6(sK1))) ) | $spl182),
  inference(cnf_transformation,[],[f17540_D])).
fof(f17540_D,plain,(
  ( ! [X20,X18] : (s(bool,X18) = s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,X20) != s(bool,sK6(sK1))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f19457,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) = s(bool,X23) | s(bool,X23) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl174)),
  inference(subsumption_resolution,[],[f19432,f2548])).
fof(f19432,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s(bool,X21) = s(bool,X23) | s(bool,X23) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl174)),
  inference(superposition,[],[f18352,f7351])).
fof(f19456,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) = s(bool,X20) | s(bool,X20) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl174)),
  inference(subsumption_resolution,[],[f19431,f2548])).
fof(f19431,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s(bool,X18) = s(bool,X20) | s(bool,X20) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl174)),
  inference(superposition,[],[f18352,f7351])).
fof(f19234,plain,(
  ( ! [X21] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X21),s('type.semanticPrimitives.eq_result',sK1))))))) ) | $spl236),
  inference(cnf_transformation,[],[f19234_D])).
fof(f19234_D,plain,(
  ( ! [X21] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X21),s('type.semanticPrimitives.eq_result',sK1))))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f19235,plain,(
  ( ! [X23,X22] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X23),s('type.semanticPrimitives.eq_result',sK1)))))) | s(bool,sK5) = s(bool,X22)) ) | ($spl12 | $spl78 | $spl86 | $spl94)),
  inference(subsumption_resolution,[],[f19180,f9367])).
fof(f19180,plain,(
  ( ! [X23,X22] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X23),s('type.semanticPrimitives.eq_result',sK1)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X22))) | s(bool,sK5) = s(bool,X22)) ) | ($spl12 | $spl86)),
  inference(superposition,[],[f1261,f10134])).
fof(f10134,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | s(bool,sK5) = s(bool,X4)) ) | ($spl12 | $spl86)),
  inference(resolution,[],[f8217,f2548])).
fof(f19233,plain,(
  ( ! [X21,X20] : (s(bool,sK5) = s(bool,X20) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X21),s('type.semanticPrimitives.eq_result',sK1))))))) ) | ($spl12 | $spl76 | $spl78 | $spl86 | $spl94)),
  inference(forward_demodulation,[],[f19232,f9082])).
fof(f19232,plain,(
  ( ! [X21,X20] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X21),s('type.semanticPrimitives.eq_result',sK1)))))) | s(bool,X20) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl78 | $spl86 | $spl94)),
  inference(subsumption_resolution,[],[f19179,f9367])).
fof(f19179,plain,(
  ( ! [X21,X20] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X21),s('type.semanticPrimitives.eq_result',sK1)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,X20) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f1261,f11533])).
fof(f11533,plain,(
  ( ! [X28] : (s(bool,X28) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X28)))) ) | ($spl12 | $spl76 | $spl86)),
  inference(subsumption_resolution,[],[f11517,f172])).
fof(f11517,plain,(
  ( ! [X28] : (~p(s(bool,sK5)) | s(bool,X28) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X28)))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f10132,f9082])).
fof(f10132,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | $spl86),
  inference(resolution,[],[f8217,f138])).
fof(f14176,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X31) = s(bool,sK6(sK1))) ) | $spl144),
  inference(cnf_transformation,[],[f14176_D])).
fof(f14176_D,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X31) = s(bool,sK6(sK1))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f19158,plain,(
  ( ! [X68,X69] : (s(bool,X69) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X69) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X68)) ) | ($spl12 | $spl137 | $spl142)),
  inference(subsumption_resolution,[],[f19032,f2548])).
fof(f19032,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | s(bool,X69) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X69) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X68)) ) | ($spl12 | $spl137 | $spl142)),
  inference(superposition,[],[f14173,f18337])).
fof(f17313,plain,(
  ( ! [X20,X18] : (s(bool,sK5) != s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,X18) != s(bool,X20)) ) | $spl180),
  inference(cnf_transformation,[],[f17313_D])).
fof(f17313_D,plain,(
  ( ! [X20,X18] : (s(bool,sK5) != s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,X18) != s(bool,X20)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f18708,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s(bool,sK5) != s(bool,X23) | s(bool,X21) != s(bool,X23) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl85 | $spl93 | $spl174)),
  inference(subsumption_resolution,[],[f18685,f14343])).
fof(f18685,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s(bool,sK5) != s(bool,X23) | s(bool,X21) != s(bool,X23) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl85 | $spl93 | $spl174)),
  inference(superposition,[],[f18473,f7351])).
fof(f18707,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s(bool,sK5) != s(bool,X20) | s(bool,X18) != s(bool,X20) | s(bool,sK5) = s(bool,X18)) ) | ($spl12 | $spl85 | $spl93 | $spl174)),
  inference(subsumption_resolution,[],[f18684,f14343])).
fof(f18684,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s(bool,sK5) != s(bool,X20) | s(bool,X18) != s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl85 | $spl93 | $spl174)),
  inference(superposition,[],[f18473,f7351])).
fof(f18295,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | s(bool,sK5) = s(bool,X22)) ) | ($spl12 | $spl80 | $spl152)),
  inference(forward_demodulation,[],[f18294,f7479])).
fof(f18294,plain,(
  ( ! [X21,X22] : (s(bool,sK5) = s(bool,X22) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X21)))))) ) | ($spl12 | $spl152)),
  inference(subsumption_resolution,[],[f18238,f2548])).
fof(f18238,plain,(
  ( ! [X21,X22] : (s(bool,sK5) = s(bool,X22) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X21))))) | p(s(bool,X22))) ) | $spl152),
  inference(superposition,[],[f135,f14504])).
fof(f18293,plain,(
  ( ! [X19,X20] : (p(s(bool,X19)) | s(bool,sK5) = s(bool,X20)) ) | ($spl12 | $spl80 | $spl152)),
  inference(forward_demodulation,[],[f18292,f7479])).
fof(f18292,plain,(
  ( ! [X19,X20] : (s(bool,sK5) = s(bool,X20) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19)))))) ) | ($spl12 | $spl152)),
  inference(subsumption_resolution,[],[f18237,f2548])).
fof(f18237,plain,(
  ( ! [X19,X20] : (s(bool,sK5) = s(bool,X20) | p(s(bool,X20)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19)))))) ) | $spl152),
  inference(superposition,[],[f135,f14504])).
fof(f18288,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,sK5) = s(bool,X8)) ) | ($spl12 | $spl80 | $spl152)),
  inference(forward_demodulation,[],[f18287,f7479])).
fof(f18287,plain,(
  ( ! [X8,X7] : (s(bool,sK5) = s(bool,X8) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X7)))))) ) | ($spl12 | $spl152)),
  inference(subsumption_resolution,[],[f18218,f2548])).
fof(f18218,plain,(
  ( ! [X8,X7] : (s(bool,sK5) = s(bool,X8) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X7))))) | p(s(bool,X8))) ) | $spl152),
  inference(superposition,[],[f14504,f135])).
fof(f18286,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | s(bool,sK5) = s(bool,X6)) ) | ($spl12 | $spl80 | $spl152)),
  inference(forward_demodulation,[],[f18285,f7479])).
fof(f18285,plain,(
  ( ! [X6,X5] : (s(bool,sK5) = s(bool,X6) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))))) ) | ($spl12 | $spl152)),
  inference(subsumption_resolution,[],[f18217,f2548])).
fof(f18217,plain,(
  ( ! [X6,X5] : (s(bool,sK5) = s(bool,X6) | p(s(bool,X6)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))))) ) | $spl152),
  inference(superposition,[],[f14504,f135])).
fof(f16843,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,sK6(sK1)) | s(bool,X4) = s(bool,sK6(sK1)) | ~p(s(bool,X5))) ) | ($spl80 | $spl85 | $spl144)),
  inference(forward_demodulation,[],[f16842,f7479])).
fof(f16846,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6(sK1)) | s(bool,X6) = s(bool,sK6(sK1)) | ~p(s(bool,X7))) ) | ($spl80 | $spl85 | $spl144)),
  inference(forward_demodulation,[],[f16845,f7479])).
fof(f17585,plain,(
  s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | $spl214),
  inference(cnf_transformation,[],[f17585_D])).
fof(f17585_D,plain,(
  s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f16839,plain,(
  ( ! [X31] : (s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X31)) ) | ($spl12 | $spl137 | $spl144)),
  inference(trivial_inequality_removal,[],[f16804])).
fof(f16804,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,X31) | s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X31)) ) | ($spl12 | $spl137 | $spl144)),
  inference(superposition,[],[f14176,f14909])).
fof(f14909,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X5)) ) | ($spl12 | $spl137)),
  inference(trivial_inequality_removal,[],[f14904])).
fof(f14904,plain,(
  ( ! [X5] : (s(bool,sK5) != s(bool,sK5) | s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X5)) ) | ($spl12 | $spl137)),
  inference(superposition,[],[f12947,f7351])).
fof(f12947,plain,(
  s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | $spl137),
  inference(cnf_transformation,[],[f12947_D])).
fof(f12947_D,plain,(
  s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f17588,plain,(
  ( ! [X18] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X18) = s(bool,sK6(sK1))) ) | ($spl11 | $spl12 | $spl76 | $spl85 | $spl86 | $spl144)),
  inference(forward_demodulation,[],[f16849,f9082])).
fof(f16849,plain,(
  ( ! [X18] : (s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X18) = s(bool,sK6(sK1))) ) | ($spl11 | $spl12 | $spl76 | $spl85 | $spl86 | $spl144)),
  inference(subsumption_resolution,[],[f16797,f14314])).
fof(f14314,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(bool,X5) = s(bool,sK6(sK1))) ) | ($spl11 | $spl12 | $spl85)),
  inference(resolution,[],[f8215,f10084])).
fof(f16797,plain,(
  ( ! [X18] : (s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X18) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18)))) ) | ($spl12 | $spl76 | $spl86 | $spl144)),
  inference(superposition,[],[f14176,f11533])).
fof(f17589,plain,(
  ( ! [X17] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X17) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85 | $spl144)),
  inference(forward_demodulation,[],[f16841,f9082])).
fof(f16841,plain,(
  ( ! [X17] : (s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X17) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85 | $spl144)),
  inference(duplicate_literal_removal,[],[f16796])).
fof(f16796,plain,(
  ( ! [X17] : (s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X17) = s(bool,sK6(sK1)) | s(bool,X17) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85 | $spl144)),
  inference(superposition,[],[f14176,f10781])).
fof(f18007,plain,(
  $false | ($spl3 | $spl11 | $spl12 | $spl82)),
  inference(subsumption_resolution,[],[f18006,f142])).
fof(f18006,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl11 | $spl12 | $spl82)),
  inference(subsumption_resolution,[],[f17969,f148])).
fof(f17969,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl12 | $spl82)),
  inference(superposition,[],[f8084,f17777])).
fof(f17777,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1)) ) | ($spl12 | $spl82)),
  inference(backward_demodulation,[],[f17620,f1259])).
fof(f17620,plain,(
  ( ! [X6] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X6))))) ) | ($spl12 | $spl82)),
  inference(resolution,[],[f7480,f2548])).
fof(f17998,plain,(
  $false | ($spl3 | $spl11 | $spl12 | $spl82)),
  inference(subsumption_resolution,[],[f17997,f142])).
fof(f17997,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl11 | $spl12 | $spl82)),
  inference(subsumption_resolution,[],[f17960,f148])).
fof(f17960,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl12 | $spl82)),
  inference(superposition,[],[f17777,f8084])).
fof(f17763,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16) | s(bool,sK5) = s(bool,sK6(X16))) ) | $spl234),
  inference(cnf_transformation,[],[f17763_D])).
fof(f17763_D,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16) | s(bool,sK5) = s(bool,sK6(X16))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f14783,plain,(
  ( ! [X10] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X10)))) ) | $spl160),
  inference(cnf_transformation,[],[f14783_D])).
fof(f14783_D,plain,(
  ( ! [X10] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X10)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f17994,plain,(
  ( ! [X17,X18] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X17) | s(bool,sK5) = s(bool,sK6(X17))) ) | ($spl12 | $spl73 | $spl76 | $spl82 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f17953,f14536])).
fof(f14536,plain,(
  ( ! [X8] : (s(bool,sK5) != s(bool,X8) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8)))) ) | ($spl12 | $spl73 | $spl76 | $spl85 | $spl93)),
  inference(resolution,[],[f14343,f11194])).
fof(f11194,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl12 | $spl73 | $spl76)),
  inference(subsumption_resolution,[],[f11193,f9217])).
fof(f11193,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,X2))) ) | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f11176,f172])).
fof(f11176,plain,(
  ( ! [X2] : (~p(s(bool,sK5)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,X2))) ) | ($spl12 | $spl76)),
  inference(superposition,[],[f1236,f9082])).
fof(f17953,plain,(
  ( ! [X17,X18] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X17) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,sK6(X17))) ) | ($spl12 | $spl82)),
  inference(superposition,[],[f17777,f7351])).
fof(f17993,plain,(
  ( ! [X15,X16] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X15) | s(bool,sK5) = s(bool,sK6(X15))) ) | ($spl12 | $spl73 | $spl76 | $spl82 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f17952,f14536])).
fof(f17952,plain,(
  ( ! [X15,X16] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X15) | s(bool,sK5) = s(bool,sK6(X15)) | s(bool,sK5) = s(bool,X16)) ) | ($spl12 | $spl82)),
  inference(superposition,[],[f17777,f7351])).
fof(f17992,plain,(
  ( ! [X8,X7] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X7) | p(s(bool,sK6(X7)))) ) | ($spl12 | $spl73 | $spl76 | $spl82)),
  inference(subsumption_resolution,[],[f17949,f11194])).
fof(f17949,plain,(
  ( ! [X8,X7] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X7) | p(s(bool,sK6(X7))) | p(s(bool,X8))) ) | ($spl12 | $spl82)),
  inference(superposition,[],[f17777,f135])).
fof(f17991,plain,(
  ( ! [X6,X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X5) | p(s(bool,sK6(X5)))) ) | ($spl12 | $spl73 | $spl76 | $spl82)),
  inference(subsumption_resolution,[],[f17948,f11194])).
fof(f17948,plain,(
  ( ! [X6,X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X5) | p(s(bool,X6)) | p(s(bool,sK6(X5)))) ) | ($spl12 | $spl82)),
  inference(superposition,[],[f17777,f135])).
fof(f17935,plain,(
  $false | ($spl12 | $spl78 | $spl82 | $spl94 | $spl230)),
  inference(subsumption_resolution,[],[f17934,f9367])).
fof(f17934,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl12 | $spl82 | $spl230)),
  inference(subsumption_resolution,[],[f17908,f172])).
fof(f17908,plain,(
  ( ! [X2] : (~p(s(bool,sK5)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl12 | $spl82 | $spl230)),
  inference(superposition,[],[f17761,f17620])).
fof(f17933,plain,(
  $false | ($spl12 | $spl73 | $spl76 | $spl230)),
  inference(subsumption_resolution,[],[f17932,f9217])).
fof(f17932,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl12 | $spl76 | $spl230)),
  inference(subsumption_resolution,[],[f17907,f172])).
fof(f17907,plain,(
  ~p(s(bool,sK5)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl12 | $spl76 | $spl230)),
  inference(superposition,[],[f17761,f9082])).
fof(f17570,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | $spl200),
  inference(cnf_transformation,[],[f17570_D])).
fof(f17570_D,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f17764,plain,(
  ( ! [X20] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X20) | s(bool,sK5) = s(bool,sK6(X20))) ) | ($spl3 | $spl12 | $spl85)),
  inference(forward_demodulation,[],[f17721,f8084])).
fof(f17721,plain,(
  ( ! [X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X20) | s(bool,sK5) = s(bool,sK6(X20))) ) | ($spl12 | $spl85)),
  inference(superposition,[],[f1259,f10061])).
fof(f17565,plain,(
  ( ! [X39] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X39))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X39)))))) | s(bool,sK5) = s(bool,X39)) ) | $spl196),
  inference(cnf_transformation,[],[f17565_D])).
fof(f17565_D,plain,(
  ( ! [X39] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X39))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X39)))))) | s(bool,sK5) = s(bool,X39)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f17720,plain,(
  ( ! [X19,X18] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X18) | s(bool,sK5) = s(bool,X19) | s(bool,sK5) = s(bool,sK6(X18))) ) | $spl12),
  inference(superposition,[],[f1259,f7351])).
fof(f17719,plain,(
  ( ! [X17,X16] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X17))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X17)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16) | s(bool,sK5) = s(bool,sK6(X16)) | s(bool,sK5) = s(bool,X17)) ) | $spl12),
  inference(superposition,[],[f1259,f7351])).
fof(f17716,plain,(
  ( ! [X10,X11] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X11))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X11)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X10) | ~p(s(bool,X11)) | ~p(s(bool,sK6(X10)))) )),
  inference(superposition,[],[f1259,f138])).
fof(f17715,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X9))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X8) | ~p(s(bool,sK6(X8))) | ~p(s(bool,X9))) )),
  inference(superposition,[],[f1259,f138])).
fof(f17714,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X7))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X7)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X6) | p(s(bool,sK6(X6))) | p(s(bool,X7))) )),
  inference(superposition,[],[f1259,f135])).
fof(f17713,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | p(s(bool,X5)) | p(s(bool,sK6(X4)))) )),
  inference(superposition,[],[f1259,f135])).
fof(f17707,plain,(
  $false | ($spl3 | $spl104)),
  inference(trivial_inequality_removal,[],[f17699])).
fof(f17699,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',sK1) | ($spl3 | $spl104)),
  inference(superposition,[],[f10008,f8084])).
fof(f17671,plain,(
  $false | ($spl82 | $spl217)),
  inference(subsumption_resolution,[],[f17587,f7480])).
fof(f14504,plain,(
  ( ! [X64] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X64))))) ) | $spl152),
  inference(cnf_transformation,[],[f14504_D])).
fof(f14504_D,plain,(
  ( ! [X64] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X64))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f17669,plain,(
  ( ! [X17,X18] : (s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X17))))) ) | ($spl12 | $spl82)),
  inference(subsumption_resolution,[],[f17661,f2548])).
fof(f17661,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X17))))) ) | ($spl12 | $spl82)),
  inference(superposition,[],[f7480,f7351])).
fof(f17668,plain,(
  ( ! [X15,X16] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X15)))) | s(bool,sK5) = s(bool,X16)) ) | ($spl12 | $spl82)),
  inference(subsumption_resolution,[],[f17660,f2548])).
fof(f17660,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X15)))) | s(bool,sK5) = s(bool,X16)) ) | ($spl12 | $spl82)),
  inference(superposition,[],[f7480,f7351])).
fof(f17576,plain,(
  ( ! [X19] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19))))))) ) | $spl212),
  inference(cnf_transformation,[],[f17576_D])).
fof(f17576_D,plain,(
  ( ! [X19] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19))))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f17612,plain,(
  ( ! [X10,X9] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X9),s(bool,X10)))) ) | $spl226),
  inference(cnf_transformation,[],[f17612_D])).
fof(f17612_D,plain,(
  ( ! [X10,X9] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X9),s(bool,X10)))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f11296,plain,(
  ( ! [X10,X8,X9] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X9),s(bool,X10))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X8))))))) ) | ($spl12 | $spl30)),
  inference(superposition,[],[f4529,f132])).
fof(f4529,plain,(
  ( ! [X59,X60,X58] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X58))),s(fun(bool,bool),X59),s(bool,X60)))) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f4505,f1886])).
fof(f4505,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,X58)))) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X58))),s(fun(bool,bool),X59),s(bool,X60)))) ) | $spl12),
  inference(superposition,[],[f2548,f133])).
fof(f15817,plain,(
  ( ! [X0,X1] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f15775,f14343])).
fof(f15775,plain,(
  ( ! [X0,X1] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) | s(bool,sK5) = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl134),
  inference(superposition,[],[f12935,f135])).
fof(f15819,plain,(
  ( ! [X2,X3] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) | p(s(bool,X2)) | p(s(bool,X3))) ) | ($spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f15776,f14343])).
fof(f15776,plain,(
  ( ! [X2,X3] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3)))) | s(bool,sK5) = s(bool,X2) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl134),
  inference(superposition,[],[f12935,f135])).
fof(f15825,plain,(
  ( ! [X8,X7] : (s(bool,sK5) != s(bool,X8) | s(bool,sK5) = s(bool,X7) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X7)))))) ) | ($spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f15803,f14343])).
fof(f15803,plain,(
  ( ! [X8,X7] : (s(bool,sK5) != s(bool,X8) | s(bool,sK5) = s(bool,X7) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X7))))) | ~p(s(bool,X8))) ) | $spl134),
  inference(superposition,[],[f12935,f138])).
fof(f15827,plain,(
  ( ! [X10,X9] : (s(bool,sK5) != s(bool,X10) | s(bool,sK5) = s(bool,X9) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))))) ) | ($spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f15804,f14343])).
fof(f15804,plain,(
  ( ! [X10,X9] : (s(bool,sK5) != s(bool,X10) | s(bool,sK5) = s(bool,X9) | ~p(s(bool,X10)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X9)))))) ) | $spl134),
  inference(superposition,[],[f12935,f138])).
fof(f16842,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X4) = s(bool,sK6(sK1)) | ~p(s(bool,X5))) ) | ($spl85 | $spl144)),
  inference(subsumption_resolution,[],[f16789,f10059])).
fof(f16789,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X4) = s(bool,sK6(sK1)) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f16845,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X6) = s(bool,sK6(sK1)) | ~p(s(bool,X7))) ) | ($spl85 | $spl144)),
  inference(subsumption_resolution,[],[f16790,f10059])).
fof(f16790,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X6) = s(bool,sK6(sK1)) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f16861,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,sK6(sK1)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) | ~p(s(bool,X10))) ) | ($spl85 | $spl144)),
  inference(subsumption_resolution,[],[f16828,f10111])).
fof(f16828,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,X11) = s(bool,sK6(sK1)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) | ~p(s(bool,X10))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f16863,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,sK6(sK1)) | ~p(s(bool,X12)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) ) | ($spl85 | $spl144)),
  inference(subsumption_resolution,[],[f16829,f10111])).
fof(f16829,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,X13) = s(bool,sK6(sK1)) | ~p(s(bool,X12)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) ) | $spl144),
  inference(superposition,[],[f14176,f138])).
fof(f17583,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,sK6(sK1)) | s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) | ($spl12 | $spl76 | $spl85 | $spl144)),
  inference(subsumption_resolution,[],[f16835,f10781])).
fof(f16835,plain,(
  ( ! [X23] : (s(bool,X23) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X23) = s(bool,sK6(sK1)) | s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) | ($spl12 | $spl76 | $spl85 | $spl144)),
  inference(superposition,[],[f14176,f10781])).
fof(f16869,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))))) ) | ($spl12 | $spl76 | $spl85 | $spl86 | $spl144)),
  inference(subsumption_resolution,[],[f16836,f10781])).
fof(f16836,plain,(
  ( ! [X24] : (s(bool,X24) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X24) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))))) ) | ($spl12 | $spl76 | $spl86 | $spl144)),
  inference(superposition,[],[f14176,f11533])).
fof(f17582,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X16) | s('type.semanticPrimitives.eq_result',X16) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X16))))) ) | $spl3),
  inference(forward_demodulation,[],[f16942,f8084])).
fof(f16942,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X16) | s('type.semanticPrimitives.eq_result',X16) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X16)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))))) ) | $spl3),
  inference(superposition,[],[f1229,f8084])).
fof(f17581,plain,(
  ( ! [X17,X18] : (s(bool,sK5) = s(bool,X17) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X18) | s('type.semanticPrimitives.eq_result',X18) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X18)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X17))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X17))))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(forward_demodulation,[],[f16943,f9082])).
fof(f16943,plain,(
  ( ! [X17,X18] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X18) | s('type.semanticPrimitives.eq_result',X18) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X18)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X17))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X17)))))) | s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f1229,f11533])).
fof(f16944,plain,(
  ( ! [X19,X20] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X20) | s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X20)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19)))))) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl86)),
  inference(superposition,[],[f1229,f10134])).
fof(f17580,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X16) | s('type.semanticPrimitives.eq_result',X16) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X16))))) ) | $spl3),
  inference(forward_demodulation,[],[f16955,f8084])).
fof(f16955,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X16) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s('type.semanticPrimitives.eq_result',X16) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X16))))) ) | $spl3),
  inference(superposition,[],[f1229,f8084])).
fof(f17579,plain,(
  ( ! [X17,X18] : (s(bool,sK5) = s(bool,X17) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X18) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X17))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X17)))))) | s('type.semanticPrimitives.eq_result',X18) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X18))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(forward_demodulation,[],[f16956,f9082])).
fof(f16956,plain,(
  ( ! [X17,X18] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X18) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X17))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X17)))))) | s('type.semanticPrimitives.eq_result',X18) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X18)))) | s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f1229,f11533])).
fof(f16957,plain,(
  ( ! [X19,X20] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X20) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19)))))) | s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X20)))) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl86)),
  inference(superposition,[],[f1229,f10134])).
fof(f17578,plain,(
  ( ! [X13] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X13) | s('type.semanticPrimitives.eq_result',X13) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X13))))) ) | $spl3),
  inference(forward_demodulation,[],[f17015,f8084])).
fof(f17015,plain,(
  ( ! [X13] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X13) | s('type.semanticPrimitives.eq_result',X13) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X13)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))))) ) | $spl3),
  inference(superposition,[],[f8084,f1229])).
fof(f17575,plain,(
  ( ! [X20] : (s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X20))))) ) | $spl210),
  inference(cnf_transformation,[],[f17575_D])).
fof(f17575_D,plain,(
  ( ! [X20] : (s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X20))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f17123,plain,(
  ( ! [X19,X20] : (s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X20)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19))))))) ) | ($spl78 | $spl94)),
  inference(subsumption_resolution,[],[f17017,f132])).
fof(f17017,plain,(
  ( ! [X19,X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',X20) | s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X20)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19))))))) ) | ($spl78 | $spl94)),
  inference(superposition,[],[f9367,f1229])).
fof(f17018,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X22) | s(bool,sK5) = s(bool,X21) | s('type.semanticPrimitives.eq_result',X22) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X22)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X21))))))) ) | ($spl12 | $spl86)),
  inference(superposition,[],[f10134,f1229])).
fof(f17124,plain,(
  ( ! [X24,X23] : (s(bool,sK5) = s(bool,X23) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X24) | s('type.semanticPrimitives.eq_result',X24) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X24)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X23))))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(forward_demodulation,[],[f17019,f9082])).
fof(f17019,plain,(
  ( ! [X24,X23] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X24) | s(bool,X23) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result',X24) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X24)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X23))))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f11533,f1229])).
fof(f17573,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X25))))))) ) | $spl206),
  inference(cnf_transformation,[],[f17573_D])).
fof(f17573_D,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X25))))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f17020,plain,(
  ( ! [X26,X25] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',X26) | s(bool,X25) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result',X26) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X26)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X25))))))) ) | ($spl11 | $spl12 | $spl85)),
  inference(superposition,[],[f14314,f1229])).
fof(f17571,plain,(
  ( ! [X27] : (s(bool,sK5) != s(bool,X27) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X27))))))) ) | $spl202),
  inference(cnf_transformation,[],[f17571_D])).
fof(f17571_D,plain,(
  ( ! [X27] : (s(bool,sK5) != s(bool,X27) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X27))))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f17021,plain,(
  ( ! [X28,X27] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',X28) | s(bool,sK5) != s(bool,X27) | s('type.semanticPrimitives.eq_result',X28) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X28)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X27))))))) ) | ($spl11 | $spl12 | $spl85 | $spl93)),
  inference(superposition,[],[f14534,f1229])).
fof(f17577,plain,(
  ( ! [X13] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X13) | s('type.semanticPrimitives.eq_result',X13) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X13))))) ) | $spl3),
  inference(forward_demodulation,[],[f17040,f8084])).
fof(f17040,plain,(
  ( ! [X13] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X13) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s('type.semanticPrimitives.eq_result',X13) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X13))))) ) | $spl3),
  inference(superposition,[],[f8084,f1229])).
fof(f17131,plain,(
  ( ! [X19,X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19)))))) | s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X20))))) ) | ($spl78 | $spl94)),
  inference(subsumption_resolution,[],[f17042,f132])).
fof(f17042,plain,(
  ( ! [X19,X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',X20) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X19)))))) | s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X20))))) ) | ($spl78 | $spl94)),
  inference(superposition,[],[f9367,f1229])).
fof(f17043,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X22) | s(bool,sK5) = s(bool,X21) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X21)))))) | s('type.semanticPrimitives.eq_result',X22) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X22))))) ) | ($spl12 | $spl86)),
  inference(superposition,[],[f10134,f1229])).
fof(f17132,plain,(
  ( ! [X24,X23] : (s(bool,sK5) = s(bool,X23) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X24) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X23)))))) | s('type.semanticPrimitives.eq_result',X24) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X24))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(forward_demodulation,[],[f17044,f9082])).
fof(f17044,plain,(
  ( ! [X24,X23] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X24) | s(bool,X23) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X23)))))) | s('type.semanticPrimitives.eq_result',X24) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X24))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f11533,f1229])).
fof(f17045,plain,(
  ( ! [X26,X25] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',X26) | s(bool,X25) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X25)))))) | s('type.semanticPrimitives.eq_result',X26) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X26))))) ) | ($spl11 | $spl12 | $spl85)),
  inference(superposition,[],[f14314,f1229])).
fof(f17046,plain,(
  ( ! [X28,X27] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',X28) | s(bool,sK5) != s(bool,X27) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X27)))))) | s('type.semanticPrimitives.eq_result',X28) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X28))))) ) | ($spl11 | $spl12 | $spl85 | $spl93)),
  inference(superposition,[],[f14534,f1229])).
fof(f17568,plain,(
  ( ! [X30,X28,X29] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))))) | s(X28,X29) = s(X28,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X28),X30),s(X28,X29)))) ) | $spl3),
  inference(forward_demodulation,[],[f17336,f8084])).
fof(f17336,plain,(
  ( ! [X30,X28,X29] : (s(X28,X29) = s(X28,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X28),X30),s(X28,X29))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))))) ) | $spl3),
  inference(superposition,[],[f1244,f8084])).
fof(f17567,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,sK5) = s(bool,X35) | s(X36,X37) = s(X36,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X36),X38),s(X36,X37))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X35))))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(forward_demodulation,[],[f17338,f9082])).
fof(f17338,plain,(
  ( ! [X37,X35,X38,X36] : (s(X36,X37) = s(X36,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X36),X38),s(X36,X37))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X35)))))) | s(bool,X35) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f1244,f11533])).
fof(f17339,plain,(
  ( ! [X39,X41,X42,X40] : (s(X40,X41) = s(X40,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X40),X42),s(X40,X41))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X39))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X39)))))) | s(bool,sK5) = s(bool,X39)) ) | ($spl12 | $spl86)),
  inference(superposition,[],[f1244,f10134])).
fof(f17564,plain,(
  ( ! [X39,X41,X40] : (s(bool,sK6(sK1)) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X39))),s(fun(bool,bool),X40),s(bool,X41)))) ) | $spl194),
  inference(cnf_transformation,[],[f17564_D])).
fof(f17564_D,plain,(
  ( ! [X39,X41,X40] : (s(bool,sK6(sK1)) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X39))),s(fun(bool,bool),X40),s(bool,X41)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f16853,plain,(
  ( ! [X39,X41,X40] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK6(sK1)) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X39))),s(fun(bool,bool),X40),s(bool,X41)))) ) | ($spl12 | $spl30 | $spl144)),
  inference(forward_demodulation,[],[f16808,f4419])).
fof(f16808,plain,(
  ( ! [X39,X41,X40] : (s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) != s(bool,happ(s(fun(bool,bool),X40),s(bool,X39))) | s(bool,sK6(sK1)) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X39))),s(fun(bool,bool),X40),s(bool,X41)))) ) | $spl144),
  inference(superposition,[],[f14176,f133])).
fof(f15814,plain,(
  ( ! [X16] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X16)) ) | ($spl12 | $spl85 | $spl134)),
  inference(duplicate_literal_removal,[],[f15783])).
fof(f15783,plain,(
  ( ! [X16] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X16) | s(bool,sK5) = s(bool,X16)) ) | ($spl12 | $spl85 | $spl134)),
  inference(superposition,[],[f12935,f10061])).
fof(f17561,plain,(
  ( ! [X67] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X67))))) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl93 | $spl134)),
  inference(subsumption_resolution,[],[f16662,f15476])).
fof(f15476,plain,(
  ( ! [X27] : (s(bool,sK5) != s(bool,X27) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X27))))) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl93)),
  inference(forward_demodulation,[],[f15450,f9082])).
fof(f15450,plain,(
  ( ! [X27] : (s(bool,sK5) != s(bool,X27) | s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X27))))) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f15052,f7479])).
fof(f16662,plain,(
  ( ! [X67] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X67)))) | s(bool,sK5) = s(bool,X67)) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl134)),
  inference(superposition,[],[f12935,f16101])).
fof(f16101,plain,(
  ( ! [X75] : (s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X75)))) | s(bool,sK5) = s(bool,X75)) ) | ($spl12 | $spl76 | $spl80 | $spl85 | $spl134)),
  inference(subsumption_resolution,[],[f15978,f7479])).
fof(f15978,plain,(
  ( ! [X75] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X75) | s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X75))))) ) | ($spl12 | $spl76 | $spl85 | $spl134)),
  inference(superposition,[],[f12935,f10781])).
fof(f17560,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl192),
  inference(cnf_transformation,[],[f17560_D])).
fof(f17560_D,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f16791,plain,(
  ( ! [X8,X9] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X8) = s(bool,sK6(sK1)) | s(bool,X8) = s(bool,X9) | s(bool,sK5) = s(bool,X9)) ) | ($spl12 | $spl144)),
  inference(superposition,[],[f14176,f7351])).
fof(f16792,plain,(
  ( ! [X10,X11] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X10) = s(bool,sK6(sK1)) | s(bool,X10) = s(bool,X11) | s(bool,sK5) = s(bool,X11)) ) | ($spl12 | $spl144)),
  inference(superposition,[],[f14176,f7351])).
fof(f17557,plain,(
  ( ! [X29] : (s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X29)))) | s(bool,X29) = s(bool,sK6(sK1))) ) | $spl190),
  inference(cnf_transformation,[],[f17557_D])).
fof(f17557_D,plain,(
  ( ! [X29] : (s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X29)))) | s(bool,X29) = s(bool,sK6(sK1))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f16802,plain,(
  ( ! [X29] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X29)))) | s(bool,X29) = s(bool,sK6(sK1))) ) | ($spl12 | $spl80 | $spl85 | $spl144)),
  inference(superposition,[],[f14176,f10254])).
fof(f10254,plain,(
  ( ! [X11] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X11)))) | s(bool,X11) = s(bool,sK6(sK1))) ) | ($spl12 | $spl80 | $spl85)),
  inference(superposition,[],[f10061,f7479])).
fof(f17556,plain,(
  ( ! [X33] : (s(bool,sK6(sK1)) = s(bool,sK6(X33)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X33)) ) | $spl188),
  inference(cnf_transformation,[],[f17556_D])).
fof(f17556_D,plain,(
  ( ! [X33] : (s(bool,sK6(sK1)) = s(bool,sK6(X33)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X33)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f16806,plain,(
  ( ! [X33] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK6(sK1)) = s(bool,sK6(X33)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X33) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33)) ) | ($spl12 | $spl86 | $spl144)),
  inference(superposition,[],[f14176,f10498])).
fof(f10498,plain,(
  ( ! [X0] : (s(bool,sK5) = s(bool,sK6(X0)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl12 | $spl86)),
  inference(superposition,[],[f10134,f132])).
fof(f17554,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X17)) ) | $spl186),
  inference(cnf_transformation,[],[f17554_D])).
fof(f17554_D,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X17)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f16865,plain,(
  ( ! [X14,X15] : (s(bool,X14) = s(bool,sK6(sK1)) | s(bool,X15) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X15)) ) | ($spl12 | $spl85 | $spl144)),
  inference(subsumption_resolution,[],[f16830,f10061])).
fof(f16830,plain,(
  ( ! [X14,X15] : (s(bool,sK5) != s(bool,X14) | s(bool,X14) = s(bool,sK6(sK1)) | s(bool,X15) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X15)) ) | ($spl12 | $spl144)),
  inference(superposition,[],[f14176,f7351])).
fof(f16866,plain,(
  ( ! [X17,X16] : (s(bool,X16) = s(bool,sK6(sK1)) | s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X17)) ) | ($spl12 | $spl85 | $spl144)),
  inference(subsumption_resolution,[],[f16831,f10061])).
fof(f16831,plain,(
  ( ! [X17,X16] : (s(bool,sK5) != s(bool,X16) | s(bool,X16) = s(bool,sK6(sK1)) | s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X17)) ) | ($spl12 | $spl144)),
  inference(superposition,[],[f14176,f7351])).
fof(f12946,plain,(
  s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | $spl136),
  inference(cnf_transformation,[],[f12946_D])).
fof(f12946_D,plain,(
  s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f16300,plain,(
  ( ! [X20,X18] : (s(bool,X18) != s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,X20) = s(bool,sK6(sK1))) ) | $spl172),
  inference(cnf_transformation,[],[f16300_D])).
fof(f16300_D,plain,(
  ( ! [X20,X18] : (s(bool,X18) != s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,X20) = s(bool,sK6(sK1))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f16832,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X19) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X18)) ) | ($spl12 | $spl144)),
  inference(superposition,[],[f14176,f7351])).
fof(f16833,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X21) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X20) | s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) | ($spl12 | $spl144)),
  inference(superposition,[],[f14176,f7351])).
fof(f17549,plain,(
  ( ! [X56,X55] : (s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X56) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X55)) ) | ($spl12 | $spl137 | $spl174)),
  inference(subsumption_resolution,[],[f17530,f2548])).
fof(f17530,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,X56) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X56) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X55)) ) | ($spl12 | $spl137 | $spl174)),
  inference(superposition,[],[f16871,f14909])).
fof(f16871,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,X0) != s(bool,sK6(sK1))) ) | $spl174),
  inference(resolution,[],[f16844,f135])).
fof(f17541,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) = s(bool,X23) | s(bool,X23) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl174)),
  inference(subsumption_resolution,[],[f17515,f2548])).
fof(f17515,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s(bool,X21) = s(bool,X23) | s(bool,X23) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl174)),
  inference(superposition,[],[f16871,f7351])).
fof(f17539,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) = s(bool,X20) | s(bool,X20) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl174)),
  inference(subsumption_resolution,[],[f17514,f2548])).
fof(f17514,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s(bool,X18) = s(bool,X20) | s(bool,X20) != s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl174)),
  inference(superposition,[],[f16871,f7351])).
fof(f17314,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s(bool,sK5) != s(bool,X23) | s(bool,X21) != s(bool,X23) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl85 | $spl93 | $spl174)),
  inference(subsumption_resolution,[],[f17287,f14343])).
fof(f17287,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s(bool,sK5) != s(bool,X23) | s(bool,X21) != s(bool,X23) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl85 | $spl93 | $spl174)),
  inference(superposition,[],[f17190,f7351])).
fof(f17190,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,sK5) != s(bool,X4) | s(bool,X3) != s(bool,X4)) ) | ($spl85 | $spl93 | $spl174)),
  inference(subsumption_resolution,[],[f17168,f8215])).
fof(f17168,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | s(bool,sK5) != s(bool,X4) | p(s(bool,X3)) | p(s(bool,sK6(sK1)))) ) | ($spl85 | $spl93 | $spl174)),
  inference(superposition,[],[f16878,f135])).
fof(f16878,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,sK6(sK1)) | s(bool,sK5) != s(bool,X12)) ) | ($spl85 | $spl93 | $spl174)),
  inference(resolution,[],[f16844,f14343])).
fof(f17312,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s(bool,sK5) != s(bool,X20) | s(bool,X18) != s(bool,X20) | s(bool,sK5) = s(bool,X18)) ) | ($spl12 | $spl85 | $spl93 | $spl174)),
  inference(subsumption_resolution,[],[f17286,f14343])).
fof(f17286,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s(bool,sK5) != s(bool,X20) | s(bool,X18) != s(bool,X20) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl85 | $spl93 | $spl174)),
  inference(superposition,[],[f17190,f7351])).
fof(f17108,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) | s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) ) | $spl176),
  inference(cnf_transformation,[],[f17108_D])).
fof(f17108_D,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) | s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f14180,plain,(
  ( ! [X26] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26)))) ) | $spl146),
  inference(cnf_transformation,[],[f14180_D])).
fof(f14180_D,plain,(
  ( ! [X26] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f16987,plain,(
  ( ! [X2,X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | s('type.semanticPrimitives.eq_result',X3) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result',X3) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X3)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f1229,f11533])).
fof(f17109,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | s(bool,X0) = s(bool,sK6(sK1))) ) | $spl178),
  inference(cnf_transformation,[],[f17109_D])).
fof(f17109_D,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | s(bool,X0) = s(bool,sK6(sK1))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f16986,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) | s(bool,X0) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85)),
  inference(superposition,[],[f1229,f10781])).
fof(f16963,plain,(
  ( ! [X2,X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | s('type.semanticPrimitives.eq_result',X3) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result',X3) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X3)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl12 | $spl76 | $spl86)),
  inference(superposition,[],[f1229,f11533])).
fof(f16962,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X1)))) | s(bool,X0) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85)),
  inference(superposition,[],[f1229,f10781])).
fof(f16569,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X66) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X65)) ) | ($spl12 | $spl137 | $spl142)),
  inference(subsumption_resolution,[],[f16461,f2548])).
fof(f16461,plain,(
  ( ! [X66,X65] : (p(s(bool,X65)) | s(bool,X66) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,X66) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X65)) ) | ($spl12 | $spl137 | $spl142)),
  inference(superposition,[],[f14173,f14909])).
fof(f16301,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X23) | s(bool,X23) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl142)),
  inference(subsumption_resolution,[],[f16279,f2548])).
fof(f16279,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X23) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl142)),
  inference(superposition,[],[f14173,f7351])).
fof(f16299,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) != s(bool,X20) | s(bool,X20) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl142)),
  inference(subsumption_resolution,[],[f16278,f2548])).
fof(f16278,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X20) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl142)),
  inference(superposition,[],[f14173,f7351])).
fof(f15877,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X41) = s(bool,sK6(sK1))) ) | $spl170),
  inference(cnf_transformation,[],[f15877_D])).
fof(f15877_D,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X41) = s(bool,sK6(sK1))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f15876,plain,(
  ( ! [X41,X40] : (s(bool,X41) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X41) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X40)))) ) | ($spl12 | $spl76 | $spl85 | $spl86)),
  inference(subsumption_resolution,[],[f15857,f8217])).
fof(f15857,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | s(bool,X41) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X41) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X40)))) ) | ($spl12 | $spl76 | $spl85 | $spl86)),
  inference(superposition,[],[f10111,f11533])).
fof(f15597,plain,(
  ( ! [X41] : (s(bool,X41) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X41)) ) | $spl168),
  inference(cnf_transformation,[],[f15597_D])).
fof(f15597_D,plain,(
  ( ! [X41] : (s(bool,X41) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X41)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f15596,plain,(
  ( ! [X41,X40] : (s(bool,X41) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X41) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X40)))) ) | ($spl12 | $spl76 | $spl85 | $spl86 | $spl93)),
  inference(subsumption_resolution,[],[f15580,f8217])).
fof(f15580,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | s(bool,X41) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) = s(bool,X41) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X40)))) ) | ($spl12 | $spl76 | $spl85 | $spl86 | $spl93)),
  inference(superposition,[],[f15467,f11533])).
fof(f12178,plain,(
  ( ! [X139,X137] : (p(s(bool,X139)) | s(bool,sK5) = s(bool,sK6(X137)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X137) | s('type.semanticPrimitives.eq_result',X137) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X139)))) ) | $spl132),
  inference(cnf_transformation,[],[f12178_D])).
fof(f12178_D,plain,(
  ( ! [X139,X137] : (p(s(bool,X139)) | s(bool,sK5) = s(bool,sK6(X137)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X137) | s('type.semanticPrimitives.eq_result',X137) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X139)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f15553,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X32)) | s('type.semanticPrimitives.eq_result',X30) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X32))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X30) | s(bool,sK5) = s(bool,X31) | s(bool,sK5) = s(bool,sK6(X30))) ) | $spl12),
  inference(subsumption_resolution,[],[f15531,f2548])).
fof(f15531,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s('type.semanticPrimitives.eq_result',X30) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X32))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X30) | s(bool,sK5) = s(bool,X31) | s(bool,sK5) = s(bool,sK6(X30))) ) | $spl12),
  inference(superposition,[],[f1234,f7351])).
fof(f15552,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X29)) | s('type.semanticPrimitives.eq_result',X27) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X29))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X27) | s(bool,sK5) = s(bool,sK6(X27)) | s(bool,sK5) = s(bool,X28)) ) | $spl12),
  inference(subsumption_resolution,[],[f15530,f2548])).
fof(f15530,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s('type.semanticPrimitives.eq_result',X27) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X29))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X27) | s(bool,sK5) = s(bool,sK6(X27)) | s(bool,sK5) = s(bool,X28)) ) | $spl12),
  inference(superposition,[],[f1234,f7351])).
fof(f12169,plain,(
  ( ! [X33,X31] : (p(s(bool,sK6(X33))) | s(bool,sK5) = s(bool,X31) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X31)))) ) | $spl126),
  inference(cnf_transformation,[],[f12169_D])).
fof(f12169_D,plain,(
  ( ! [X33,X31] : (p(s(bool,sK6(X33))) | s(bool,sK5) = s(bool,X31) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X31)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f15541,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,sK6(X23))) | s('type.semanticPrimitives.eq_result',X23) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X23) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | $spl12),
  inference(subsumption_resolution,[],[f15504,f2548])).
fof(f15504,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | p(s(bool,sK6(X23))) | s('type.semanticPrimitives.eq_result',X23) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X23) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | $spl12),
  inference(superposition,[],[f1234,f7351])).
fof(f15540,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,sK6(X20))) | s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X20) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | $spl12),
  inference(subsumption_resolution,[],[f15503,f2548])).
fof(f15503,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | p(s(bool,sK6(X20))) | s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X20) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | $spl12),
  inference(superposition,[],[f1234,f7351])).
fof(f15240,plain,(
  ( ! [X35,X36,X34] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X34))),s(fun(bool,bool),X35),s(bool,X36)))) ) | $spl166),
  inference(cnf_transformation,[],[f15240_D])).
fof(f15240_D,plain,(
  ( ! [X35,X36,X34] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X34))),s(fun(bool,bool),X35),s(bool,X36)))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f15274,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,sK5) = s(bool,X77) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X74))),s(fun(bool,bool),X75),s(bool,X76)))) ) | ($spl12 | $spl30)),
  inference(duplicate_literal_removal,[],[f15273])).
fof(f15273,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,sK5) = s(bool,X77) | s(bool,sK5) = s(bool,X77) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X74))),s(fun(bool,bool),X75),s(bool,X76)))) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f15204,f4419])).
fof(f15204,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X77) = s(bool,happ(s(fun(bool,bool),X75),s(bool,X74))) | s(bool,sK5) = s(bool,X77) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X74))),s(fun(bool,bool),X75),s(bool,X76)))) ) | $spl12),
  inference(superposition,[],[f7351,f133])).
fof(f15272,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,sK5) = s(bool,X73) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X70))),s(fun(bool,bool),X71),s(bool,X72)))) ) | ($spl12 | $spl30)),
  inference(duplicate_literal_removal,[],[f15271])).
fof(f15271,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,sK5) = s(bool,X73) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X70))),s(fun(bool,bool),X71),s(bool,X72))) | s(bool,sK5) = s(bool,X73)) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f15203,f4419])).
fof(f15203,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,X73) = s(bool,happ(s(fun(bool,bool),X71),s(bool,X70))) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X70))),s(fun(bool,bool),X71),s(bool,X72))) | s(bool,sK5) = s(bool,X73)) ) | $spl12),
  inference(superposition,[],[f7351,f133])).
fof(f15242,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,sK5) = s(bool,X41) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X38))),s(fun(bool,bool),X39),s(bool,X40)))) ) | ($spl12 | $spl30)),
  inference(duplicate_literal_removal,[],[f15241])).
fof(f15241,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,sK5) = s(bool,X41) | s(bool,sK5) = s(bool,X41) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X38))),s(fun(bool,bool),X39),s(bool,X40)))) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f15183,f4419])).
fof(f15183,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X41) = s(bool,happ(s(fun(bool,bool),X39),s(bool,X38))) | s(bool,sK5) = s(bool,X41) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X38))),s(fun(bool,bool),X39),s(bool,X40)))) ) | $spl12),
  inference(superposition,[],[f133,f7351])).
fof(f15239,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,sK5) = s(bool,X37) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X34))),s(fun(bool,bool),X35),s(bool,X36)))) ) | ($spl12 | $spl30)),
  inference(duplicate_literal_removal,[],[f15238])).
fof(f15238,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,sK5) = s(bool,X37) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X34))),s(fun(bool,bool),X35),s(bool,X36))) | s(bool,sK5) = s(bool,X37)) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f15182,f4419])).
fof(f15182,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X37) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X34))),s(fun(bool,bool),X35),s(bool,X36))) | s(bool,sK5) = s(bool,X37)) ) | $spl12),
  inference(superposition,[],[f133,f7351])).
fof(f15106,plain,(
  $false | ($spl11 | $spl96)),
  inference(trivial_inequality_removal,[],[f15094])).
fof(f15094,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',sK1) | ($spl11 | $spl96)),
  inference(superposition,[],[f148,f9079])).
fof(f15051,plain,(
  ( ! [X45] : (s(bool,X45) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) != s(bool,X45)) ) | $spl164),
  inference(cnf_transformation,[],[f15051_D])).
fof(f15051_D,plain,(
  ( ! [X45] : (s(bool,X45) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) != s(bool,X45)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f15050,plain,(
  ( ! [X45,X44] : (s(bool,X45) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) != s(bool,X45) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X44)))) ) | ($spl12 | $spl76 | $spl85 | $spl86 | $spl93)),
  inference(subsumption_resolution,[],[f15032,f8217])).
fof(f15032,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | s(bool,X45) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) != s(bool,X45) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X44)))) ) | ($spl12 | $spl76 | $spl85 | $spl86 | $spl93)),
  inference(superposition,[],[f14531,f11533])).
fof(f14812,plain,(
  $false | ($spl40 | $spl85)),
  inference(resolution,[],[f1899,f8215])).
fof(f14801,plain,(
  ( ! [X70,X69] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X70))) | p(s(bool,X69))) ) | ($spl12 | $spl73 | $spl76 | $spl78 | $spl80 | $spl94 | $spl148)),
  inference(subsumption_resolution,[],[f14800,f11194])).
fof(f14800,plain,(
  ( ! [X70,X69] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X70))) | p(s(bool,X69)) | p(s(bool,X70))) ) | ($spl12 | $spl76 | $spl78 | $spl80 | $spl94 | $spl148)),
  inference(forward_demodulation,[],[f14720,f9082])).
fof(f14720,plain,(
  ( ! [X70,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X70))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))))) ) | ($spl78 | $spl80 | $spl94 | $spl148)),
  inference(superposition,[],[f10837,f14209])).
fof(f10837,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f10817,f9367])).
fof(f10817,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) ) | $spl80),
  inference(superposition,[],[f1234,f7479])).
fof(f14799,plain,(
  ( ! [X68,X67] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X68))) | p(s(bool,X67))) ) | ($spl12 | $spl73 | $spl76 | $spl78 | $spl80 | $spl94 | $spl148)),
  inference(subsumption_resolution,[],[f14798,f11194])).
fof(f14798,plain,(
  ( ! [X68,X67] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X68))) | p(s(bool,X67)) | p(s(bool,X68))) ) | ($spl12 | $spl76 | $spl78 | $spl80 | $spl94 | $spl148)),
  inference(forward_demodulation,[],[f14719,f9082])).
fof(f14719,plain,(
  ( ! [X68,X67] : (p(s(bool,X67)) | p(s(bool,X68)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X68))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))))) ) | ($spl78 | $spl80 | $spl94 | $spl148)),
  inference(superposition,[],[f10837,f14209])).
fof(f14792,plain,(
  ( ! [X34] : (s('type.semanticPrimitives.eq_result',X34) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X34) | p(s(bool,sK6(X34)))) ) | $spl162),
  inference(cnf_transformation,[],[f14792_D])).
fof(f14792_D,plain,(
  ( ! [X34] : (s('type.semanticPrimitives.eq_result',X34) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X34) | p(s(bool,sK6(X34)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f14791,plain,(
  ( ! [X33,X34] : (s('type.semanticPrimitives.eq_result',X34) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | p(s(bool,X33)) | p(s(bool,sK6(X34))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X34)) ) | ($spl12 | $spl76 | $spl148)),
  inference(forward_demodulation,[],[f14698,f9082])).
fof(f14698,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | p(s(bool,sK6(X34))) | s('type.semanticPrimitives.eq_result',X34) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X34)) ) | $spl148),
  inference(superposition,[],[f1234,f14209])).
fof(f14782,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X10)))) ) | ($spl12 | $spl73 | $spl76 | $spl148)),
  inference(subsumption_resolution,[],[f14781,f9217])).
fof(f14781,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X10))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) ) | ($spl12 | $spl73 | $spl76 | $spl148)),
  inference(subsumption_resolution,[],[f14686,f11194])).
fof(f14686,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X10))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) ) | $spl148),
  inference(superposition,[],[f1234,f14209])).
fof(f14508,plain,(
  ( ! [X72] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X72))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | ~p(s(bool,X72))) ) | $spl158),
  inference(cnf_transformation,[],[f14508_D])).
fof(f14508_D,plain,(
  ( ! [X72] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X72))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | ~p(s(bool,X72))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f14478,plain,(
  ( ! [X72,X71] : (~p(s(bool,X71)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X72))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | ~p(s(bool,X72))) ) | ($spl78 | $spl80 | $spl94 | $spl116)),
  inference(superposition,[],[f11197,f10329])).
fof(f11197,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X5))) ) | ($spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f11178,f9367])).
fof(f11178,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X6))) ) | $spl80),
  inference(superposition,[],[f1236,f7479])).
fof(f14507,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | $spl156),
  inference(cnf_transformation,[],[f14507_D])).
fof(f14507_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f14477,plain,(
  ( ! [X70] : (~p(s(bool,X70)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl12 | $spl73 | $spl76 | $spl116)),
  inference(superposition,[],[f11194,f10329])).
fof(f14505,plain,(
  ( ! [X67] : (p(s(bool,X67)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X67))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | $spl154),
  inference(cnf_transformation,[],[f14505_D])).
fof(f14505_D,plain,(
  ( ! [X67] : (p(s(bool,X67)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X67))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f14476,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | p(s(bool,X69)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X69))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl78 | $spl80 | $spl94 | $spl116)),
  inference(superposition,[],[f10837,f10329])).
fof(f14475,plain,(
  ( ! [X66,X67] : (p(s(bool,X66)) | p(s(bool,X67)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X67))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl78 | $spl80 | $spl94 | $spl116)),
  inference(superposition,[],[f10837,f10329])).
fof(f14473,plain,(
  ( ! [X64] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X64)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl12 | $spl80 | $spl86 | $spl116)),
  inference(superposition,[],[f10493,f10329])).
fof(f10493,plain,(
  ( ! [X11] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X11)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X11)))) ) | ($spl12 | $spl80 | $spl86)),
  inference(superposition,[],[f10134,f7479])).
fof(f14469,plain,(
  ( ! [X57] : (~p(s(bool,X57)) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl11 | $spl12 | $spl116)),
  inference(superposition,[],[f10056,f10329])).
fof(f14466,plain,(
  ( ! [X54] : (p(s(bool,X54)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl86 | $spl116)),
  inference(superposition,[],[f8217,f10329])).
fof(f14503,plain,(
  ( ! [X33] : (p(s(bool,sK6(X33))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | $spl150),
  inference(cnf_transformation,[],[f14503_D])).
fof(f14503_D,plain,(
  ( ! [X33] : (p(s(bool,sK6(X33))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f14455,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | p(s(bool,sK6(X33))) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33)) ) | $spl116),
  inference(superposition,[],[f1234,f10329])).
fof(f14499,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6)))) ) | ($spl3 | $spl86 | $spl116)),
  inference(subsumption_resolution,[],[f14498,f142])).
fof(f14498,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6)))) ) | ($spl86 | $spl116)),
  inference(subsumption_resolution,[],[f14441,f8217])).
fof(f14441,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X6))) ) | $spl116),
  inference(superposition,[],[f1236,f10329])).
fof(f14494,plain,(
  ( ! [X65] : (s(bool,sK5) = s(bool,X65) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | ($spl12 | $spl76 | $spl86 | $spl116)),
  inference(forward_demodulation,[],[f14414,f9082])).
fof(f14414,plain,(
  ( ! [X65] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s(bool,X65) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl86 | $spl116)),
  inference(superposition,[],[f11533,f10329])).
fof(f14406,plain,(
  ( ! [X53] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X53)))) ) | ($spl12 | $spl80 | $spl86 | $spl116)),
  inference(superposition,[],[f10493,f10329])).
fof(f14405,plain,(
  ( ! [X52] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s(bool,sK5) = s(bool,X52)) ) | ($spl12 | $spl86 | $spl116)),
  inference(superposition,[],[f10134,f10329])).
fof(f14289,plain,(
  $false | ($spl3 | $spl11 | $spl92)),
  inference(subsumption_resolution,[],[f14288,f142])).
fof(f14288,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl11 | $spl92)),
  inference(subsumption_resolution,[],[f14231,f148])).
fof(f14231,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | $spl92),
  inference(superposition,[],[f132,f8536])).
fof(f14277,plain,(
  $false | ($spl12 | $spl85 | $spl92)),
  inference(subsumption_resolution,[],[f14211,f172])).
fof(f14211,plain,(
  ~p(s(bool,sK5)) | ($spl85 | $spl92)),
  inference(backward_demodulation,[],[f8536,f8215])).
fof(f14276,plain,(
  $false | ($spl3 | $spl11 | $spl92)),
  inference(subsumption_resolution,[],[f14210,f148])).
fof(f14210,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl3 | $spl92)),
  inference(backward_demodulation,[],[f8536,f8084])).
fof(f14209,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | $spl148),
  inference(cnf_transformation,[],[f14209_D])).
fof(f14209_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f8535,plain,(
  s(bool,sK5) != s(bool,sK6(sK1)) | $spl93),
  inference(cnf_transformation,[],[f8535_D])).
fof(f8535_D,plain,(
  s(bool,sK5) != s(bool,sK6(sK1)) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f11037,plain,(
  ( ! [X0] : (s(bool,sK5) != s(bool,sK6(sK1)) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl85)),
  inference(forward_demodulation,[],[f10968,f9082])).
fof(f10968,plain,(
  ( ! [X0] : (s(bool,sK6(sK1)) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl85)),
  inference(equality_factoring,[],[f10781])).
fof(f8536,plain,(
  s(bool,sK5) = s(bool,sK6(sK1)) | $spl92),
  inference(cnf_transformation,[],[f8536_D])).
fof(f8536_D,plain,(
  s(bool,sK5) = s(bool,sK6(sK1)) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f11818,plain,(
  ( ! [X94,X95] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X94))) | s(bool,sK5) = s(bool,sK6(sK1)) | s(bool,X94) = s(bool,X95) | s(bool,sK5) = s(bool,X95)) ) | ($spl11 | $spl12 | $spl85)),
  inference(superposition,[],[f10114,f7351])).
fof(f10114,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(bool,X5) = s(bool,sK6(sK1))) ) | ($spl11 | $spl12 | $spl85)),
  inference(resolution,[],[f10059,f10056])).
fof(f11941,plain,(
  ( ! [X94,X95] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X94))) | s(bool,sK5) = s(bool,sK6(sK1)) | s(bool,X94) = s(bool,X95) | s(bool,sK5) = s(bool,X95)) ) | ($spl11 | $spl12 | $spl85)),
  inference(superposition,[],[f10114,f7351])).
fof(f12690,plain,(
  ( ! [X26] : (s(bool,sK5) = s(bool,sK6(sK1)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26)))) ) | ($spl12 | $spl73 | $spl76 | $spl85 | $spl86)),
  inference(forward_demodulation,[],[f12635,f9082])).
fof(f12635,plain,(
  ( ! [X26] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26))) | s(bool,sK6(sK1)) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26)))) ) | ($spl12 | $spl73 | $spl76 | $spl85 | $spl86)),
  inference(superposition,[],[f11023,f11533])).
fof(f14157,plain,(
  ( ! [X30,X31] : (s(bool,X31) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X30) | s(bool,X31) = s(bool,sK6(sK1))) ) | ($spl3 | $spl11 | $spl12 | $spl76 | $spl80 | $spl85 | $spl93)),
  inference(duplicate_literal_removal,[],[f14137])).
fof(f14137,plain,(
  ( ! [X30,X31] : (s(bool,X31) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X30) | s(bool,X31) = s(bool,sK6(sK1)) | s(bool,sK5) = s(bool,X30)) ) | ($spl3 | $spl11 | $spl12 | $spl76 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f10657,f12975])).
fof(f12975,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X5)) ) | ($spl3 | $spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(trivial_inequality_removal,[],[f12969])).
fof(f12969,plain,(
  ( ! [X5] : (s(bool,sK5) != s(bool,sK5) | s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X5)) ) | ($spl3 | $spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f12929,f7351])).
fof(f12929,plain,(
  s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ($spl3 | $spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(trivial_inequality_removal,[],[f12921])).
fof(f12921,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',sK1) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | ($spl3 | $spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f10527,f8084])).
fof(f10527,plain,(
  ( ! [X11] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X11))) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X11))))) ) | ($spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f10341,f7479])).
fof(f10341,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(bool,sK5) != s(bool,X5)) ) | ($spl11 | $spl12 | $spl85 | $spl93)),
  inference(resolution,[],[f10336,f10056])).
fof(f10336,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK5) != s(bool,X1)) ) | ($spl85 | $spl93)),
  inference(subsumption_resolution,[],[f10331,f8215])).
fof(f10331,plain,(
  ( ! [X1] : (s(bool,sK5) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,sK6(sK1)))) ) | $spl93),
  inference(superposition,[],[f8535,f135])).
fof(f10657,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | s(bool,sK5) = s(bool,X12) | s(bool,X11) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(resolution,[],[f10627,f10059])).
fof(f10627,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,X5) | s(bool,sK5) = s(bool,X5)) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(forward_demodulation,[],[f10626,f9082])).
fof(f10626,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ~p(s(bool,X4))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f10610,f172])).
fof(f10610,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,X5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ~p(s(bool,sK5)) | ~p(s(bool,X4))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(superposition,[],[f10601,f138])).
fof(f10601,plain,(
  ( ! [X17] : (s(bool,sK5) != s(bool,X17) | s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f10592,f172])).
fof(f10592,plain,(
  ( ! [X17] : (~p(s(bool,sK5)) | s(bool,X17) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s(bool,sK5) != s(bool,X17)) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(superposition,[],[f10338,f9082])).
fof(f10338,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,sK5) != s(bool,X0)) ) | ($spl85 | $spl93)),
  inference(resolution,[],[f10336,f138])).
fof(f14173,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,X2) | s(bool,X2) = s(bool,sK6(sK1))) ) | $spl142),
  inference(cnf_transformation,[],[f14173_D])).
fof(f14173_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,X2) | s(bool,X2) = s(bool,sK6(sK1))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f14174,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,sK5) = s(bool,X3) | s(bool,X5) = s(bool,sK6(sK1)) | p(s(bool,X4))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f14127,f2548])).
fof(f14127,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,sK5) = s(bool,X3) | s(bool,X5) = s(bool,sK6(sK1)) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(superposition,[],[f10657,f135])).
fof(f14172,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,sK5) = s(bool,X0) | s(bool,X2) = s(bool,sK6(sK1)) | p(s(bool,X1))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f14126,f2548])).
fof(f14126,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,sK5) = s(bool,X0) | s(bool,X2) = s(bool,sK6(sK1)) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl12 | $spl76 | $spl85 | $spl93)),
  inference(superposition,[],[f10657,f135])).
fof(f13615,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X23)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f13565,f2548])).
fof(f13565,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(superposition,[],[f10837,f7351])).
fof(f13614,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f13564,f2548])).
fof(f13564,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | p(s(bool,X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(superposition,[],[f10837,f7351])).
fof(f13595,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X23)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f13536,f2548])).
fof(f13536,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X21))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X22) | s(bool,sK5) = s(bool,X21)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(superposition,[],[f10837,f7351])).
fof(f13593,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(subsumption_resolution,[],[f13535,f2548])).
fof(f13535,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X19)) | p(s(bool,X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X18))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X18) | s(bool,sK5) = s(bool,X19)) ) | ($spl12 | $spl78 | $spl80 | $spl94)),
  inference(superposition,[],[f10837,f7351])).
fof(f13331,plain,(
  ( ! [X61,X60] : (s(bool,X61) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X61))) | s(bool,sK5) = s(bool,X60)) ) | ($spl3 | $spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f13208,f2548])).
fof(f13208,plain,(
  ( ! [X61,X60] : (p(s(bool,X60)) | s(bool,X61) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X61))) | s(bool,sK5) = s(bool,X60)) ) | ($spl3 | $spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f10084,f12975])).
fof(f12945,plain,(
  ( ! [X35] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35))) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))))) ) | ($spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f12916,f10341])).
fof(f12916,plain,(
  ( ! [X35] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35))) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) | s(bool,sK5) = s(bool,X35)) ) | ($spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f10527,f10061])).
fof(f12935,plain,(
  ( ! [X12] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X12)))) | s(bool,sK5) = s(bool,X12)) ) | $spl134),
  inference(cnf_transformation,[],[f12935_D])).
fof(f12935_D,plain,(
  ( ! [X12] : (s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X12)))) | s(bool,sK5) = s(bool,X12)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f12936,plain,(
  ( ! [X14,X15] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X15))) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X14)))) | s(bool,sK5) = s(bool,X14)) ) | ($spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f12901,f10341])).
fof(f12901,plain,(
  ( ! [X14,X15] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X15))) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X14)))) | s(bool,sK5) = s(bool,X15) | s(bool,sK5) = s(bool,X14)) ) | ($spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f10527,f7351])).
fof(f12934,plain,(
  ( ! [X12,X13] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X13))) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X12)))) | s(bool,sK5) = s(bool,X12)) ) | ($spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(subsumption_resolution,[],[f12900,f10341])).
fof(f12900,plain,(
  ( ! [X12,X13] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X13))) | s(bool,sK5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X12)))) | s(bool,sK5) = s(bool,X12) | s(bool,sK5) = s(bool,X13)) ) | ($spl11 | $spl12 | $spl80 | $spl85 | $spl93)),
  inference(superposition,[],[f10527,f7351])).
fof(f12189,plain,(
  ( ! [X138,X139,X137] : (p(s(bool,X139)) | s('type.semanticPrimitives.eq_result',X137) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X139))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X137) | s(bool,sK5) = s(bool,X138) | s(bool,sK5) = s(bool,sK6(X137))) ) | $spl12),
  inference(subsumption_resolution,[],[f11765,f2548])).
fof(f11765,plain,(
  ( ! [X138,X139,X137] : (p(s(bool,X138)) | p(s(bool,X139)) | s('type.semanticPrimitives.eq_result',X137) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X139))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X137) | s(bool,sK5) = s(bool,X138) | s(bool,sK5) = s(bool,sK6(X137))) ) | $spl12),
  inference(superposition,[],[f1234,f7351])).
fof(f12176,plain,(
  ( ! [X135] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X135) | s(bool,sK5) = s(bool,sK6(X135)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X135)) ) | $spl130),
  inference(cnf_transformation,[],[f12176_D])).
fof(f12176_D,plain,(
  ( ! [X135] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X135) | s(bool,sK5) = s(bool,sK6(X135)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X135)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f12188,plain,(
  ( ! [X136,X135] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X135) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X135) | s(bool,sK5) = s(bool,X136) | s(bool,sK5) = s(bool,sK6(X135))) ) | ($spl3 | $spl12 | $spl85)),
  inference(subsumption_resolution,[],[f11764,f2548])).
fof(f11764,plain,(
  ( ! [X136,X135] : (p(s(bool,X136)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X135) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X135) | s(bool,sK5) = s(bool,X136) | s(bool,sK5) = s(bool,sK6(X135))) ) | ($spl3 | $spl12 | $spl85)),
  inference(superposition,[],[f10828,f7351])).
fof(f10828,plain,(
  ( ! [X19] : (p(s(bool,sK6(X19))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X19) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X19)) ) | ($spl3 | $spl85)),
  inference(forward_demodulation,[],[f10794,f8084])).
fof(f10794,plain,(
  ( ! [X19] : (p(s(bool,sK6(X19))) | s('type.semanticPrimitives.eq_result',X19) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X19)) ) | $spl85),
  inference(resolution,[],[f1234,f8215])).
fof(f12185,plain,(
  ( ! [X66,X64,X65] : (s(bool,X64) = s(bool,X66) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X66))) | s(bool,sK5) = s(bool,X65) | s(bool,sK5) = s(bool,X64)) ) | ($spl11 | $spl12)),
  inference(subsumption_resolution,[],[f11732,f2548])).
fof(f11732,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,X65)) | s(bool,X64) = s(bool,X66) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X66))) | s(bool,sK5) = s(bool,X65) | s(bool,sK5) = s(bool,X64)) ) | ($spl11 | $spl12)),
  inference(superposition,[],[f10084,f7351])).
fof(f12183,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,sK6(X33))) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X31))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s(bool,sK5) = s(bool,X32) | s(bool,sK5) = s(bool,X31)) ) | $spl12),
  inference(subsumption_resolution,[],[f11719,f2548])).
fof(f11719,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | p(s(bool,sK6(X33))) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X31))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s(bool,sK5) = s(bool,X32) | s(bool,sK5) = s(bool,X31)) ) | $spl12),
  inference(superposition,[],[f1234,f7351])).
fof(f12177,plain,(
  ( ! [X138,X139,X137] : (p(s(bool,X139)) | s('type.semanticPrimitives.eq_result',X137) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X139))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X137) | s(bool,sK5) = s(bool,sK6(X137)) | s(bool,sK5) = s(bool,X138)) ) | $spl12),
  inference(subsumption_resolution,[],[f11690,f2548])).
fof(f11690,plain,(
  ( ! [X138,X139,X137] : (p(s(bool,X138)) | p(s(bool,X139)) | s('type.semanticPrimitives.eq_result',X137) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X139))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X137) | s(bool,sK5) = s(bool,sK6(X137)) | s(bool,sK5) = s(bool,X138)) ) | $spl12),
  inference(superposition,[],[f1234,f7351])).
fof(f12175,plain,(
  ( ! [X136,X135] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X135) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X135) | s(bool,sK5) = s(bool,sK6(X135)) | s(bool,sK5) = s(bool,X136)) ) | ($spl3 | $spl12 | $spl85)),
  inference(subsumption_resolution,[],[f11689,f2548])).
fof(f11689,plain,(
  ( ! [X136,X135] : (p(s(bool,X136)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X135) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X135) | s(bool,sK5) = s(bool,sK6(X135)) | s(bool,sK5) = s(bool,X136)) ) | ($spl3 | $spl12 | $spl85)),
  inference(superposition,[],[f10828,f7351])).
fof(f12171,plain,(
  ( ! [X66,X64,X65] : (s(bool,X64) = s(bool,X66) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X66))) | s(bool,sK5) = s(bool,X64) | s(bool,sK5) = s(bool,X65)) ) | ($spl11 | $spl12)),
  inference(subsumption_resolution,[],[f11657,f2548])).
fof(f11657,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,X65)) | s(bool,X64) = s(bool,X66) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X66))) | s(bool,sK5) = s(bool,X64) | s(bool,sK5) = s(bool,X65)) ) | ($spl11 | $spl12)),
  inference(superposition,[],[f10084,f7351])).
fof(f12168,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,sK6(X33))) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X31))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s(bool,sK5) = s(bool,X31) | s(bool,sK5) = s(bool,X32)) ) | $spl12),
  inference(subsumption_resolution,[],[f11644,f2548])).
fof(f11644,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | p(s(bool,sK6(X33))) | s('type.semanticPrimitives.eq_result',X33) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X31))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X33) | s(bool,sK5) = s(bool,X31) | s(bool,sK5) = s(bool,X32)) ) | $spl12),
  inference(superposition,[],[f1234,f7351])).
fof(f11532,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27)))) ) | $spl124),
  inference(cnf_transformation,[],[f11532_D])).
fof(f11532_D,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f11531,plain,(
  ( ! [X26,X27] : (s(bool,X27) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) | s(bool,X26) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85 | $spl86)),
  inference(subsumption_resolution,[],[f11516,f10059])).
fof(f11516,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | s(bool,X27) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) | s(bool,X26) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85 | $spl86)),
  inference(superposition,[],[f10132,f10781])).
fof(f11369,plain,(
  ( ! [X6,X7] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X6),s(bool,X7)))) ) | $spl122),
  inference(cnf_transformation,[],[f11369_D])).
fof(f11369_D,plain,(
  ( ! [X6,X7] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X6),s(bool,X7)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f11295,plain,(
  ( ! [X6,X7,X5] : (s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X6),s(bool,X7))) | s(bool,sK5) = s(bool,X5)) ) | ($spl12 | $spl30 | $spl86)),
  inference(superposition,[],[f4529,f10134])).
fof(f11191,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,sK6(sK1))) ) | ($spl12 | $spl73 | $spl76 | $spl85)),
  inference(subsumption_resolution,[],[f11190,f9217])).
fof(f11190,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85)),
  inference(subsumption_resolution,[],[f11175,f10059])).
fof(f11175,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,sK6(sK1))) ) | ($spl12 | $spl76 | $spl85)),
  inference(superposition,[],[f1236,f10781])).
fof(f10833,plain,(
  ( ! [X21] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X21) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X21) | p(s(bool,sK6(X21)))) ) | $spl118),
  inference(cnf_transformation,[],[f10833_D])).
fof(f10833_D,plain,(
  ( ! [X21] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X21) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X21) | p(s(bool,sK6(X21)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f10832,plain,(
  ( ! [X21,X20] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result',X21) | p(s(bool,sK6(X21))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X21) | s(bool,sK5) = s(bool,X20)) ) | ($spl3 | $spl12 | $spl85)),
  inference(forward_demodulation,[],[f10831,f8084])).
fof(f10831,plain,(
  ( ! [X21,X20] : (p(s(bool,sK6(X21))) | s('type.semanticPrimitives.eq_result',X21) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X21) | s(bool,sK5) = s(bool,X20)) ) | ($spl12 | $spl85)),
  inference(subsumption_resolution,[],[f10812,f2548])).
fof(f10812,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | p(s(bool,sK6(X21))) | s('type.semanticPrimitives.eq_result',X21) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X21) | s(bool,sK5) = s(bool,X20)) ) | ($spl12 | $spl85)),
  inference(superposition,[],[f1234,f10061])).
fof(f10313,plain,(
  ( ! [X0] : (s(bool,sK5) != s(bool,sK6(sK1)) | s(bool,X0) = s(bool,sK6(sK1))) ) | ($spl12 | $spl85)),
  inference(equality_factoring,[],[f10061])).
fof(f10271,plain,(
  ( ! [X16] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK6(sK1)))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))) | s(bool,sK5) = s(bool,X16)) ) | ($spl6 | $spl12 | $spl85)),
  inference(superposition,[],[f145,f10061])).
fof(f10244,plain,(
  ( ! [X2,X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | p(s(bool,X3))) ) | ($spl6 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f10235,f10056])).
fof(f10243,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X1))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X1))) ) | ($spl6 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f10234,f10056])).
fof(f10185,plain,(
  $false | ($spl11 | $spl86 | $spl88)),
  inference(trivial_inequality_removal,[],[f10177])).
fof(f10177,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result',sK1) | ($spl11 | $spl86 | $spl88)),
  inference(superposition,[],[f148,f10162])).
fof(f10162,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl86 | $spl88)),
  inference(subsumption_resolution,[],[f8218,f8217])).
fof(f10161,plain,(
  $false | ($spl5 | $spl106)),
  inference(subsumption_resolution,[],[f10156,f144])).
fof(f10156,plain,(
  s(sK0,sK2) = s(sK0,sK3) | $spl106),
  inference(superposition,[],[f131,f10011])).
fof(f10160,plain,(
  $false | ($spl5 | $spl106)),
  inference(subsumption_resolution,[],[f10155,f144])).
fof(f10155,plain,(
  s(sK0,sK2) = s(sK0,sK3) | $spl106),
  inference(superposition,[],[f10011,f131])).
fof(f9926,plain,(
  ( ! [X4,X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl6),
  inference(superposition,[],[f145,f138])).
fof(f9927,plain,(
  ( ! [X6,X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl6),
  inference(superposition,[],[f145,f138])).
fof(f10055,plain,(
  $false | ($spl85 | $spl86 | $spl98)),
  inference(subsumption_resolution,[],[f8215,f10025])).
fof(f10025,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl86 | $spl98)),
  inference(subsumption_resolution,[],[f8217,f9938])).
fof(f10011,plain,(
  s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK3))) | $spl106),
  inference(cnf_transformation,[],[f10011_D])).
fof(f10011_D,plain,(
  s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK3))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f9912,plain,(
  s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK3))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | $spl0),
  inference(superposition,[],[f139,f132])).
fof(f10022,plain,(
  $false | ($spl3 | $spl11 | $spl12 | $spl84)),
  inference(subsumption_resolution,[],[f8601,f148])).
fof(f8601,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl3 | $spl12 | $spl84)),
  inference(backward_demodulation,[],[f8591,f8084])).
fof(f8591,plain,(
  s(bool,sK5) = s(bool,sK6(sK1)) | ($spl12 | $spl84)),
  inference(resolution,[],[f8216,f2548])).
fof(f149,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | $spl10),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f9078,plain,(
  ( ! [X6] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | p(s(bool,X6))) ) | ($spl12 | $spl84 | $spl90)),
  inference(forward_demodulation,[],[f9077,f8591])).
fof(f9077,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | $spl90),
  inference(forward_demodulation,[],[f9014,f131])).
fof(f9014,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X5),s(bool,X6)))) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1))))) ) | $spl90),
  inference(superposition,[],[f8511,f132])).
fof(f9919,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK3))) | ($spl0 | $spl12 | $spl84)),
  inference(forward_demodulation,[],[f9912,f8591])).
fof(f10009,plain,(
  ( ! [X6,X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X7))) ) | ($spl6 | $spl98)),
  inference(subsumption_resolution,[],[f9927,f9938])).
fof(f10007,plain,(
  ( ! [X4,X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | ~p(s(bool,X5))) ) | ($spl6 | $spl98)),
  inference(subsumption_resolution,[],[f9926,f9938])).
fof(f10006,plain,(
  $false | ($spl0 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f9913,f146])).
fof(f9913,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | ($spl0 | $spl10)),
  inference(superposition,[],[f139,f8090])).
fof(f8090,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(bool,X0),X1),s(bool,sK5))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X0),X1),s(X0,X2)))) ) | $spl10),
  inference(superposition,[],[f133,f149])).
fof(f9976,plain,(
  ( ! [X4,X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl6 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f9926,f8096])).
fof(f8096,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f8087,f172])).
fof(f8087,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK5)) | ~p(s(bool,X2))) ) | $spl10),
  inference(superposition,[],[f149,f138])).
fof(f9974,plain,(
  ( ! [X6,X7] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X7))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl6 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f9927,f8096])).
fof(f9966,plain,(
  $false | ($spl6 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f9937,f149])).
fof(f9937,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl6 | $spl8)),
  inference(trivial_inequality_removal,[],[f9936])).
fof(f9936,plain,(
  s(sK0,sK2) != s(sK0,sK2) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl6 | $spl8)),
  inference(superposition,[],[f145,f147])).
fof(f9965,plain,(
  $false | ($spl6 | $spl8 | $spl10 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f9964,f149])).
fof(f9964,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl6 | $spl8 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f9963,f4419])).
fof(f9963,plain,(
  ( ! [X21] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,happ(s(fun(bool,bool),X21),s(bool,sK5)))))) ) | ($spl6 | $spl8 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f9962,f8090])).
fof(f9962,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X21),s(bool,X22)))))) ) | ($spl6 | $spl8 | $spl10 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f9961,f147])).
fof(f9961,plain,(
  ( ! [X21,X22] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X21),s(bool,X22)))))) ) | ($spl6 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f9935,f4419])).
fof(f9935,plain,(
  ( ! [X21,X22] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,happ(s(fun(bool,bool),X21),s(bool,sK5))))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X21),s(bool,X22)))))) ) | ($spl6 | $spl10)),
  inference(superposition,[],[f145,f8090])).
fof(f9959,plain,(
  $false | ($spl6 | $spl8 | $spl10 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f9958,f149])).
fof(f9958,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl6 | $spl8 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f9957,f4419])).
fof(f9957,plain,(
  ( ! [X17,X16] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,happ(s(fun(bool,bool),X17),s(bool,X16)))))) ) | ($spl6 | $spl8 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f9956,f133])).
fof(f9956,plain,(
  ( ! [X17,X18,X16] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))),s(fun(bool,bool),X17),s(bool,X18)))))) ) | ($spl6 | $spl8 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f9955,f147])).
fof(f9955,plain,(
  ( ! [X17,X18,X16] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))),s(fun(bool,bool),X17),s(bool,X18)))))) ) | ($spl6 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f9933,f4419])).
fof(f9933,plain,(
  ( ! [X17,X18,X16] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,happ(s(fun(bool,bool),X17),s(bool,X16))))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))),s(fun(bool,bool),X17),s(bool,X18)))))) ) | $spl6),
  inference(superposition,[],[f145,f133])).
fof(f9952,plain,(
  $false | ($spl6 | $spl8 | $spl10 | $spl12 | $spl76)),
  inference(subsumption_resolution,[],[f9951,f149])).
fof(f9951,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl6 | $spl8 | $spl12 | $spl76)),
  inference(forward_demodulation,[],[f9950,f9082])).
fof(f9950,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | ($spl6 | $spl8 | $spl12 | $spl76)),
  inference(subsumption_resolution,[],[f9930,f147])).
fof(f9930,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | ($spl6 | $spl12 | $spl76)),
  inference(superposition,[],[f145,f9082])).
fof(f9949,plain,(
  $false | ($spl6 | $spl8 | $spl10 | $spl12 | $spl84)),
  inference(subsumption_resolution,[],[f9948,f149])).
fof(f9948,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl6 | $spl8 | $spl12 | $spl84)),
  inference(forward_demodulation,[],[f9947,f8591])).
fof(f9947,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | ($spl6 | $spl8 | $spl12 | $spl84)),
  inference(subsumption_resolution,[],[f9929,f147])).
fof(f9929,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(sK1)))) | ($spl6 | $spl12 | $spl84)),
  inference(superposition,[],[f145,f8591])).
fof(f9946,plain,(
  $false | ($spl6 | $spl8 | $spl10 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f9945,f149])).
fof(f9945,plain,(
  s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl6 | $spl8 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f9944,f4419])).
fof(f9944,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))))) ) | ($spl6 | $spl8 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f9928,f147])).
fof(f9928,plain,(
  ( ! [X8,X9] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))))) ) | ($spl6 | $spl12 | $spl30)),
  inference(superposition,[],[f145,f4419])).
fof(f9943,plain,(
  ( ! [X6,X7] : (~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl6 | $spl8 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f9942,f8096])).
fof(f9942,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl6 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f9927,f2753])).
fof(f2753,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3)))) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f2747,f172])).
fof(f2747,plain,(
  ( ! [X3] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | ~p(s(bool,sK5)) | ~p(s(bool,X3))) ) | $spl8),
  inference(superposition,[],[f147,f138])).
fof(f9941,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl6 | $spl8 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f9940,f8096])).
fof(f9940,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl6 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f9926,f2753])).
fof(f9925,plain,(
  ( ! [X2,X3] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X3))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f145,f135])).
fof(f9924,plain,(
  ( ! [X0,X1] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X1))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl6),
  inference(superposition,[],[f145,f135])).
fof(f9911,plain,(
  $false | ($spl1 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f9863,f147])).
fof(f9863,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | ($spl1 | $spl10)),
  inference(superposition,[],[f140,f8090])).
fof(f9608,plain,(
  $false | ($spl1 | $spl8 | $spl10 | $spl12 | $spl78 | $spl82 | $spl94)),
  inference(trivial_inequality_removal,[],[f9606])).
fof(f9606,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ($spl1 | $spl8 | $spl10 | $spl12 | $spl78 | $spl82 | $spl94)),
  inference(superposition,[],[f9551,f147])).
fof(f9551,plain,(
  ( ! [X124] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X124)))) ) | ($spl1 | $spl10 | $spl12 | $spl78 | $spl82 | $spl94)),
  inference(superposition,[],[f140,f9386])).
fof(f9386,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl10 | $spl12 | $spl78 | $spl82 | $spl94)),
  inference(backward_demodulation,[],[f9384,f133])).
fof(f9384,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) ) | ($spl10 | $spl12 | $spl78 | $spl82 | $spl94)),
  inference(forward_demodulation,[],[f9383,f149])).
fof(f9383,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) ) | ($spl12 | $spl78 | $spl82 | $spl94)),
  inference(forward_demodulation,[],[f9382,f9199])).
fof(f9199,plain,(
  ( ! [X4] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X4))))) ) | ($spl12 | $spl82)),
  inference(resolution,[],[f7480,f2548])).
fof(f9382,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1))))))) ) | ($spl78 | $spl94)),
  inference(trivial_inequality_removal,[],[f9381])).
fof(f9381,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1))))))) ) | ($spl78 | $spl94)),
  inference(superposition,[],[f9367,f132])).
fof(f8800,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | ($spl3 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f8785,f8096])).
fof(f8785,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | ($spl3 | $spl10)),
  inference(superposition,[],[f8156,f135])).
fof(f8156,plain,(
  s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f8155,f149])).
fof(f8155,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f8146,f142])).
fof(f8146,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | $spl10),
  inference(superposition,[],[f149,f132])).
fof(f8802,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | ($spl3 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f8786,f8096])).
fof(f8786,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | p(s(bool,X1))) ) | ($spl3 | $spl10)),
  inference(superposition,[],[f8156,f135])).
fof(f8787,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | ~p(s(bool,X2))) ) | ($spl3 | $spl10)),
  inference(superposition,[],[f8156,f138])).
fof(f8788,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | ($spl3 | $spl10)),
  inference(superposition,[],[f8156,f138])).
fof(f9042,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | $spl73),
  inference(superposition,[],[f5391,f135])).
fof(f9043,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | p(s(bool,X1))) ) | $spl73),
  inference(superposition,[],[f5391,f135])).
fof(f9044,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | ~p(s(bool,X2))) ) | $spl73),
  inference(superposition,[],[f5391,f138])).
fof(f9045,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | $spl73),
  inference(superposition,[],[f5391,f138])).
fof(f9074,plain,(
  $false | ($spl1 | $spl3 | $spl8 | $spl10 | $spl12 | $spl77)),
  inference(trivial_inequality_removal,[],[f9072])).
fof(f9072,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ($spl1 | $spl3 | $spl8 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f8920,f147])).
fof(f8920,plain,(
  ( ! [X107] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X107)))) ) | ($spl1 | $spl3 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f140,f8808])).
fof(f8808,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(backward_demodulation,[],[f8801,f133])).
fof(f8801,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(subsumption_resolution,[],[f8800,f6373])).
fof(f8511,plain,(
  ( ! [X12,X11] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X11),s(bool,X12))))) ) | $spl90),
  inference(cnf_transformation,[],[f8511_D])).
fof(f8511_D,plain,(
  ( ! [X12,X11] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X11),s(bool,X12))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8955,plain,(
  ( ! [X26,X24,X23] : (s(bool,sK5) = s(bool,X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X23),s(bool,X24))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(subsumption_resolution,[],[f8954,f2548])).
fof(f8954,plain,(
  ( ! [X26,X24,X23] : (s(bool,sK5) = s(bool,X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X23),s(bool,X24)))) | p(s(bool,X26))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(forward_demodulation,[],[f8901,f4419])).
fof(f8901,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X25))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X23),s(bool,X24)))) | p(s(bool,X26))) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f135,f8808])).
fof(f8953,plain,(
  ( ! [X19,X22,X20] : (s(bool,sK5) = s(bool,X22) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X19),s(bool,X20))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(subsumption_resolution,[],[f8952,f2548])).
fof(f8952,plain,(
  ( ! [X19,X22,X20] : (s(bool,sK5) = s(bool,X22) | p(s(bool,X22)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X19),s(bool,X20))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(forward_demodulation,[],[f8900,f4419])).
fof(f8900,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X21))) | p(s(bool,X22)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X19),s(bool,X20))))) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f135,f8808])).
fof(f8949,plain,(
  ( ! [X62,X60,X63] : (s(bool,sK5) = s(bool,X63) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X60),s(bool,X62))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(duplicate_literal_removal,[],[f8948])).
fof(f8948,plain,(
  ( ! [X62,X60,X63] : (s(bool,sK5) = s(bool,X63) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X60),s(bool,X62)))) | s(bool,sK5) = s(bool,X63)) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(forward_demodulation,[],[f8887,f4419])).
fof(f8887,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X60),s(bool,X62)))) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61))) | s(bool,sK5) = s(bool,X63)) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f4406,f8808])).
fof(f8940,plain,(
  ( ! [X37,X36,X34] : (s(bool,sK5) = s(bool,X37) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X34),s(bool,X36))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(subsumption_resolution,[],[f8939,f2548])).
fof(f8939,plain,(
  ( ! [X37,X36,X34] : (s(bool,sK5) = s(bool,X37) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X34),s(bool,X36)))) | p(s(bool,X37))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(forward_demodulation,[],[f8880,f4419])).
fof(f8880,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X34),s(bool,X36)))) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X34),s(bool,X35))) | p(s(bool,X37))) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f135,f8808])).
fof(f8938,plain,(
  ( ! [X30,X33,X32] : (s(bool,sK5) = s(bool,X33) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X30),s(bool,X32))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(subsumption_resolution,[],[f8937,f2548])).
fof(f8937,plain,(
  ( ! [X30,X33,X32] : (s(bool,sK5) = s(bool,X33) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X30),s(bool,X32)))) | p(s(bool,X33))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(forward_demodulation,[],[f8879,f4419])).
fof(f8879,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X30),s(bool,X32)))) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X30),s(bool,X31))) | p(s(bool,X33))) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f135,f8808])).
fof(f8931,plain,(
  ( ! [X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(subsumption_resolution,[],[f8930,f2548])).
fof(f8930,plain,(
  ( ! [X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X15),s(bool,X16)))) | p(s(bool,X17))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(forward_demodulation,[],[f8869,f4419])).
fof(f8869,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X18))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X15),s(bool,X16)))) | p(s(bool,X17))) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f8808,f135])).
fof(f8929,plain,(
  ( ! [X12,X13,X11] : (s(bool,sK5) = s(bool,X13) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(subsumption_resolution,[],[f8928,f2548])).
fof(f8928,plain,(
  ( ! [X12,X13,X11] : (s(bool,sK5) = s(bool,X13) | p(s(bool,X13)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl3 | $spl10 | $spl12 | $spl30 | $spl77)),
  inference(forward_demodulation,[],[f8868,f4419])).
fof(f8868,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X13) = s(bool,happ(s(fun(bool,bool),X11),s(bool,X14))) | p(s(bool,X13)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl3 | $spl10 | $spl12 | $spl77)),
  inference(superposition,[],[f8808,f135])).
fof(f8587,plain,(
  $false | ($spl1 | $spl8 | $spl10 | $spl12 | $spl86)),
  inference(trivial_inequality_removal,[],[f8585])).
fof(f8585,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ($spl1 | $spl8 | $spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f8501,f147])).
fof(f8501,plain,(
  ( ! [X113] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X113)))) ) | ($spl1 | $spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f140,f8350])).
fof(f8350,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl10 | $spl12 | $spl86)),
  inference(backward_demodulation,[],[f8349,f133])).
fof(f8349,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | ($spl10 | $spl12 | $spl86)),
  inference(subsumption_resolution,[],[f8217,f8096])).
fof(f8541,plain,(
  ( ! [X26,X24,X23] : (s(bool,sK5) = s(bool,X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X23),s(bool,X24))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(subsumption_resolution,[],[f8540,f2548])).
fof(f8540,plain,(
  ( ! [X26,X24,X23] : (s(bool,sK5) = s(bool,X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X23),s(bool,X24)))) | p(s(bool,X26))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(forward_demodulation,[],[f8480,f4419])).
fof(f8480,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X25))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X23),s(bool,X24)))) | p(s(bool,X26))) ) | ($spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f135,f8350])).
fof(f8539,plain,(
  ( ! [X19,X22,X20] : (s(bool,sK5) = s(bool,X22) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X19),s(bool,X20))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(subsumption_resolution,[],[f8538,f2548])).
fof(f8538,plain,(
  ( ! [X19,X22,X20] : (s(bool,sK5) = s(bool,X22) | p(s(bool,X22)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X19),s(bool,X20))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(forward_demodulation,[],[f8479,f4419])).
fof(f8479,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X21))) | p(s(bool,X22)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X19),s(bool,X20))))) ) | ($spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f135,f8350])).
fof(f8534,plain,(
  ( ! [X71,X73] : (s(bool,sK5) = s(bool,sK6(sK1)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X71),s(bool,X73))))) ) | ($spl10 | $spl12 | $spl30 | $spl85 | $spl86)),
  inference(forward_demodulation,[],[f8467,f4419])).
fof(f8467,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X71),s(bool,X73)))) | s(bool,sK6(sK1)) = s(bool,happ(s(fun(bool,bool),X71),s(bool,X72)))) ) | ($spl10 | $spl12 | $spl85 | $spl86)),
  inference(superposition,[],[f8223,f8350])).
fof(f8223,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,sK6(sK1))) ) | $spl85),
  inference(resolution,[],[f8215,f135])).
fof(f8531,plain,(
  ( ! [X62,X60,X63] : (s(bool,sK5) = s(bool,X63) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X60),s(bool,X62))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(duplicate_literal_removal,[],[f8530])).
fof(f8530,plain,(
  ( ! [X62,X60,X63] : (s(bool,sK5) = s(bool,X63) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X60),s(bool,X62)))) | s(bool,sK5) = s(bool,X63)) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(forward_demodulation,[],[f8464,f4419])).
fof(f8464,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X60),s(bool,X62)))) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61))) | s(bool,sK5) = s(bool,X63)) ) | ($spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f4406,f8350])).
fof(f8522,plain,(
  ( ! [X37,X36,X34] : (s(bool,sK5) = s(bool,X37) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X34),s(bool,X36))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(subsumption_resolution,[],[f8521,f2548])).
fof(f8521,plain,(
  ( ! [X37,X36,X34] : (s(bool,sK5) = s(bool,X37) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X34),s(bool,X36)))) | p(s(bool,X37))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(forward_demodulation,[],[f8457,f4419])).
fof(f8457,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X34),s(bool,X36)))) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X34),s(bool,X35))) | p(s(bool,X37))) ) | ($spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f135,f8350])).
fof(f8520,plain,(
  ( ! [X30,X33,X32] : (s(bool,sK5) = s(bool,X33) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X30),s(bool,X32))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(subsumption_resolution,[],[f8519,f2548])).
fof(f8519,plain,(
  ( ! [X30,X33,X32] : (s(bool,sK5) = s(bool,X33) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X30),s(bool,X32)))) | p(s(bool,X33))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(forward_demodulation,[],[f8456,f4419])).
fof(f8456,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X30),s(bool,X32)))) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X30),s(bool,X31))) | p(s(bool,X33))) ) | ($spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f135,f8350])).
fof(f8513,plain,(
  ( ! [X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(subsumption_resolution,[],[f8512,f2548])).
fof(f8512,plain,(
  ( ! [X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X15),s(bool,X16)))) | p(s(bool,X17))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(forward_demodulation,[],[f8445,f4419])).
fof(f8445,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X18))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X15),s(bool,X16)))) | p(s(bool,X17))) ) | ($spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f8350,f135])).
fof(f8510,plain,(
  ( ! [X12,X13,X11] : (s(bool,sK5) = s(bool,X13) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(subsumption_resolution,[],[f8509,f2548])).
fof(f8509,plain,(
  ( ! [X12,X13,X11] : (s(bool,sK5) = s(bool,X13) | p(s(bool,X13)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl10 | $spl12 | $spl30 | $spl86)),
  inference(forward_demodulation,[],[f8444,f4419])).
fof(f8444,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X13) = s(bool,happ(s(fun(bool,bool),X11),s(bool,X14))) | p(s(bool,X13)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl10 | $spl12 | $spl86)),
  inference(superposition,[],[f8350,f135])).
fof(f8207,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK6(sK1)))) ) | $spl3),
  inference(superposition,[],[f8084,f138])).
fof(f8206,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK6(sK1))) | ~p(s(bool,X2))) ) | $spl3),
  inference(superposition,[],[f8084,f138])).
fof(f8205,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6(sK1))) | p(s(bool,X1))) ) | $spl3),
  inference(superposition,[],[f8084,f135])).
fof(f8204,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK6(sK1)))) ) | $spl3),
  inference(superposition,[],[f8084,f135])).
fof(f8184,plain,(
  $false | ($spl3 | $spl10 | $spl72)),
  inference(subsumption_resolution,[],[f8174,f142])).
fof(f8174,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl10 | $spl72)),
  inference(backward_demodulation,[],[f6368,f149])).
fof(f6368,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | $spl72),
  inference(duplicate_literal_removal,[],[f6364])).
fof(f6364,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | $spl72),
  inference(superposition,[],[f132,f5392])).
fof(f8141,plain,(
  $false | ($spl12 | $spl77 | $spl80)),
  inference(subsumption_resolution,[],[f8140,f172])).
fof(f8140,plain,(
  ~p(s(bool,sK5)) | ($spl77 | $spl80)),
  inference(forward_demodulation,[],[f6373,f7479])).
fof(f8136,plain,(
  $false | ($spl3 | $spl10 | $spl12 | $spl72 | $spl76)),
  inference(subsumption_resolution,[],[f8126,f142])).
fof(f8126,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl10 | $spl12 | $spl72 | $spl76)),
  inference(backward_demodulation,[],[f6572,f149])).
fof(f6572,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl12 | $spl72 | $spl76)),
  inference(forward_demodulation,[],[f5392,f6487])).
fof(f6487,plain,(
  s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))) | ($spl12 | $spl76)),
  inference(resolution,[],[f6374,f2548])).
fof(f8110,plain,(
  $false | ($spl3 | $spl10 | $spl12 | $spl16)),
  inference(subsumption_resolution,[],[f8100,f142])).
fof(f8100,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | ($spl10 | $spl12 | $spl16)),
  inference(backward_demodulation,[],[f4448,f149])).
fof(f4448,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl12 | $spl16)),
  inference(resolution,[],[f174,f172])).
fof(f8082,plain,(
  $false | ($spl1 | $spl2 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8043,f147])).
fof(f8043,plain,(
  s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | ($spl1 | $spl2 | $spl10)),
  inference(superposition,[],[f160,f166])).
fof(f166,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(bool,X0),X1),s(bool,sK5))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X1),s(X0,X2)))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f133,f161])).
fof(f161,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | ($spl2 | $spl10)),
  inference(forward_demodulation,[],[f149,f141])).
fof(f160,plain,(
  s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK3))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f140,f141])).
fof(f8060,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | s(bool,sK5) = s(bool,X34)) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f8059,f131])).
fof(f8059,plain,(
  ( ! [X33,X34,X32] : (s(bool,sK5) = s(bool,X34) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X32),s(bool,X33))))) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f8058,f2548])).
fof(f8058,plain,(
  ( ! [X33,X34,X32] : (s(bool,sK5) = s(bool,X34) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X32),s(bool,X33)))) | p(s(bool,X34))) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f8022,f4419])).
fof(f8022,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X32),s(bool,sK5))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X32),s(bool,X33)))) | p(s(bool,X34))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f135,f166])).
fof(f8057,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | s(bool,sK5) = s(bool,X31)) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f8056,f131])).
fof(f8056,plain,(
  ( ! [X30,X31,X29] : (s(bool,sK5) = s(bool,X31) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X29),s(bool,X30))))) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f8055,f2548])).
fof(f8055,plain,(
  ( ! [X30,X31,X29] : (s(bool,sK5) = s(bool,X31) | p(s(bool,X31)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X29),s(bool,X30))))) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f8021,f4419])).
fof(f8021,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) = s(bool,happ(s(fun(bool,bool),X29),s(bool,sK5))) | p(s(bool,X31)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X29),s(bool,X30))))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f135,f166])).
fof(f8050,plain,(
  ( ! [X15,X16] : (p(s(bool,X15)) | s(bool,sK5) = s(bool,X16)) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f8049,f131])).
fof(f8049,plain,(
  ( ! [X14,X15,X16] : (s(bool,sK5) = s(bool,X16) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X14),s(bool,X15))))) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f8048,f2548])).
fof(f8048,plain,(
  ( ! [X14,X15,X16] : (s(bool,sK5) = s(bool,X16) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X14),s(bool,X15)))) | p(s(bool,X16))) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f8011,f4419])).
fof(f8011,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) = s(bool,happ(s(fun(bool,bool),X14),s(bool,sK5))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X14),s(bool,X15)))) | p(s(bool,X16))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f166,f135])).
fof(f8047,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,sK5) = s(bool,X13)) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f8046,f131])).
fof(f8046,plain,(
  ( ! [X12,X13,X11] : (s(bool,sK5) = s(bool,X13) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(subsumption_resolution,[],[f8045,f2548])).
fof(f8045,plain,(
  ( ! [X12,X13,X11] : (s(bool,sK5) = s(bool,X13) | p(s(bool,X13)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl2 | $spl10 | $spl12 | $spl30)),
  inference(forward_demodulation,[],[f8010,f4419])).
fof(f8010,plain,(
  ( ! [X12,X13,X11] : (s(bool,X13) = s(bool,happ(s(fun(bool,bool),X11),s(bool,sK5))) | p(s(bool,X13)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f166,f135])).
fof(f7775,plain,(
  ( ! [X30,X28] : (p(s(bool,X28)) | s(bool,sK5) = s(bool,X30)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7774,f131])).
fof(f7774,plain,(
  ( ! [X30,X28,X27] : (s(bool,sK5) = s(bool,X30) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X28))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(subsumption_resolution,[],[f7773,f2548])).
fof(f7773,plain,(
  ( ! [X30,X28,X27] : (s(bool,sK5) = s(bool,X30) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X28)))) | p(s(bool,X30))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7700,f4419])).
fof(f7700,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X30) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X29))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X28)))) | p(s(bool,X30))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(superposition,[],[f135,f7571])).
fof(f7571,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(backward_demodulation,[],[f7558,f133])).
fof(f7558,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(duplicate_literal_removal,[],[f7557])).
fof(f7557,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(forward_demodulation,[],[f7534,f161])).
fof(f7534,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | ($spl12 | $spl82)),
  inference(superposition,[],[f132,f7495])).
fof(f7495,plain,(
  ( ! [X5] : (s(bool,sK5) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X5))))) ) | ($spl12 | $spl82)),
  inference(resolution,[],[f7480,f2548])).
fof(f7772,plain,(
  ( ! [X26,X24] : (p(s(bool,X24)) | s(bool,sK5) = s(bool,X26)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7771,f131])).
fof(f7771,plain,(
  ( ! [X26,X24,X23] : (s(bool,sK5) = s(bool,X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X24))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(subsumption_resolution,[],[f7770,f2548])).
fof(f7770,plain,(
  ( ! [X26,X24,X23] : (s(bool,sK5) = s(bool,X26) | p(s(bool,X26)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X24))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7699,f4419])).
fof(f7699,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X25))) | p(s(bool,X26)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X24))))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(superposition,[],[f135,f7571])).
fof(f7768,plain,(
  ( ! [X62,X63] : (s(bool,sK5) = s(bool,X63) | p(s(bool,X62))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(duplicate_literal_removal,[],[f7767])).
fof(f7767,plain,(
  ( ! [X62,X63] : (s(bool,sK5) = s(bool,X63) | p(s(bool,X62)) | s(bool,sK5) = s(bool,X63)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7766,f4419])).
fof(f7766,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X62)) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61))) | s(bool,sK5) = s(bool,X63)) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(forward_demodulation,[],[f7684,f131])).
fof(f7684,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X60),s(bool,X62)))) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61))) | s(bool,sK5) = s(bool,X63)) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(superposition,[],[f4406,f7571])).
fof(f7753,plain,(
  ( ! [X37,X36] : (s(bool,sK5) = s(bool,X37) | p(s(bool,X36))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(subsumption_resolution,[],[f7752,f2548])).
fof(f7752,plain,(
  ( ! [X37,X36] : (s(bool,sK5) = s(bool,X37) | p(s(bool,X36)) | p(s(bool,X37))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7751,f4419])).
fof(f7751,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X36)) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X34),s(bool,X35))) | p(s(bool,X37))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(forward_demodulation,[],[f7677,f131])).
fof(f7677,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X34),s(bool,X36)))) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X34),s(bool,X35))) | p(s(bool,X37))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(superposition,[],[f135,f7571])).
fof(f7750,plain,(
  ( ! [X33,X32] : (s(bool,sK5) = s(bool,X33) | p(s(bool,X32))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(subsumption_resolution,[],[f7749,f2548])).
fof(f7749,plain,(
  ( ! [X33,X32] : (s(bool,sK5) = s(bool,X33) | p(s(bool,X32)) | p(s(bool,X33))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7748,f4419])).
fof(f7748,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X32)) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X30),s(bool,X31))) | p(s(bool,X33))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(forward_demodulation,[],[f7676,f131])).
fof(f7676,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X30),s(bool,X32)))) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X30),s(bool,X31))) | p(s(bool,X33))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(superposition,[],[f135,f7571])).
fof(f7740,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,sK5) = s(bool,X21)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7739,f131])).
fof(f7739,plain,(
  ( ! [X21,X19,X20] : (s(bool,sK5) = s(bool,X21) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X20))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(subsumption_resolution,[],[f7738,f2548])).
fof(f7738,plain,(
  ( ! [X21,X19,X20] : (s(bool,sK5) = s(bool,X21) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X20)))) | p(s(bool,X21))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7665,f4419])).
fof(f7665,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X22))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X20)))) | p(s(bool,X21))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(superposition,[],[f7571,f135])).
fof(f7737,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,sK5) = s(bool,X17)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7736,f131])).
fof(f7736,plain,(
  ( ! [X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(subsumption_resolution,[],[f7735,f2548])).
fof(f7735,plain,(
  ( ! [X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,X17)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl82)),
  inference(forward_demodulation,[],[f7664,f4419])).
fof(f7664,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X18))) | p(s(bool,X17)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl2 | $spl10 | $spl12 | $spl82)),
  inference(superposition,[],[f7571,f135])).
fof(f7479,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) ) | $spl80),
  inference(cnf_transformation,[],[f7479_D])).
fof(f7479_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f7480,plain,(
  ( ! [X1] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1)))))) ) | $spl82),
  inference(cnf_transformation,[],[f7480_D])).
fof(f7480_D,plain,(
  ( ! [X1] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f7482,plain,(
  ( ! [X2,X3] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3))))) | s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) ) | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f7481,f6526])).
fof(f6526,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) ) | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f6505,f172])).
fof(f6505,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) | ~p(s(bool,sK5)) | ~p(s(bool,X2))) ) | ($spl12 | $spl76)),
  inference(superposition,[],[f6487,f138])).
fof(f7481,plain,(
  ( ! [X2,X3] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3))))) | s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) | p(s(bool,X2))) ) | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f7462,f6500])).
fof(f6500,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2)))))) ) | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f6492,f172])).
fof(f6492,plain,(
  ( ! [X2] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2))))) | ~p(s(bool,sK5)) | ~p(s(bool,X2))) ) | $spl76),
  inference(superposition,[],[f6374,f138])).
fof(f7462,plain,(
  ( ! [X2,X3] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X3))))) | s(bool,X2) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) | p(s(bool,X2)) | p(s(bool,X3))) ) | ($spl12 | $spl76)),
  inference(superposition,[],[f7416,f135])).
fof(f7416,plain,(
  ( ! [X0] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) ) | ($spl12 | $spl76)),
  inference(factoring,[],[f6542])).
fof(f6542,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | ($spl12 | $spl76)),
  inference(resolution,[],[f6500,f135])).
fof(f7478,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1))))) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0))))) ) | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f7477,f6526])).
fof(f7477,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1))))) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) | p(s(bool,X0))) ) | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f7461,f6500])).
fof(f7461,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1))))) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl12 | $spl76)),
  inference(superposition,[],[f7416,f135])).
fof(f7148,plain,(
  $false | ($spl1 | $spl2 | $spl10 | $spl12 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f7114,f6950])).
fof(f6950,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X7) = s(X5,happ(s(fun(bool,X5),X6),s(bool,X8)))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl38)),
  inference(superposition,[],[f6819,f131])).
fof(f6819,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl38)),
  inference(backward_demodulation,[],[f6803,f133])).
fof(f6803,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl38)),
  inference(duplicate_literal_removal,[],[f6795])).
fof(f6795,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl38)),
  inference(superposition,[],[f161,f6602])).
fof(f6602,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl12 | $spl30 | $spl38)),
  inference(backward_demodulation,[],[f6601,f132])).
fof(f6601,plain,(
  ( ! [X17] : (s(bool,sK5) = s(bool,X17)) ) | ($spl12 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f6600,f2548])).
fof(f6600,plain,(
  ( ! [X17] : (s(bool,sK5) = s(bool,X17) | p(s(bool,X17))) ) | ($spl12 | $spl30 | $spl38)),
  inference(forward_demodulation,[],[f1898,f4419])).
fof(f7114,plain,(
  ( ! [X208,X209] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),X208),s(bool,X209)))) ) | ($spl1 | $spl2 | $spl10 | $spl12 | $spl30 | $spl38)),
  inference(superposition,[],[f160,f6950])).
fof(f6485,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | ($spl16 | $spl73)),
  inference(subsumption_resolution,[],[f6477,f174])).
fof(f6477,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | $spl73),
  inference(superposition,[],[f5391,f138])).
fof(f6484,plain,(
  ( ! [X2] : (~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | ~p(s(bool,X2))) ) | ($spl16 | $spl73)),
  inference(subsumption_resolution,[],[f6476,f174])).
fof(f6476,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | ~p(s(bool,X2))) ) | $spl73),
  inference(superposition,[],[f5391,f138])).
fof(f6475,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | p(s(bool,X1))) ) | $spl73),
  inference(superposition,[],[f5391,f135])).
fof(f6474,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | $spl73),
  inference(superposition,[],[f5391,f135])).
fof(f1687,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5)))) ) | $spl28),
  inference(cnf_transformation,[],[f1687_D])).
fof(f1687_D,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f6375,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | ($spl16 | $spl72)),
  inference(subsumption_resolution,[],[f6360,f174])).
fof(f6360,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))) | p(s(bool,X1))) ) | $spl72),
  inference(superposition,[],[f5392,f135])).
fof(f6372,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | ($spl16 | $spl72)),
  inference(subsumption_resolution,[],[f6359,f174])).
fof(f6359,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) ) | $spl72),
  inference(superposition,[],[f5392,f135])).
fof(f6330,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))),s(fun(bool,bool),X24),s(bool,X25))))) ) | $spl74),
  inference(cnf_transformation,[],[f6330_D])).
fof(f6330_D,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))),s(fun(bool,bool),X24),s(bool,X25))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f6338,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,sK5) = s(bool,X13) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X10))),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f6337,f2548])).
fof(f6337,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,sK5) = s(bool,X13) | p(s(bool,X13)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X10))),s(fun(bool,bool),X11),s(bool,X12))))) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f5933,f4419])).
fof(f5933,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X13) = s(bool,happ(s(fun(bool,bool),X11),s(bool,X10))) | p(s(bool,X13)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X10))),s(fun(bool,bool),X11),s(bool,X12))))) )),
  inference(superposition,[],[f133,f135])).
fof(f6336,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X14))),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f6335,f2548])).
fof(f6335,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X14))),s(fun(bool,bool),X15),s(bool,X16)))) | p(s(bool,X17))) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f5934,f4419])).
fof(f5934,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X14))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X14))),s(fun(bool,bool),X15),s(bool,X16)))) | p(s(bool,X17))) )),
  inference(superposition,[],[f133,f135])).
fof(f6332,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,sK5) = s(bool,X22) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))),s(fun(bool,bool),X20),s(bool,X21))))) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f6331,f2548])).
fof(f6331,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,sK5) = s(bool,X22) | p(s(bool,X22)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))),s(fun(bool,bool),X20),s(bool,X21))))) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f5943,f4419])).
fof(f5943,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X19))) | p(s(bool,X22)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X19))),s(fun(bool,bool),X20),s(bool,X21))))) )),
  inference(superposition,[],[f135,f133])).
fof(f6329,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,sK5) = s(bool,X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))),s(fun(bool,bool),X24),s(bool,X25))))) ) | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f6328,f2548])).
fof(f6328,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,sK5) = s(bool,X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))),s(fun(bool,bool),X24),s(bool,X25)))) | p(s(bool,X26))) ) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f5944,f4419])).
fof(f5944,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),X24),s(bool,X23))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))),s(fun(bool,bool),X24),s(bool,X25)))) | p(s(bool,X26))) )),
  inference(superposition,[],[f135,f133])).
fof(f6317,plain,(
  $false | ($spl1 | $spl2 | $spl16 | $spl40)),
  inference(subsumption_resolution,[],[f6283,f6121])).
fof(f6121,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X7) = s(X5,happ(s(fun(bool,X5),X6),s(bool,X8)))) ) | ($spl16 | $spl40)),
  inference(superposition,[],[f5966,f131])).
fof(f5966,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl16 | $spl40)),
  inference(backward_demodulation,[],[f5895,f133])).
fof(f5895,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl16 | $spl40)),
  inference(subsumption_resolution,[],[f174,f1899])).
fof(f6283,plain,(
  ( ! [X210,X209] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),X209),s(bool,X210)))) ) | ($spl1 | $spl2 | $spl16 | $spl40)),
  inference(superposition,[],[f160,f6121])).
fof(f5891,plain,(
  $false | ($spl1 | $spl2 | $spl10 | $spl12 | $spl56)),
  inference(subsumption_resolution,[],[f5857,f5545])).
fof(f5545,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X7) = s(X5,happ(s(fun(bool,X5),X6),s(bool,X8)))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f5490,f131])).
fof(f5490,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(forward_demodulation,[],[f133,f5382])).
fof(f5382,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(duplicate_literal_removal,[],[f5374])).
fof(f5374,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f161,f4673])).
fof(f4673,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1)) ) | ($spl12 | $spl56)),
  inference(duplicate_literal_removal,[],[f4657])).
fof(f4657,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1)) ) | ($spl12 | $spl56)),
  inference(superposition,[],[f132,f4436])).
fof(f4436,plain,(
  ( ! [X0] : (s(bool,sK5) = s(bool,sK6(X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl12 | $spl56)),
  inference(resolution,[],[f2006,f2548])).
fof(f5857,plain,(
  ( ! [X216,X217] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),X216),s(bool,X217)))) ) | ($spl1 | $spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f160,f5545])).
fof(f5651,plain,(
  ( ! [X30,X28] : (p(s(bool,X28)) | s(bool,sK5) = s(bool,X30)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5650,f131])).
fof(f5650,plain,(
  ( ! [X30,X28,X27] : (s(bool,sK5) = s(bool,X30) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X28))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(subsumption_resolution,[],[f5649,f2548])).
fof(f5649,plain,(
  ( ! [X30,X28,X27] : (s(bool,sK5) = s(bool,X30) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X28)))) | p(s(bool,X30))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5585,f4419])).
fof(f5585,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X30) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X29))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X28)))) | p(s(bool,X30))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f135,f5490])).
fof(f5648,plain,(
  ( ! [X26,X24] : (p(s(bool,X24)) | s(bool,sK5) = s(bool,X26)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5647,f131])).
fof(f5647,plain,(
  ( ! [X26,X24,X23] : (s(bool,sK5) = s(bool,X26) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X24))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(subsumption_resolution,[],[f5646,f2548])).
fof(f5646,plain,(
  ( ! [X26,X24,X23] : (s(bool,sK5) = s(bool,X26) | p(s(bool,X26)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X24))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5584,f4419])).
fof(f5584,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X25))) | p(s(bool,X26)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X24))))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f135,f5490])).
fof(f5634,plain,(
  ( ! [X37,X36] : (s(bool,sK5) = s(bool,X37) | p(s(bool,X36))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(subsumption_resolution,[],[f5633,f2548])).
fof(f5633,plain,(
  ( ! [X37,X36] : (s(bool,sK5) = s(bool,X37) | p(s(bool,X36)) | p(s(bool,X37))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5632,f4419])).
fof(f5632,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X36)) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X34),s(bool,X35))) | p(s(bool,X37))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(forward_demodulation,[],[f5563,f131])).
fof(f5563,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X34),s(bool,X36)))) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X34),s(bool,X35))) | p(s(bool,X37))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f135,f5490])).
fof(f5631,plain,(
  ( ! [X33,X32] : (s(bool,sK5) = s(bool,X33) | p(s(bool,X32))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(subsumption_resolution,[],[f5630,f2548])).
fof(f5630,plain,(
  ( ! [X33,X32] : (s(bool,sK5) = s(bool,X33) | p(s(bool,X32)) | p(s(bool,X33))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5629,f4419])).
fof(f5629,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X32)) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X30),s(bool,X31))) | p(s(bool,X33))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(forward_demodulation,[],[f5562,f131])).
fof(f5562,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X30),s(bool,X32)))) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X30),s(bool,X31))) | p(s(bool,X33))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f135,f5490])).
fof(f5621,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,sK5) = s(bool,X21)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5620,f131])).
fof(f5620,plain,(
  ( ! [X21,X19,X20] : (s(bool,sK5) = s(bool,X21) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X20))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(subsumption_resolution,[],[f5619,f2548])).
fof(f5619,plain,(
  ( ! [X21,X19,X20] : (s(bool,sK5) = s(bool,X21) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X20)))) | p(s(bool,X21))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5548,f4419])).
fof(f5548,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X22))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X20)))) | p(s(bool,X21))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f5490,f135])).
fof(f5618,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,sK5) = s(bool,X17)) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5617,f131])).
fof(f5617,plain,(
  ( ! [X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(subsumption_resolution,[],[f5616,f2548])).
fof(f5616,plain,(
  ( ! [X17,X15,X16] : (s(bool,sK5) = s(bool,X17) | p(s(bool,X17)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl2 | $spl10 | $spl12 | $spl30 | $spl56)),
  inference(forward_demodulation,[],[f5547,f4419])).
fof(f5547,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X18))) | p(s(bool,X17)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X15),s(bool,X16))))) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(superposition,[],[f5490,f135])).
fof(f2212,plain,(
  ( ! [X34] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X34)) ) | $spl62),
  inference(cnf_transformation,[],[f2212_D])).
fof(f2212_D,plain,(
  ( ! [X34] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X34)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5392,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | $spl72),
  inference(cnf_transformation,[],[f5392_D])).
fof(f5392_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f5396,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1)) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(forward_demodulation,[],[f5381,f161])).
fof(f5381,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))))) ) | ($spl12 | $spl56)),
  inference(duplicate_literal_removal,[],[f5375])).
fof(f5375,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1)) ) | ($spl12 | $spl56)),
  inference(superposition,[],[f132,f4673])).
fof(f5390,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4)) ) | ($spl2 | $spl10 | $spl12 | $spl56)),
  inference(forward_demodulation,[],[f5384,f161])).
fof(f5384,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))))) ) | ($spl12 | $spl56)),
  inference(duplicate_literal_removal,[],[f5364])).
fof(f5364,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK5))))))) ) | ($spl12 | $spl56)),
  inference(superposition,[],[f4673,f132])).
fof(f5227,plain,(
  $false | ($spl1 | $spl2 | $spl10 | $spl12 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f5193,f5027])).
fof(f5027,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X7) = s(X5,happ(s(fun(bool,X5),X6),s(bool,X8)))) ) | ($spl2 | $spl10 | $spl12 | $spl38 | $spl46)),
  inference(superposition,[],[f4896,f131])).
fof(f4896,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl2 | $spl10 | $spl12 | $spl38 | $spl46)),
  inference(backward_demodulation,[],[f4880,f133])).
fof(f4880,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl12 | $spl38 | $spl46)),
  inference(duplicate_literal_removal,[],[f4872])).
fof(f4872,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl12 | $spl38 | $spl46)),
  inference(superposition,[],[f161,f4685])).
fof(f4685,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl12 | $spl38 | $spl46)),
  inference(backward_demodulation,[],[f4683,f132])).
fof(f4683,plain,(
  ( ! [X17] : (s(bool,sK5) = s(bool,X17)) ) | ($spl12 | $spl38 | $spl46)),
  inference(subsumption_resolution,[],[f4682,f2548])).
fof(f4682,plain,(
  ( ! [X17] : (s(bool,sK5) = s(bool,X17) | p(s(bool,X17))) ) | ($spl38 | $spl46)),
  inference(forward_demodulation,[],[f1898,f1915])).
fof(f5193,plain,(
  ( ! [X208,X207] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),X207),s(bool,X208)))) ) | ($spl1 | $spl2 | $spl10 | $spl12 | $spl38 | $spl46)),
  inference(superposition,[],[f160,f5027])).
fof(f4605,plain,(
  ( ! [X41,X42] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,happ(s(fun(bool,bool),X41),s(bool,X42)))))) ) | $spl70),
  inference(cnf_transformation,[],[f4605_D])).
fof(f4605_D,plain,(
  ( ! [X41,X42] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,happ(s(fun(bool,bool),X41),s(bool,X42)))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4590,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,X43)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,happ(s(fun(bool,bool),X41),s(bool,X42)))))) ) | ($spl16 | $spl68)),
  inference(superposition,[],[f174,f3383])).
fof(f4389,plain,(
  $false | ($spl1 | $spl2 | $spl8 | $spl10 | $spl66)),
  inference(subsumption_resolution,[],[f4385,f3474])).
fof(f3474,plain,(
  ( ! [X0] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X0)))) ) | ($spl8 | $spl66)),
  inference(superposition,[],[f147,f3380])).
fof(f4385,plain,(
  ( ! [X115] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X115)))) ) | ($spl1 | $spl2 | $spl10 | $spl66)),
  inference(superposition,[],[f160,f4255])).
fof(f4255,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl2 | $spl10 | $spl66)),
  inference(forward_demodulation,[],[f133,f3475])).
fof(f3475,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) ) | ($spl2 | $spl10 | $spl66)),
  inference(superposition,[],[f161,f3380])).
fof(f4101,plain,(
  $false | ($spl1 | $spl2 | $spl10 | $spl12 | $spl36)),
  inference(subsumption_resolution,[],[f4070,f3918])).
fof(f3918,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X7) = s(X5,happ(s(fun(bool,X5),X6),s(bool,X8)))) ) | ($spl2 | $spl10 | $spl12 | $spl36)),
  inference(superposition,[],[f3864,f131])).
fof(f3864,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl2 | $spl10 | $spl12 | $spl36)),
  inference(forward_demodulation,[],[f133,f3549])).
fof(f3549,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl12 | $spl36)),
  inference(duplicate_literal_removal,[],[f3542])).
fof(f3542,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl12 | $spl36)),
  inference(superposition,[],[f161,f3518])).
fof(f3518,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl12 | $spl36)),
  inference(forward_demodulation,[],[f132,f3367])).
fof(f3367,plain,(
  ( ! [X6] : (s(bool,sK5) = s(bool,X6)) ) | ($spl12 | $spl36)),
  inference(resolution,[],[f1891,f172])).
fof(f4070,plain,(
  ( ! [X194,X193] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),X193),s(bool,X194)))) ) | ($spl1 | $spl2 | $spl10 | $spl12 | $spl36)),
  inference(superposition,[],[f160,f3918])).
fof(f3863,plain,(
  $false | ($spl1 | $spl2 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f3831,f3283])).
fof(f3283,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X7) = s(X5,happ(s(fun(bool,X5),X6),s(bool,X8)))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f131])).
fof(f3171,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f133,f2278])).
fof(f2278,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2)))) ) | ($spl14 | $spl16)),
  inference(subsumption_resolution,[],[f174,f173])).
fof(f3831,plain,(
  ( ! [X200,X199] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),X199),s(bool,X200)))) ) | ($spl1 | $spl2 | $spl14 | $spl16)),
  inference(superposition,[],[f160,f3283])).
fof(f3383,plain,(
  ( ! [X19,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X20)))) ) | $spl68),
  inference(cnf_transformation,[],[f3383_D])).
fof(f3383_D,plain,(
  ( ! [X19,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X20)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3382,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,X21)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X20)))) ) | ($spl14 | $spl16 | $spl36)),
  inference(forward_demodulation,[],[f3374,f131])).
fof(f3374,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X21)))) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X20)))) ) | ($spl14 | $spl16 | $spl36)),
  inference(superposition,[],[f1891,f3171])).
fof(f1890,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X14))))) ) | $spl34),
  inference(cnf_transformation,[],[f1890_D])).
fof(f1890_D,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X14))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3380,plain,(
  ( ! [X15,X13] : (s(bool,X13) = s(bool,X15)) ) | $spl66),
  inference(cnf_transformation,[],[f3380_D])).
fof(f3380_D,plain,(
  ( ! [X15,X13] : (s(bool,X13) = s(bool,X15)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3379,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X13) = s(bool,X15) | ~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X14))))) ) | ($spl14 | $spl16 | $spl36)),
  inference(forward_demodulation,[],[f3372,f131])).
fof(f3372,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X14)))) | s(bool,X15) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X12),s(bool,X13)))) ) | ($spl14 | $spl16 | $spl36)),
  inference(superposition,[],[f1891,f3171])).
fof(f1886,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))))) ) | $spl30),
  inference(cnf_transformation,[],[f1886_D])).
fof(f1886_D,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1887,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8)) ) | $spl32),
  inference(cnf_transformation,[],[f1887_D])).
fof(f1887_D,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3364,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X7) = s(bool,X8) | p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3241,f131])).
fof(f3241,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X7) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X5),s(bool,X8))) | p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f135])).
fof(f3363,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X11) = s(bool,X12) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,X11))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3242,f131])).
fof(f3242,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X11) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X9),s(bool,X12))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,X11))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f135])).
fof(f1891,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,X16)) ) | $spl36),
  inference(cnf_transformation,[],[f1891_D])).
fof(f1891_D,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,X16)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3362,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X15) = s(bool,X16) | ~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X14)))) | ~p(s(bool,X15))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3243,f131])).
fof(f3243,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X15) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X13),s(bool,X16))) | ~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X14)))) | ~p(s(bool,X15))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f138])).
fof(f3361,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X19) = s(bool,X20) | ~p(s(bool,X19)) | ~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3244,f131])).
fof(f3244,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X19) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X17),s(bool,X20))) | ~p(s(bool,X19)) | ~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f138])).
fof(f1898,plain,(
  ( ! [X17,X15,X18] : (s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X18))) | p(s(bool,X17))) ) | $spl38),
  inference(cnf_transformation,[],[f1898_D])).
fof(f1898_D,plain,(
  ( ! [X17,X15,X18] : (s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X18))) | p(s(bool,X17))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f3360,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,X10)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),X9),s(bool,X12))) | p(s(bool,X11))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3284,f131])).
fof(f3284,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X11) = s(bool,happ(s(fun(bool,bool),X9),s(bool,X12))) | p(s(bool,X11)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X9),s(bool,X10))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f135])).
fof(f3359,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X14)) | s(bool,X15) = s(bool,happ(s(fun(bool,bool),X13),s(bool,X16))) | p(s(bool,X15))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3285,f131])).
fof(f3285,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X15) = s(bool,happ(s(fun(bool,bool),X13),s(bool,X16))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X13),s(bool,X14)))) | p(s(bool,X15))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f135])).
fof(f1902,plain,(
  ( ! [X26,X23,X25] : (~p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26)))) ) | $spl42),
  inference(cnf_transformation,[],[f1902_D])).
fof(f1902_D,plain,(
  ( ! [X26,X23,X25] : (~p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3358,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,X18)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X20))) | ~p(s(bool,X19))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3286,f131])).
fof(f3286,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X20))) | ~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X17),s(bool,X18)))) | ~p(s(bool,X19))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f138])).
fof(f3357,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,X22)) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X21),s(bool,X24))) | ~p(s(bool,X23))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3287,f131])).
fof(f3287,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X23) = s(bool,happ(s(fun(bool,bool),X21),s(bool,X24))) | ~p(s(bool,X23)) | ~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X21),s(bool,X22))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f3171,f138])).
fof(f3356,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X7) = s(bool,X8) | p(s(bool,X8)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3296,f131])).
fof(f3296,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X8) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X5),s(bool,X7))) | p(s(bool,X8)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f135,f3171])).
fof(f3355,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X11) = s(bool,X12) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,X12))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3297,f131])).
fof(f3297,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X12) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X9),s(bool,X11))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,X12))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f135,f3171])).
fof(f3354,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X15)) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X13),s(bool,X14))) | p(s(bool,X16))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3298,f131])).
fof(f3298,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X13),s(bool,X15)))) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X13),s(bool,X14))) | p(s(bool,X16))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f135,f3171])).
fof(f3353,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X19)) | s(bool,X20) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | p(s(bool,X20))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3299,f131])).
fof(f3299,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X17),s(bool,X19)))) | s(bool,X20) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))) | p(s(bool,X20))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f135,f3171])).
fof(f3352,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X23) = s(bool,X24) | ~p(s(bool,happ(s(fun(bool,bool),X21),s(bool,X22)))) | ~p(s(bool,X24))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3300,f131])).
fof(f3300,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X24) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X21),s(bool,X23))) | ~p(s(bool,happ(s(fun(bool,bool),X21),s(bool,X22)))) | ~p(s(bool,X24))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f138,f3171])).
fof(f3351,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X27) = s(bool,X28) | ~p(s(bool,X28)) | ~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3301,f131])).
fof(f3301,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X28) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X25),s(bool,X27))) | ~p(s(bool,X28)) | ~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f138,f3171])).
fof(f3350,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X31)) | s(bool,X32) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | ~p(s(bool,X32))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3302,f131])).
fof(f3302,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X29),s(bool,X31)))) | s(bool,X32) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | ~p(s(bool,X32))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f138,f3171])).
fof(f3349,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,happ(s(fun(bool,bool),X33),s(bool,X34))) | ~p(s(bool,X36))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3303,f131])).
fof(f3303,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X33),s(bool,X35)))) | s(bool,X36) = s(bool,happ(s(fun(bool,bool),X33),s(bool,X34))) | ~p(s(bool,X36))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f138,f3171])).
fof(f3348,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,X10)) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),X9),s(bool,X11))) | p(s(bool,X12))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3319,f131])).
fof(f3319,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),X9),s(bool,X11))) | p(s(bool,X12)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X9),s(bool,X10))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f135,f3171])).
fof(f3347,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X14)) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X13),s(bool,X15))) | p(s(bool,X16))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3320,f131])).
fof(f3320,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X16) = s(bool,happ(s(fun(bool,bool),X13),s(bool,X15))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X13),s(bool,X14)))) | p(s(bool,X16))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f135,f3171])).
fof(f3346,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X18) = s(bool,X20) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X19)))) | p(s(bool,X20))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3321,f131])).
fof(f3321,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X19)))) | s(bool,X20) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X17),s(bool,X18))) | p(s(bool,X20))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f135,f3171])).
fof(f3345,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) = s(bool,X24) | p(s(bool,happ(s(fun(bool,bool),X21),s(bool,X23)))) | p(s(bool,X24))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3322,f131])).
fof(f3322,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,happ(s(fun(bool,bool),X21),s(bool,X23)))) | s(bool,X24) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X21),s(bool,X22))) | p(s(bool,X24))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f135,f3171])).
fof(f3344,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,X26)) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X27))) | ~p(s(bool,X28))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3323,f131])).
fof(f3323,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X28) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X27))) | ~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X25),s(bool,X26)))) | ~p(s(bool,X28))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f138,f3171])).
fof(f3343,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X30)) | s(bool,X32) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X31))) | ~p(s(bool,X32))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3324,f131])).
fof(f3324,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X32) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X31))) | ~p(s(bool,X32)) | ~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X29),s(bool,X30))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f138,f3171])).
fof(f3342,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) = s(bool,X36) | ~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X35)))) | ~p(s(bool,X36))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3325,f131])).
fof(f3325,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X35)))) | s(bool,X36) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X33),s(bool,X34))) | ~p(s(bool,X36))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f138,f3171])).
fof(f3341,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) = s(bool,X40) | ~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X39)))) | ~p(s(bool,X40))) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f3326,f131])).
fof(f3326,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X39)))) | s(bool,X40) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X37),s(bool,X38))) | ~p(s(bool,X40))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f138,f3171])).
fof(f3340,plain,(
  $false | ($spl1 | $spl2 | $spl8 | $spl12 | $spl14 | $spl16 | $spl32)),
  inference(subsumption_resolution,[],[f3339,f2744])).
fof(f2744,plain,(
  ( ! [X0] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X0)))) ) | ($spl8 | $spl12 | $spl32)),
  inference(superposition,[],[f147,f2674])).
fof(f2674,plain,(
  ( ! [X1] : (s(bool,sK5) = s(bool,X1)) ) | ($spl12 | $spl32)),
  inference(subsumption_resolution,[],[f2548,f1887])).
fof(f3339,plain,(
  ( ! [X84] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X84)))) ) | ($spl1 | $spl2 | $spl14 | $spl16)),
  inference(superposition,[],[f160,f3171])).
fof(f3125,plain,(
  $false | ($spl1 | $spl2 | $spl8 | $spl12 | $spl32 | $spl62)),
  inference(subsumption_resolution,[],[f3123,f2744])).
fof(f3123,plain,(
  ( ! [X104] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X104)))) ) | ($spl1 | $spl2 | $spl62)),
  inference(superposition,[],[f2870,f2938])).
fof(f2938,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | $spl62),
  inference(forward_demodulation,[],[f133,f2212])).
fof(f2870,plain,(
  ( ! [X0] : (s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X0),s(fun(bool,sK0),sK4),s(sK0,sK3)))) ) | ($spl1 | $spl2 | $spl62)),
  inference(superposition,[],[f160,f2212])).
fof(f3124,plain,(
  $false | ($spl1 | $spl2 | $spl8 | $spl12 | $spl32 | $spl62)),
  inference(subsumption_resolution,[],[f3122,f2744])).
fof(f3122,plain,(
  ( ! [X103] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X103)))) ) | ($spl1 | $spl2 | $spl62)),
  inference(superposition,[],[f160,f2938])).
fof(f2865,plain,(
  $false | ($spl1 | $spl2 | $spl8 | $spl28)),
  inference(trivial_inequality_removal,[],[f2864])).
fof(f2864,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ($spl1 | $spl2 | $spl8 | $spl28)),
  inference(superposition,[],[f1884,f147])).
fof(f1884,plain,(
  ( ! [X138] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X138)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(superposition,[],[f160,f1690])).
fof(f1690,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | $spl28),
  inference(backward_demodulation,[],[f1687,f133])).
fof(f2853,plain,(
  $false | ($spl2 | $spl6 | $spl8 | $spl28)),
  inference(trivial_inequality_removal,[],[f2852])).
fof(f2852,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ($spl2 | $spl6 | $spl8 | $spl28)),
  inference(superposition,[],[f2841,f147])).
fof(f2841,plain,(
  ( ! [X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5)))) ) | ($spl2 | $spl6 | $spl28)),
  inference(subsumption_resolution,[],[f2281,f1712])).
fof(f1712,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | $spl28),
  inference(duplicate_literal_removal,[],[f1708])).
fof(f1708,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | $spl28),
  inference(superposition,[],[f132,f1687])).
fof(f2281,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5)))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f145,f141])).
fof(f2840,plain,(
  $false | ($spl12 | $spl32 | $spl34)),
  inference(subsumption_resolution,[],[f1890,f2691])).
fof(f2691,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl12 | $spl32)),
  inference(superposition,[],[f172,f2674])).
fof(f2657,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f2653,f144])).
fof(f2653,plain,(
  s(sK0,sK2) = s(sK0,sK3) | ($spl0 | $spl2)),
  inference(superposition,[],[f131,f2282])).
fof(f2282,plain,(
  s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK3))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f139,f141])).
fof(f2656,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f2652,f144])).
fof(f2652,plain,(
  s(sK0,sK2) = s(sK0,sK3) | ($spl0 | $spl2)),
  inference(superposition,[],[f2282,f131])).
fof(f2561,plain,(
  ( ! [X4,X5] : (s(bool,sK5) != s(bool,X5) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | ~p(s(bool,X5))) ) | ($spl24 | $spl56)),
  inference(subsumption_resolution,[],[f1993,f2006])).
fof(f1993,plain,(
  ( ! [X4,X5] : (s(bool,sK5) != s(bool,X5) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | ~p(s(bool,sK6(X4))) | ~p(s(bool,X5))) ) | $spl24),
  inference(superposition,[],[f1679,f138])).
fof(f1679,plain,(
  ( ! [X0] : (s(bool,sK5) != s(bool,sK6(X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f1673])).
fof(f1673,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s(bool,sK5) != s(bool,sK6(X0))) ) | $spl24),
  inference(superposition,[],[f132,f1654])).
fof(f2560,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,X7) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X6) | ~p(s(bool,X7))) ) | ($spl24 | $spl56)),
  inference(subsumption_resolution,[],[f1994,f2006])).
fof(f1994,plain,(
  ( ! [X6,X7] : (s(bool,sK5) != s(bool,X7) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X6) | ~p(s(bool,X7)) | ~p(s(bool,sK6(X6)))) ) | $spl24),
  inference(superposition,[],[f1679,f138])).
fof(f2556,plain,(
  $false | ($spl12 | $spl26)),
  inference(trivial_inequality_removal,[],[f2547])).
fof(f2547,plain,(
  s(bool,sK5) != s(bool,sK5) | ($spl12 | $spl26)),
  inference(resolution,[],[f172,f1686])).
fof(f2545,plain,(
  $false | ($spl2 | $spl5 | $spl10 | $spl18)),
  inference(subsumption_resolution,[],[f2543,f1112])).
fof(f1112,plain,(
  ( ! [X4,X0,X3] : (s(X0,X3) = s(X0,X4)) ) | ($spl2 | $spl10 | $spl18)),
  inference(superposition,[],[f977,f977])).
fof(f977,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,X7) = s(X5,happ(s(fun(bool,X5),X6),s(bool,X8)))) ) | ($spl2 | $spl10 | $spl18)),
  inference(superposition,[],[f745,f131])).
fof(f745,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X2),s(X0,X3)))) ) | ($spl2 | $spl10 | $spl18)),
  inference(backward_demodulation,[],[f705,f133])).
fof(f705,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl18)),
  inference(duplicate_literal_removal,[],[f697])).
fof(f697,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl2 | $spl10 | $spl18)),
  inference(superposition,[],[f161,f579])).
fof(f579,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | $spl18),
  inference(backward_demodulation,[],[f552,f132])).
fof(f2543,plain,(
  ( ! [X137] : (s(sK0,sK2) != s(sK0,X137)) ) | ($spl2 | $spl5 | $spl10 | $spl18)),
  inference(superposition,[],[f144,f1112])).
fof(f2544,plain,(
  $false | ($spl2 | $spl5 | $spl10 | $spl18)),
  inference(subsumption_resolution,[],[f2502,f1112])).
fof(f2502,plain,(
  ( ! [X137] : (s(sK0,sK2) != s(sK0,X137)) ) | ($spl2 | $spl5 | $spl10 | $spl18)),
  inference(superposition,[],[f144,f1112])).
fof(f2279,plain,(
  $false | ($spl12 | $spl44)),
  inference(subsumption_resolution,[],[f172,f1903])).
fof(f2272,plain,(
  $false | ($spl30 | $spl44)),
  inference(subsumption_resolution,[],[f1886,f1903])).
fof(f2222,plain,(
  $false | ($spl13 | $spl40)),
  inference(resolution,[],[f1899,f171])).
fof(f2213,plain,(
  ( ! [X35,X33] : (~p(s(bool,X33)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35))) | s(bool,X33) = s(bool,X35)) ) | $spl64),
  inference(cnf_transformation,[],[f2213_D])).
fof(f2213_D,plain,(
  ( ! [X35,X33] : (~p(s(bool,X33)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35))) | s(bool,X33) = s(bool,X35)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2211,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X34) | s(bool,X33) = s(bool,X35) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35)))) ) | ($spl13 | $spl14 | $spl24 | $spl56)),
  inference(subsumption_resolution,[],[f2197,f1679])).
fof(f2197,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | s(bool,sK5) = s(bool,sK6(X34)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X34) | s(bool,X33) = s(bool,X35) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35)))) ) | ($spl13 | $spl14 | $spl56)),
  inference(superposition,[],[f2011,f260])).
fof(f260,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,X5) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X4))) | s(bool,sK5) = s(bool,X5)) ) | ($spl13 | $spl14)),
  inference(resolution,[],[f208,f175])).
fof(f175,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,sK5) = s(bool,X0)) ) | $spl13),
  inference(resolution,[],[f171,f135])).
fof(f208,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) ) | $spl14),
  inference(resolution,[],[f173,f138])).
fof(f2011,plain,(
  ( ! [X4,X3] : (~p(s(bool,X4)) | s(bool,X4) = s(bool,sK6(X3)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X3)) ) | $spl56),
  inference(resolution,[],[f2006,f138])).
fof(f2210,plain,(
  ( ! [X29,X27] : (s(bool,X27) = s(bool,sK6(X29)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X29)) ) | $spl60),
  inference(cnf_transformation,[],[f2210_D])).
fof(f2210_D,plain,(
  ( ! [X29,X27] : (s(bool,X27) = s(bool,sK6(X29)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X29)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2209,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X28)) | s(bool,X27) = s(bool,sK6(X29)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X29) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27)))) ) | ($spl13 | $spl14 | $spl26 | $spl56)),
  inference(subsumption_resolution,[],[f2195,f1686])).
fof(f2195,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X28)) | s(bool,X27) = s(bool,sK6(X29)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X29) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X27))) | s(bool,sK5) = s(bool,X28)) ) | ($spl13 | $spl14 | $spl56)),
  inference(superposition,[],[f2011,f260])).
fof(f2206,plain,(
  ( ! [X17,X15] : (s(bool,X15) = s(bool,sK6(X17)) | s(bool,sK5) = s(bool,X15) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X17)) ) | $spl58),
  inference(cnf_transformation,[],[f2206_D])).
fof(f2206_D,plain,(
  ( ! [X17,X15] : (s(bool,X15) = s(bool,sK6(X17)) | s(bool,sK5) = s(bool,X15) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X17)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2208,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) = s(bool,sK6(X26)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X26) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) | s(bool,sK5) = s(bool,X24)) ) | ($spl13 | $spl14 | $spl56)),
  inference(subsumption_resolution,[],[f2194,f173])).
fof(f2194,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X25)) | s(bool,X24) = s(bool,sK6(X26)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X26) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) | s(bool,sK5) = s(bool,X24)) ) | ($spl13 | $spl14 | $spl56)),
  inference(superposition,[],[f2011,f260])).
fof(f2207,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X19)) | s(bool,X18) = s(bool,sK6(X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X20) | s(bool,sK5) = s(bool,X18)) ) | ($spl13 | $spl26 | $spl56)),
  inference(subsumption_resolution,[],[f2192,f1686])).
fof(f2192,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X19)) | s(bool,X18) = s(bool,sK6(X20)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X20) | s(bool,sK5) = s(bool,X19) | s(bool,sK5) = s(bool,X18)) ) | ($spl13 | $spl56)),
  inference(superposition,[],[f2011,f227])).
fof(f227,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,sK5) = s(bool,X2) | s(bool,sK5) = s(bool,X3)) ) | $spl13),
  inference(resolution,[],[f183,f175])).
fof(f183,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,sK5) = s(bool,X0)) ) | $spl13),
  inference(resolution,[],[f175,f138])).
fof(f2205,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X16)) | s(bool,X15) = s(bool,sK6(X17)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X17) | s(bool,sK5) = s(bool,X15)) ) | ($spl13 | $spl26 | $spl56)),
  inference(subsumption_resolution,[],[f2191,f1686])).
fof(f2191,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X16)) | s(bool,X15) = s(bool,sK6(X17)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X17) | s(bool,sK5) = s(bool,X15) | s(bool,sK5) = s(bool,X16)) ) | ($spl13 | $spl56)),
  inference(superposition,[],[f2011,f227])).
fof(f2007,plain,(
  ( ! [X2,X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X2) | p(s(bool,sK6(X2))) | p(s(bool,X3))) ) | ($spl13 | $spl24)),
  inference(subsumption_resolution,[],[f1992,f175])).
fof(f1992,plain,(
  ( ! [X2,X3] : (s(bool,sK5) != s(bool,X3) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X2) | p(s(bool,sK6(X2))) | p(s(bool,X3))) ) | $spl24),
  inference(superposition,[],[f1679,f135])).
fof(f2005,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | p(s(bool,X1)) | p(s(bool,sK6(X0)))) ) | ($spl13 | $spl24)),
  inference(subsumption_resolution,[],[f1991,f175])).
fof(f1991,plain,(
  ( ! [X0,X1] : (s(bool,sK5) != s(bool,X1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | p(s(bool,X1)) | p(s(bool,sK6(X0)))) ) | $spl24),
  inference(superposition,[],[f1679,f135])).
fof(f1933,plain,(
  ( ! [X63] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X63)))) ) | $spl54),
  inference(cnf_transformation,[],[f1933_D])).
fof(f1933_D,plain,(
  ( ! [X63] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X63)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1932,plain,(
  ( ! [X64,X62,X63] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X63))) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X64))))) ) | ($spl8 | $spl13 | $spl28)),
  inference(forward_demodulation,[],[f1867,f131])).
fof(f1867,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X64)))) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X62),s(bool,X63)))))) ) | ($spl8 | $spl13 | $spl28)),
  inference(superposition,[],[f206,f1690])).
fof(f206,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X0)))) ) | ($spl8 | $spl13)),
  inference(subsumption_resolution,[],[f198,f171])).
fof(f198,plain,(
  ( ! [X0] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X0))) | p(s(bool,sK5)) | p(s(bool,X0))) ) | $spl8),
  inference(superposition,[],[f147,f135])).
fof(f1931,plain,(
  ( ! [X61,X59] : (s(bool,X59) = s(bool,X61) | s(bool,sK5) = s(bool,X61)) ) | $spl52),
  inference(cnf_transformation,[],[f1931_D])).
fof(f1931_D,plain,(
  ( ! [X61,X59] : (s(bool,X59) = s(bool,X61) | s(bool,sK5) = s(bool,X61)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1930,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X59) = s(bool,X61) | ~p(s(bool,happ(s(fun(bool,bool),X58),s(bool,X60)))) | s(bool,sK5) = s(bool,X61)) ) | ($spl13 | $spl28)),
  inference(forward_demodulation,[],[f1866,f131])).
fof(f1866,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,happ(s(fun(bool,bool),X58),s(bool,X60)))) | s(bool,X61) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X58),s(bool,X59))) | s(bool,sK5) = s(bool,X61)) ) | ($spl13 | $spl28)),
  inference(superposition,[],[f183,f1690])).
fof(f1929,plain,(
  ( ! [X57,X56,X55] : (s(bool,sK5) = s(bool,X56) | p(s(bool,happ(s(fun(bool,bool),X55),s(bool,X57))))) ) | ($spl13 | $spl28)),
  inference(forward_demodulation,[],[f1865,f131])).
fof(f1865,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,happ(s(fun(bool,bool),X55),s(bool,X57)))) | s(bool,sK5) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X55),s(bool,X56)))) ) | ($spl13 | $spl28)),
  inference(superposition,[],[f175,f1690])).
fof(f1928,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X52) = s(bool,X54) | ~p(s(bool,happ(s(fun(bool,bool),X51),s(bool,X53)))) | ~p(s(bool,X54))) ) | $spl28),
  inference(forward_demodulation,[],[f1864,f131])).
fof(f1864,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X51),s(bool,X53)))) | s(bool,X54) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X51),s(bool,X52))) | ~p(s(bool,X54))) ) | $spl28),
  inference(superposition,[],[f138,f1690])).
fof(f1927,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X48) = s(bool,X50) | ~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,X49)))) | ~p(s(bool,X50))) ) | $spl28),
  inference(forward_demodulation,[],[f1863,f131])).
fof(f1863,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,X49)))) | s(bool,X50) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X47),s(bool,X48))) | ~p(s(bool,X50))) ) | $spl28),
  inference(superposition,[],[f138,f1690])).
fof(f1926,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | s(bool,X46) = s(bool,happ(s(fun(bool,bool),X43),s(bool,X45))) | ~p(s(bool,X46))) ) | $spl28),
  inference(forward_demodulation,[],[f1862,f131])).
fof(f1862,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X46) = s(bool,happ(s(fun(bool,bool),X43),s(bool,X45))) | ~p(s(bool,X46)) | ~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X43),s(bool,X44))))) ) | $spl28),
  inference(superposition,[],[f138,f1690])).
fof(f1925,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | s(bool,X42) = s(bool,happ(s(fun(bool,bool),X39),s(bool,X41))) | ~p(s(bool,X42))) ) | $spl28),
  inference(forward_demodulation,[],[f1861,f131])).
fof(f1861,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun(bool,bool),X39),s(bool,X41))) | ~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X39),s(bool,X40)))) | ~p(s(bool,X42))) ) | $spl28),
  inference(superposition,[],[f138,f1690])).
fof(f1924,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X36) = s(bool,X38) | p(s(bool,happ(s(fun(bool,bool),X35),s(bool,X37)))) | p(s(bool,X38))) ) | $spl28),
  inference(forward_demodulation,[],[f1860,f131])).
fof(f1860,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,happ(s(fun(bool,bool),X35),s(bool,X37)))) | s(bool,X38) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X35),s(bool,X36))) | p(s(bool,X38))) ) | $spl28),
  inference(superposition,[],[f135,f1690])).
fof(f1923,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X32) = s(bool,X34) | p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X33)))) | p(s(bool,X34))) ) | $spl28),
  inference(forward_demodulation,[],[f1859,f131])).
fof(f1859,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X33)))) | s(bool,X34) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X31),s(bool,X32))) | p(s(bool,X34))) ) | $spl28),
  inference(superposition,[],[f135,f1690])).
fof(f1922,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X30) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X29))) | p(s(bool,X30))) ) | $spl28),
  inference(forward_demodulation,[],[f1858,f131])).
fof(f1858,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X30) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X29))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X28)))) | p(s(bool,X30))) ) | $spl28),
  inference(superposition,[],[f135,f1690])).
fof(f1921,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X25))) | p(s(bool,X26))) ) | $spl28),
  inference(forward_demodulation,[],[f1857,f131])).
fof(f1857,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X25))) | p(s(bool,X26)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X24))))) ) | $spl28),
  inference(superposition,[],[f135,f1690])).
fof(f1919,plain,(
  ( ! [X59,X58] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,happ(s(fun(bool,bool),X58),s(bool,X59)))))) ) | $spl50),
  inference(cnf_transformation,[],[f1919_D])).
fof(f1919_D,plain,(
  ( ! [X59,X58] : (s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,happ(s(fun(bool,bool),X58),s(bool,X59)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1918,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X60)) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,happ(s(fun(bool,bool),X58),s(bool,X59)))))) ) | ($spl8 | $spl13 | $spl28)),
  inference(forward_demodulation,[],[f1836,f131])).
fof(f1836,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X58),s(bool,X60)))) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,happ(s(fun(bool,bool),X58),s(bool,X59)))))) ) | ($spl8 | $spl13 | $spl28)),
  inference(superposition,[],[f206,f1690])).
fof(f1916,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X56)) | s(bool,X57) = s(bool,happ(s(fun(bool,bool),X54),s(bool,X55))) | s(bool,sK5) = s(bool,X57)) ) | ($spl13 | $spl28)),
  inference(forward_demodulation,[],[f1835,f131])).
fof(f1835,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X54),s(bool,X56)))) | s(bool,X57) = s(bool,happ(s(fun(bool,bool),X54),s(bool,X55))) | s(bool,sK5) = s(bool,X57)) ) | ($spl13 | $spl28)),
  inference(superposition,[],[f183,f1690])).
fof(f1915,plain,(
  ( ! [X52,X51] : (s(bool,sK5) = s(bool,happ(s(fun(bool,bool),X51),s(bool,X52)))) ) | $spl46),
  inference(cnf_transformation,[],[f1915_D])).
fof(f1915_D,plain,(
  ( ! [X52,X51] : (s(bool,sK5) = s(bool,happ(s(fun(bool,bool),X51),s(bool,X52)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1914,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X53)) | s(bool,sK5) = s(bool,happ(s(fun(bool,bool),X51),s(bool,X52)))) ) | ($spl13 | $spl28)),
  inference(forward_demodulation,[],[f1834,f131])).
fof(f1834,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X51),s(bool,X53)))) | s(bool,sK5) = s(bool,happ(s(fun(bool,bool),X51),s(bool,X52)))) ) | ($spl13 | $spl28)),
  inference(superposition,[],[f175,f1690])).
fof(f1913,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),X47),s(bool,X48))) | ~p(s(bool,X50))) ) | $spl28),
  inference(forward_demodulation,[],[f1833,f131])).
fof(f1833,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X47),s(bool,X49)))) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),X47),s(bool,X48))) | ~p(s(bool,X50))) ) | $spl28),
  inference(superposition,[],[f138,f1690])).
fof(f1912,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | s(bool,X46) = s(bool,happ(s(fun(bool,bool),X43),s(bool,X44))) | ~p(s(bool,X46))) ) | $spl28),
  inference(forward_demodulation,[],[f1832,f131])).
fof(f1832,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X43),s(bool,X45)))) | s(bool,X46) = s(bool,happ(s(fun(bool,bool),X43),s(bool,X44))) | ~p(s(bool,X46))) ) | $spl28),
  inference(superposition,[],[f138,f1690])).
fof(f1911,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) = s(bool,X42) | ~p(s(bool,X42)) | ~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X40))))) ) | $spl28),
  inference(forward_demodulation,[],[f1831,f131])).
fof(f1831,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X42) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X39),s(bool,X41))) | ~p(s(bool,X42)) | ~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X40))))) ) | $spl28),
  inference(superposition,[],[f138,f1690])).
fof(f1910,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) = s(bool,X38) | ~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,X36)))) | ~p(s(bool,X38))) ) | $spl28),
  inference(forward_demodulation,[],[f1830,f131])).
fof(f1830,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X38) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X35),s(bool,X37))) | ~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,X36)))) | ~p(s(bool,X38))) ) | $spl28),
  inference(superposition,[],[f138,f1690])).
fof(f1909,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,X33)) | s(bool,X34) = s(bool,happ(s(fun(bool,bool),X31),s(bool,X32))) | p(s(bool,X34))) ) | $spl28),
  inference(forward_demodulation,[],[f1829,f131])).
fof(f1829,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X31),s(bool,X33)))) | s(bool,X34) = s(bool,happ(s(fun(bool,bool),X31),s(bool,X32))) | p(s(bool,X34))) ) | $spl28),
  inference(superposition,[],[f135,f1690])).
fof(f1908,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X29)) | s(bool,X30) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X28))) | p(s(bool,X30))) ) | $spl28),
  inference(forward_demodulation,[],[f1828,f131])).
fof(f1828,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X29)))) | s(bool,X30) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X28))) | p(s(bool,X30))) ) | $spl28),
  inference(superposition,[],[f135,f1690])).
fof(f1907,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X25) = s(bool,X26) | p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X24)))) | p(s(bool,X26))) ) | $spl28),
  inference(forward_demodulation,[],[f1827,f131])).
fof(f1827,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X25))) | p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X24)))) | p(s(bool,X26))) ) | $spl28),
  inference(superposition,[],[f135,f1690])).
fof(f1906,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,X22) | p(s(bool,X22)) | p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X20))))) ) | $spl28),
  inference(forward_demodulation,[],[f1826,f131])).
fof(f1826,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X21))) | p(s(bool,X22)) | p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X20))))) ) | $spl28),
  inference(superposition,[],[f135,f1690])).
fof(f1904,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X28)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | ~p(s(bool,X29))) ) | $spl28),
  inference(forward_demodulation,[],[f1813,f131])).
fof(f1813,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | ~p(s(bool,X29)) | ~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X27),s(bool,X28))))) ) | $spl28),
  inference(superposition,[],[f1690,f138])).
fof(f1901,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | ~p(s(bool,X25))) ) | $spl28),
  inference(forward_demodulation,[],[f1812,f131])).
fof(f1812,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | ~p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X23),s(bool,X24)))) | ~p(s(bool,X25))) ) | $spl28),
  inference(superposition,[],[f1690,f138])).
fof(f1900,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X22))) | p(s(bool,X21))) ) | $spl28),
  inference(forward_demodulation,[],[f1811,f131])).
fof(f1811,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X22))) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X19),s(bool,X20)))) | p(s(bool,X21))) ) | $spl28),
  inference(superposition,[],[f1690,f135])).
fof(f1897,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X18))) | p(s(bool,X17))) ) | $spl28),
  inference(forward_demodulation,[],[f1810,f131])).
fof(f1810,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X17) = s(bool,happ(s(fun(bool,bool),X15),s(bool,X18))) | p(s(bool,X17)) | p(s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X15),s(bool,X16))))) ) | $spl28),
  inference(superposition,[],[f1690,f135])).
fof(f1892,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X19) = s(bool,X20) | ~p(s(bool,X19)) | ~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))))) ) | $spl28),
  inference(forward_demodulation,[],[f1760,f131])).
fof(f1760,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X19) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X17),s(bool,X20))) | ~p(s(bool,X19)) | ~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18))))) ) | $spl28),
  inference(superposition,[],[f1690,f138])).
fof(f1889,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X15) = s(bool,X16) | ~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X14)))) | ~p(s(bool,X15))) ) | $spl28),
  inference(forward_demodulation,[],[f1759,f131])).
fof(f1759,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X15) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X13),s(bool,X16))) | ~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X14)))) | ~p(s(bool,X15))) ) | $spl28),
  inference(superposition,[],[f1690,f138])).
fof(f1888,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X11) = s(bool,X12) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,X11))) ) | $spl28),
  inference(forward_demodulation,[],[f1758,f131])).
fof(f1758,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X11) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X9),s(bool,X12))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,X11))) ) | $spl28),
  inference(superposition,[],[f1690,f135])).
fof(f1885,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X7) = s(bool,X8) | p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))))) ) | $spl28),
  inference(forward_demodulation,[],[f1757,f131])).
fof(f1757,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X7) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X5),s(bool,X8))) | p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))))) ) | $spl28),
  inference(superposition,[],[f1690,f135])).
fof(f1688,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X7))) | s(bool,sK5) != s(bool,X6) | ~p(s(bool,X6))) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f1658,f173])).
fof(f1658,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X7))) | s(bool,sK5) != s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl24),
  inference(superposition,[],[f1654,f138])).
fof(f1685,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(bool,sK5) != s(bool,X4) | ~p(s(bool,X4))) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f1657,f173])).
fof(f1657,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(bool,sK5) != s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl24),
  inference(superposition,[],[f1654,f138])).
fof(f1654,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | s(bool,sK5) != s(bool,X1)) ) | $spl24),
  inference(cnf_transformation,[],[f1654_D])).
fof(f1654_D,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | s(bool,sK5) != s(bool,X1)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1587,plain,(
  ( ! [X0,X1] : (s(bool,sK5) != s(bool,X1) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | s(bool,sK5) = s(bool,X0)) ) | ($spl13 | $spl14)),
  inference(equality_factoring,[],[f260])).
fof(f1619,plain,(
  ( ! [X30,X32] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X30))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X32))) | s(bool,X30) = s(bool,X32)) ) | $spl22),
  inference(cnf_transformation,[],[f1619_D])).
fof(f1619_D,plain,(
  ( ! [X30,X32] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X30))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X32))) | s(bool,X30) = s(bool,X32)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1636,plain,(
  ( ! [X47,X48,X49] : (s(bool,X47) = s(bool,X49) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X49))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X47))) | s(bool,sK5) = s(bool,X48)) ) | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f1477,f175])).
fof(f1477,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,X48)) | s(bool,X47) = s(bool,X49) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X49))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X47))) | s(bool,sK5) = s(bool,X48)) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f208,f260])).
fof(f1600,plain,(
  ( ! [X24,X23,X25] : (s(bool,sK5) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23)))) ) | ($spl13 | $spl14)),
  inference(duplicate_literal_removal,[],[f1455])).
fof(f1455,plain,(
  ( ! [X24,X23,X25] : (s(bool,sK5) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X23))) | s(bool,sK5) = s(bool,X24)) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f260,f260])).
fof(f1608,plain,(
  ( ! [X35,X33,X34] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X33))) | s(bool,sK5) = s(bool,X34) | s(bool,X33) = s(bool,X35) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35)))) ) | ($spl13 | $spl14)),
  inference(duplicate_literal_removal,[],[f1427])).
fof(f1427,plain,(
  ( ! [X35,X33,X34] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X33))) | s(bool,sK5) = s(bool,X34) | s(bool,sK5) = s(bool,X34) | s(bool,X33) = s(bool,X35) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X35)))) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f260,f260])).
fof(f1610,plain,(
  ( ! [X26,X24,X25] : (s(bool,X25) = s(bool,X26) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26))) | s(bool,sK5) = s(bool,X24) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25)))) ) | ($spl13 | $spl14)),
  inference(duplicate_literal_removal,[],[f1407])).
fof(f1407,plain,(
  ( ! [X26,X24,X25] : (s(bool,X25) = s(bool,X26) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X26))) | s(bool,sK5) = s(bool,X24) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X25))) | s(bool,sK5) = s(bool,X24)) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f260,f260])).
fof(f1613,plain,(
  ( ! [X30,X31,X32] : (s(bool,sK5) = s(bool,X31) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X30))) | s(bool,X30) = s(bool,X32) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X32)))) ) | ($spl13 | $spl14)),
  inference(duplicate_literal_removal,[],[f1392])).
fof(f1392,plain,(
  ( ! [X30,X31,X32] : (s(bool,sK5) = s(bool,X31) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X30))) | s(bool,sK5) = s(bool,X31) | s(bool,X30) = s(bool,X32) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X32)))) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f260,f260])).
fof(f553,plain,(
  ( ! [X42,X40] : (s(bool,X40) = s(bool,X42) | s(bool,sK5) = s(bool,X40) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X42)))) ) | $spl20),
  inference(cnf_transformation,[],[f553_D])).
fof(f553_D,plain,(
  ( ! [X42,X40] : (s(bool,X40) = s(bool,X42) | s(bool,sK5) = s(bool,X40) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X42)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1285,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) = s(bool,X20) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X19) | s(bool,sK5) = s(bool,X18)) ) | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f1276,f175])).
fof(f1276,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X19)) | s(bool,X18) = s(bool,X20) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X20))) | s(bool,sK5) = s(bool,X19) | s(bool,sK5) = s(bool,X18)) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f208,f227])).
fof(f1284,plain,(
  ( ! [X17,X15,X16] : (s(bool,X15) = s(bool,X17) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X17))) | s(bool,sK5) = s(bool,X15) | s(bool,sK5) = s(bool,X16)) ) | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f1275,f175])).
fof(f1275,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X16)) | s(bool,X15) = s(bool,X17) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X17))) | s(bool,sK5) = s(bool,X15) | s(bool,sK5) = s(bool,X16)) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f208,f227])).
fof(f1206,plain,(
  $false | ($spl1 | $spl2 | $spl10 | $spl18)),
  inference(subsumption_resolution,[],[f1169,f977])).
fof(f1169,plain,(
  ( ! [X220,X219] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),X219),s(bool,X220)))) ) | ($spl1 | $spl2 | $spl10 | $spl18)),
  inference(superposition,[],[f160,f977])).
fof(f558,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,X42) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X42))) | s(bool,sK5) = s(bool,X41) | s(bool,sK5) = s(bool,X40)) ) | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f378,f175])).
fof(f378,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X41)) | s(bool,X40) = s(bool,X42) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X42))) | s(bool,sK5) = s(bool,X41) | s(bool,sK5) = s(bool,X40)) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f208,f227])).
fof(f551,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,X42) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X42))) | s(bool,sK5) = s(bool,X40) | s(bool,sK5) = s(bool,X41)) ) | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f353,f175])).
fof(f353,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X41)) | s(bool,X40) = s(bool,X42) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X42))) | s(bool,sK5) = s(bool,X40) | s(bool,sK5) = s(bool,X41)) ) | ($spl13 | $spl14)),
  inference(superposition,[],[f208,f227])).
fof(f165,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | ~p(s(bool,sK5)) | ~p(s(bool,X3))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f161,f138])).
fof(f164,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | ~p(s(bool,X2)) | ~p(s(bool,sK5))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f161,f138])).
fof(f163,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | p(s(bool,X1)) | p(s(bool,sK5))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f161,f135])).
fof(f162,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))) | p(s(bool,sK5)) | p(s(bool,X0))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f161,f135])).
fof(f159,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f158,f131])).
fof(f158,plain,(
  s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),sK4),s(sK0,sK2))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f157,f141])).
fof(f157,plain,(
  s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK2))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f140,f143])).
fof(f139,plain,(
  s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3))) | $spl0),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f141,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | $spl2),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f125,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  (s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3))) | (s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) & s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) & s(sK0,sK2) = s(sK0,sK3))) & (s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3))) | (! [X5] : (s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5)))) & (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',sK1) | s(sK0,sK2) != s(sK0,sK3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f122])).
fof(f122,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) | ? [X6] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X6))) & s(X0,X2) = s(X0,happ(s(fun(bool,X0),X4),s(bool,X6)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) & s(X0,X2) = s(X0,X3))) & (s(X0,X2) != s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) | (! [X5] : (s('type.semanticPrimitives.eq_result',X1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(X0,X2) != s(X0,happ(s(fun(bool,X0),X4),s(bool,X5)))) & (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',X1) | s(X0,X2) != s(X0,X3)))))),
  inference(rectify,[],[f121])).
fof(f121,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) | ? [X5] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) & s(X0,X2) = s(X0,happ(s(fun(bool,X0),X4),s(bool,X5)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) & s(X0,X2) = s(X0,X3))) & (s(X0,X2) != s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) | (! [X5] : (s('type.semanticPrimitives.eq_result',X1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(X0,X2) != s(X0,happ(s(fun(bool,X0),X4),s(bool,X5)))) & (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',X1) | s(X0,X2) != s(X0,X3)))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) | (? [X5] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) & s(X0,X2) = s(X0,happ(s(fun(bool,X0),X4),s(bool,X5)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) & s(X0,X2) = s(X0,X3)))) & (s(X0,X2) != s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) | (! [X5] : (s('type.semanticPrimitives.eq_result',X1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(X0,X2) != s(X0,happ(s(fun(bool,X0),X4),s(bool,X5)))) & (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',X1) | s(X0,X2) != s(X0,X3)))))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ? [X0,X1,X2,X3,X4] : (s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) <~> (? [X5] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) & s(X0,X2) = s(X0,happ(s(fun(bool,X0),X4),s(bool,X5)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) & s(X0,X2) = s(X0,X3))))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ~! [X0,X1,X2,X3,X4] : (s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) <=> (? [X5] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) & s(X0,X2) = s(X0,happ(s(fun(bool,X0),X4),s(bool,X5)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) & s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X1),s(fun(bool,X0),X4),s(X0,X3))) <=> (? [X5] : (s('type.semanticPrimitives.eq_result',X1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) & s(X0,X2) = s(X0,happ(s(fun(bool,X0),X4),s(bool,X5)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) & s(X0,X2) = s(X0,X3))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (s(X6,X8) = s(X6,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X7),s(fun(bool,X6),X10),s(X6,X9))) <=> (? [X11] : (s('type.semanticPrimitives.eq_result',X7) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X11))) & s(X6,X8) = s(X6,happ(s(fun(bool,X6),X10),s(bool,X11)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X7) & s(X6,X8) = s(X6,X9))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (s(X6,X8) = s(X6,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X7),s(fun(bool,X6),X10),s(X6,X9))) <=> (? [X11] : (s('type.semanticPrimitives.eq_result',X7) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X11))) & s(X6,X8) = s(X6,happ(s(fun(bool,X6),X10),s(bool,X11)))) | (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X7) & s(X6,X8) = s(X6,X9))))),
  file('thm.closProps.eqs|split|6','thm.closProps.eqs|split|6')).
fof(f147,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | $spl8),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f126,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK1) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f123])).
fof(f143,plain,(
  s(sK0,sK2) = s(sK0,sK3) | $spl4),
  inference(cnf_transformation,[],[f143_D])).
fof(f143_D,plain,(
  s(sK0,sK2) = s(sK0,sK3) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f127,plain,(
  s(sK0,sK2) = s(sK0,sK3) | s('type.semanticPrimitives.eq_result',sK1) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK5))) | s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f123])).
fof(f128,plain,(
  s(sK0,sK2) = s(sK0,sK3) | s(sK0,sK2) = s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,sK5))) | s(sK0,sK2) = s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f123])).
fof(f140,plain,(
  s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3))) | $spl1),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f145,plain,(
  ( ! [X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5)))) ) | $spl6),
  inference(cnf_transformation,[],[f145_D])).
fof(f145_D,plain,(
  ( ! [X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f129,plain,(
  ( ! [X5] : (s(sK0,sK2) != s(sK0,happ(s(fun(bool,sK0),sK4),s(bool,X5))) | s('type.semanticPrimitives.eq_result',sK1) != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))) | s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f142,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',sK1) | $spl3),
  inference(cnf_transformation,[],[f142_D])).
fof(f142_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f144,plain,(
  s(sK0,sK2) != s(sK0,sK3) | $spl5),
  inference(cnf_transformation,[],[f144_D])).
fof(f144_D,plain,(
  s(sK0,sK2) != s(sK0,sK3) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f130,plain,(
  s(sK0,sK2) != s(sK0,sK3) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result',sK1) | s(sK0,sK2) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK1),s(fun(bool,sK0),sK4),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f123])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 469
% Passive clauses: 4410
% Generated clauses: 26818
% Final active clauses: 41
% Final passive clauses: 589
% Input formulas: 105
% Initial clauses: 14
% 
% Selected by SInE selection: 6
% SInE iterations: 1
% 
% Duplicate literals: 857
% Trivial inequalities: 354
% Fw subsumption resolutions: 2637
% Fw demodulations: 3377
% Bw demodulations: 118
% 
% Simple tautologies: 709
% Equational tautologies: 65
% Forward subsumptions: 14002
% Fw demodulations to eq. taut.: 209
% Bw demodulations to eq. taut.: 19
% 
% Binary resolution: 1124
% Factoring: 7
% Forward superposition: 8902
% Backward superposition: 7178
% Self superposition: 513
% Equality factoring: 100
% Equality resolution: 32
% 
% Split clauses: 548
% Split components: 1103
% SAT solver clauses: 656
% SAT solver binary clauses: 23
% 
% Sat splits: 548
% Sat splitting refutations: 93
% 
% Memory used [KB]: 9978
% Time elapsed: 0.519 s
% ------------------------------
% ------------------------------
% Success in time 12.26 s
