Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (82059)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 237
% Passive clauses: 2947
% Generated clauses: 21328
% Final active clauses: 204
% Final passive clauses: 2578
% Input formulas: 261
% Initial clauses: 419
% 
% Duplicate literals: 249
% Trivial inequalities: 343
% Fw subsumption resolutions: 1161
% Fw demodulations: 3558
% Bw demodulations: 32
% 
% Simple tautologies: 290
% Equational tautologies: 28
% Forward subsumptions: 12605
% Fw demodulations to eq. taut.: 7
% 
% Binary resolution: 154
% Factoring: 1
% Forward superposition: 7429
% Backward superposition: 7772
% Self superposition: 112
% Equality factoring: 41
% Equality resolution: 22
% 
% Split clauses: 71
% Split components: 144
% SAT solver clauses: 133
% SAT solver unit clauses: 8
% SAT solver binary clauses: 61
% 
% Sat splits: 71
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 133
% TWLsolver calls for satisfiability: 24
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% (82073)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 239
% Passive clauses: 1465
% Generated clauses: 22901
% Final active clauses: 151
% Final passive clauses: 789
% Input formulas: 261
% Initial clauses: 419
% 
% Duplicate literals: 202
% Trivial inequalities: 234
% Fw subsumption resolutions: 1375
% Fw demodulations: 5144
% Bw demodulations: 488
% Condensations: 310
% Global subsumptions: 20
% 
% Simple tautologies: 312
% Equational tautologies: 12
% Forward subsumptions: 12044
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 18
% 
% Binary resolution: 178
% Factoring: 4
% Forward superposition: 7903
% Backward superposition: 4844
% Self superposition: 123
% Equality factoring: 18
% Equality resolution: 29
% 
% Split clauses: 25
% Split components: 50
% SAT solver clauses: 2052
% SAT solver unit clauses: 580
% SAT solver binary clauses: 93
% SAT solver learnt clauses: 40
% SAT solver learnt literals: 14
% 
% Sat splits: 25
% Sat splitting refutations: 29
% 
% TWLsolver clauses: 1972
% TWLsolver calls for satisfiability: 4335
% 
% Memory used [KB]: 17782
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvp_to_wordProps
% (82089)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1001
% Passive clauses: 8348
% Generated clauses: 285529
% Final active clauses: 438
% Final passive clauses: 2187
% Input formulas: 261
% Initial clauses: 419
% Discarded non-redundant clauses: 30186
% 
% Duplicate literals: 1332
% Trivial inequalities: 2665
% Fw subsumption resolutions: 9571
% Fw demodulations: 85847
% Bw demodulations: 3973
% Condensations: 1687
% Global subsumptions: 1651
% 
% Simple tautologies: 928
% Equational tautologies: 500
% Forward subsumptions: 138203
% Fw demodulations to eq. taut.: 6209
% Bw demodulations to eq. taut.: 70
% 
% Unit resulting resolution: 58992
% Factoring: 1
% Forward superposition: 64750
% Backward superposition: 49238
% Self superposition: 388
% Equality factoring: 58
% Equality resolution: 127
% 
% Split clauses: 108
% Split components: 218
% SAT solver clauses: 12486
% SAT solver unit clauses: 1439
% SAT solver binary clauses: 5935
% 
% Sat splits: 108
% Sat splitting refutations: 2269
% 
% Memory used [KB]: 189975
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f3661,plain,(
  $false),
  inference(subsumption_resolution,[],[f3584,f836])).
fof(f836,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))))))),
  inference(resolution,[],[f567,f602])).
fof(f602,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X1))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X0)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X0))))))) )),
  inference(cnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X0)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X1))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X0)))))))),
  inference(ennf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X0)))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X1))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X0)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X12] : (p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X12),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X11)))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X12))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X11)))))))),
  file('thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb','thm.bvp_to_wordProps.bvp_to_wordProps_53_prove_ltb')).
fof(f567,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK1))))))),
  inference(cnf_transformation,[],[f473])).
fof(f473,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',sK2),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),sK1)))))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f430])).
fof(f430,plain,(
  ? [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X0)))))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X0)))))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X1),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X0)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X0)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X7),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X6)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X6)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.bool.IN_2'(s('type.bvlSem.v',X7),s(fun('type.bvlSem.v',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.bvlSem.v'),X6)))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),X6)))))))))),
  file('thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb','thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb')).
fof(f3584,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))))))),
  inference(resolution,[],[f1005,f570])).
fof(f570,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f474])).
fof(f474,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f163])).
fof(f163,axiom,(
  ! [X77,X31] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X77),s('type.num.num',X31)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X31),s('type.num.num',X77)))))),
  file('thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb','thm.arithmetic.NOT_LESS')).
fof(f1005,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))),s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))))))),
  inference(forward_demodulation,[],[f916,f577])).
fof(f577,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f191])).
fof(f191,axiom,(
  ! [X77,X31] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X77),s('type.num.num',X31))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X77))),s('type.num.num',X31)))),
  file('thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb','thm.arithmetic.LESS_EQ')).
fof(f916,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))))),s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))))))),
  inference(resolution,[],[f835,f569])).
fof(f569,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f474])).
fof(f835,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))))))))),
  inference(forward_demodulation,[],[f834,f600])).
fof(f600,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X31] : s('type.num.num',X31) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X31)))),
  file('thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb','thm.arithmetic.ADD|split|0')).
fof(f834,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))))))))))),
  inference(forward_demodulation,[],[f833,f592])).
fof(f592,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X77,X31] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X77))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X77),s('type.num.num',X31)))),
  file('thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb','thm.arithmetic.ADD_SYM')).
fof(f833,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))),s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f832,f591])).
fof(f591,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X31,X77] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X77),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X77),s('type.num.num',X31)))))),
  file('thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f832,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f831,f621])).
fof(f621,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f85])).
fof(f85,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.bvp_to_wordProps.bvp_to_wordProps_9|48|_prove_ltb','thm.numeral.numeral_distrib|split|13')).
fof(f831,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f568,f592])).
fof(f568,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',sK2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.bvlSem.v|49|_size_1'(s('type.list.list'('type.bvlSem.v'),sK1))))))))),
  inference(cnf_transformation,[],[f473])).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 218
% Passive clauses: 1268
% Generated clauses: 2880
% Final active clauses: 199
% Final passive clauses: 488
% Input formulas: 261
% Initial clauses: 260
% 
% Selected by SInE selection: 137
% SInE iterations: 1
% 
% Duplicate literals: 127
% Trivial inequalities: 25
% Fw subsumption resolutions: 125
% Fw demodulations: 106
% 
% Simple tautologies: 205
% Equational tautologies: 3
% Forward subsumptions: 867
% 
% Binary resolution: 980
% Forward superposition: 1431
% 
% Split clauses: 163
% Split components: 347
% SAT solver clauses: 177
% SAT solver unit clauses: 13
% SAT solver binary clauses: 20
% 
% Sat splits: 163
% Sat splitting refutations: 13
% 
% Memory used [KB]: 3582
% Time elapsed: 0.078 s
% ------------------------------
% ------------------------------
% Success in time 11.87 s
