Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.labSem
% (13983)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 66
% Passive clauses: 2565
% Generated clauses: 11479
% Final active clauses: 62
% Final passive clauses: 2364
% Input formulas: 69
% Initial clauses: 74
% 
% Duplicate literals: 647
% Trivial inequalities: 95
% Fw subsumption resolutions: 1452
% Fw demodulations: 250
% Bw demodulations: 1
% 
% Simple tautologies: 195
% Equational tautologies: 112
% Forward subsumptions: 4943
% 
% Binary resolution: 99
% Factoring: 2
% Forward superposition: 3482
% Backward superposition: 4527
% Self superposition: 462
% Equality factoring: 328
% Equality resolution: 10
% 
% Split clauses: 379
% Split components: 760
% SAT solver clauses: 1279
% SAT solver unit clauses: 26
% SAT solver binary clauses: 192
% 
% Sat splits: 379
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1277
% TWLsolver calls for satisfiability: 132
% 
% Memory used [KB]: 8955
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.labSem
% (13997)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 39
% Passive clauses: 376
% Generated clauses: 5255
% Final active clauses: 38
% Final passive clauses: 324
% Input formulas: 69
% Initial clauses: 74
% 
% Duplicate literals: 635
% Trivial inequalities: 37
% Fw subsumption resolutions: 173
% Fw demodulations: 74
% Bw demodulations: 13
% Condensations: 965
% 
% Simple tautologies: 106
% Equational tautologies: 72
% Forward subsumptions: 1810
% 
% Binary resolution: 63
% Factoring: 2
% Forward superposition: 1125
% Backward superposition: 1599
% Self superposition: 218
% Equality factoring: 166
% Equality resolution: 7
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 664
% SAT solver unit clauses: 132
% SAT solver binary clauses: 39
% 
% Sat splits: 9
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 664
% TWLsolver calls for satisfiability: 144
% 
% Memory used [KB]: 8315
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.labSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.labSem
% SZS output start Proof for thm.labSem
fof(f6612,plain,(
  $false),
  inference(subsumption_resolution,[],[f6611,f211])).
fof(f211,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.labSem.state_fupdfupds_comp|split|7',aHL_TRUTH)).
fof(f6611,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f6609,f1214])).
fof(f1214,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(equality_resolution,[],[f612])).
fof(f612,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(condensation,[],[f611])).
fof(f611,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f608,f287])).
fof(f287,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f212,f213])).
fof(f213,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.labSem.state_fupdfupds_comp|split|7',aHL_BOOL_CASES)).
fof(f212,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.labSem.state_fupdfupds_comp|split|7',aHL_FALSITY)).
fof(f608,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f244,f285])).
fof(f285,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f213,f213])).
fof(f244,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X45,X46,X47] : (p(s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),'=_0'),s(X45,X46))),s(X45,X47)))) <=> s(X45,X46) = s(X45,X47))),
  file('thm.labSem.state_fupdfupds_comp|split|7',arity_0_HL_EQ)).
fof(f6609,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.labSem.state'(sK1,sK2)),bool),happ(s(fun(fun(sK0,'type.labSem.state'(sK1,sK2)),fun(fun(sK0,'type.labSem.state'(sK1,sK2)),bool)),'=_0'),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))))))),
  inference(backward_demodulation,[],[f6578,f1547])).
fof(f1547,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.labSem.state'(sK1,sK2)),bool),happ(s(fun(fun(sK0,'type.labSem.state'(sK1,sK2)),fun(fun(sK0,'type.labSem.state'(sK1,sK2)),bool)),'=_0'),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),sK5))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),sK4))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))))))))),
  inference(unit_resulting_resolution,[],[f210,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f209])).
fof(f210,plain,(
  s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))) != s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),sK5))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),sK4))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3)))))),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))) != s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),sK5))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),sK4))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f201])).
fof(f201,plain,(
  ? [X0,X1,X2,X3,X4,X5] : s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X5),s(fun('type.num.num','type.num.num'),X4))))),s(fun(X0,'type.labSem.state'(X1,X2)),X3))) != s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X5))),s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X4))),s(fun(X0,'type.labSem.state'(X1,X2)),X3)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X5),s(fun('type.num.num','type.num.num'),X4))))),s(fun(X0,'type.labSem.state'(X1,X2)),X3))) = s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X5))),s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X4))),s(fun(X0,'type.labSem.state'(X1,X2)),X3)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5] : s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X5),s(fun('type.num.num','type.num.num'),X4))))),s(fun(X0,'type.labSem.state'(X1,X2)),X3))) = s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X5))),s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X4))),s(fun(X0,'type.labSem.state'(X1,X2)),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11] : s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X11),s(fun('type.num.num','type.num.num'),X10))))),s(fun(X6,'type.labSem.state'(X7,X8)),X9))) = s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X11))),s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X10))),s(fun(X6,'type.labSem.state'(X7,X8)),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11] : s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X11),s(fun('type.num.num','type.num.num'),X10))))),s(fun(X6,'type.labSem.state'(X7,X8)),X9))) = s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X11))),s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X10))),s(fun(X6,'type.labSem.state'(X7,X8)),X9)))))),
  file('thm.labSem.state_fupdfupds_comp|split|7','thm.labSem.state_fupdfupds_comp|split|7')).
fof(f6578,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X4,'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X2),s(fun('type.num.num','type.num.num'),X3))))),s(fun(X4,'type.labSem.state'(X0,X1)),X5))) = s(fun(X4,'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X2))),s(fun(X4,'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X3))),s(fun(X4,'type.labSem.state'(X0,X1)),X5)))))) )),
  inference(superposition,[],[f251,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))))) = s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X3))),s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X2)))))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3] : s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))))) = s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X3))),s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X2)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3] : s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))))) = s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X3))),s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X2)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X8] : ! [X10,X11] : s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X11),s(fun('type.num.num','type.num.num'),X10))))) = s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X11))),s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_pc_fupd_1'(s(fun('type.num.num','type.num.num'),X10)))))),
  file('thm.labSem.state_fupdfupds_comp|split|7','thm.labSem.state_fupdfupds_comp|split|6')).
fof(f251,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7,X18,X16] : ! [X11,X10,X9] : s(fun(X18,X6),'const.combin.o_2'(s(fun(X7,X6),X11),s(fun(X18,X7),'const.combin.o_2'(s(fun(X16,X7),X10),s(fun(X18,X16),X9))))) = s(fun(X18,X6),'const.combin.o_2'(s(fun(X16,X6),'const.combin.o_2'(s(fun(X7,X6),X11),s(fun(X16,X7),X10))),s(fun(X18,X16),X9)))),
  file('thm.labSem.state_fupdfupds_comp|split|7','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.labSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 74
% Passive clauses: 659
% Generated clauses: 6236
% Final active clauses: 64
% Final passive clauses: 503
% Input formulas: 69
% Initial clauses: 74
% 
% Duplicate literals: 396
% Trivial inequalities: 34
% Fw subsumption resolutions: 142
% Fw demodulations: 880
% Bw demodulations: 63
% Condensations: 342
% Global subsumptions: 1
% 
% Simple tautologies: 89
% Equational tautologies: 110
% Forward subsumptions: 3539
% Fw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 443
% Factoring: 1
% Forward superposition: 1641
% Backward superposition: 2003
% Self superposition: 154
% Equality factoring: 54
% Equality resolution: 1
% 
% Split clauses: 22
% Split components: 49
% SAT solver clauses: 707
% SAT solver unit clauses: 205
% SAT solver binary clauses: 183
% 
% Sat splits: 22
% Sat splitting refutations: 1
% 
% Memory used [KB]: 11001
% Time elapsed: 0.340 s
% ------------------------------
% ------------------------------
% Success in time 1.674 s
