Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.SplayHeap
% (57499)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 564
% Passive clauses: 3204
% Generated clauses: 15130
% Final active clauses: 554
% Final passive clauses: 2370
% Input formulas: 35
% Initial clauses: 136
% 
% Duplicate literals: 729
% Trivial inequalities: 79
% Fw subsumption resolutions: 1510
% Fw demodulations: 377
% Bw demodulations: 10
% 
% Simple tautologies: 202
% Equational tautologies: 65
% Forward subsumptions: 5996
% 
% Binary resolution: 4515
% Factoring: 1116
% Forward superposition: 1885
% Backward superposition: 4369
% Self superposition: 110
% Equality factoring: 110
% Equality resolution: 16
% 
% Split clauses: 766
% Split components: 1638
% SAT solver clauses: 1034
% SAT solver unit clauses: 14
% SAT solver binary clauses: 104
% 
% Sat splits: 766
% 
% TWLsolver clauses: 1032
% TWLsolver calls for satisfiability: 300
% 
% Memory used [KB]: 7036
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.SplayHeap
% (57523)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 316
% Passive clauses: 1623
% Generated clauses: 8226
% Final active clauses: 222
% Final passive clauses: 680
% Input formulas: 35
% Initial clauses: 136
% 
% Duplicate literals: 151
% Trivial inequalities: 30
% Fw subsumption resolutions: 814
% Fw demodulations: 258
% Bw demodulations: 530
% Condensations: 198
% Global subsumptions: 13
% 
% Simple tautologies: 187
% Equational tautologies: 17
% Forward subsumptions: 3231
% Bw demodulations to eq. taut.: 145
% 
% Binary resolution: 2170
% Factoring: 182
% Forward superposition: 610
% Backward superposition: 588
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 5
% 
% Split clauses: 220
% Split components: 445
% SAT solver clauses: 1686
% SAT solver unit clauses: 151
% SAT solver binary clauses: 76
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 220
% Sat splitting refutations: 24
% 
% TWLsolver clauses: 1634
% TWLsolver calls for satisfiability: 1338
% 
% Memory used [KB]: 6780
% Time elapsed: 0.730 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.SplayHeap
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.SplayHeap
% SZS output start Proof for thm.SplayHeap
fof(f3287,plain,(
  $false),
  inference(global_subsumption,[],[f195,f194,f193,f196,f203,f211,f210,f247,f246,f341,f265,f262,f261,f260,f270,f271,f278,f277,f276,f275,f274,f273,f272,f282,f281,f280,f279,f286,f285,f284,f283,f290,f289,f288,f287,f294,f293,f292,f291,f301,f300,f299,f330,f297,f296,f295,f305,f304,f303,f302,f309,f308,f307,f306,f318,f317,f316,f329,f314,f313,f312,f322,f321,f320,f319,f326,f325,f324,f323,f342,f343,f344,f345,f346,f347,f348,f349,f350,f351,f352,f353,f354,f355,f356,f357,f358,f368,f367,f366,f365,f364,f363,f362,f361,f360,f359,f376,f380,f388,f389,f392,f397,f404,f409,f410,f413,f421,f441,f442,f443,f444,f445,f459,f460,f463,f466,f474,f477,f481,f492,f495,f510,f508,f513,f521,f564,f567,f565,f577,f575,f590,f586,f597,f608,f606,f628,f647,f655,f668,f676,f680,f686,f695,f713,f726,f734,f738,f769,f770,f779,f788,f796,f808,f816,f820,f835,f838,f836,f856,f859,f876,f875,f867,f865,f864,f883,f916,f906,f905,f959,f954,f1026,f1037,f1041,f1043,f1049,f1064,f1077,f1087,f1117,f1179,f1221,f1222,f1471,f1475,f1478,f1482,f1485,f1488,f1491,f1494,f1497,f1500,f1503,f1506,f1509,f1512,f1515,f1518,f1521,f1524,f1527,f1530,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1550,f1551,f1552,f1553,f1554,f1289,f1559,f1561,f1563,f1566,f1568,f1570,f1572,f1574,f1576,f1578,f1580,f1582,f1584,f1586,f1588,f1590,f1592,f1594,f1596,f1598,f240,f239,f238,f1555,f1556,f1557,f1065,f1066,f1025,f1178,f1288,f989,f988,f340,f215,f214,f2085,f2086,f2087,f2088,f2089,f2108,f2109,f2123,f2533,f2543,f2554,f2570,f2586,f2983,f3013,f3011,f3049,f3051,f3212,f3209,f3206,f3203,f3202,f3201,f3200,f3199,f3198,f3197,f3196,f3195,f3194,f3193,f3191,f3190,f3189,f3184,f3177,f3176,f3171,f3170,f3143,f3142,f3141,f3136,f3131,f3130,f3125,f3116,f3111,f3107,f3105,f3102,f3100,f3099,f3098,f3097,f3092,f3091,f3084,f3079,f3078,f3071,f3070,f3069,f3068,f3067,f3066,f3063,f3062,f3061,f3060,f3058,f3057,f3056,f3055,f3054,f3053,f3286,f3282])).
fof(f3282,plain,(
  p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK17))))))),
  inference(subsumption_resolution,[],[f3275,f196])).
fof(f3275,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK17))))))),
  inference(superposition,[],[f1066,f380])).
fof(f3286,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK17)))))) | s(bool,X2) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f3277,f196])).
fof(f3277,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK17)))))) | s(bool,X2) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1066,f521])).
fof(f3053,plain,(
  p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18))),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18))))))))),
  inference(unit_resulting_resolution,[],[f3049,f3049,f240])).
fof(f3054,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f3049,f211])).
fof(f3055,plain,(
  s(bool,'F_0') != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f3049,f214])).
fof(f3056,plain,(
  ( ! [X0] : (sP0(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) )),
  inference(unit_resulting_resolution,[],[f3049,f246])).
fof(f3057,plain,(
  ( ! [X0] : (sP0('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0)) )),
  inference(unit_resulting_resolution,[],[f3049,f247])).
fof(f3058,plain,(
  ~sP1('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'T_0')),
  inference(unit_resulting_resolution,[],[f196,f3049,f261])).
fof(f3060,plain,(
  ~sP1('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f196,f3049,f261])).
fof(f3061,plain,(
  ~sP1('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f3049,f3049,f261])).
fof(f3062,plain,(
  sP1('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0')),
  inference(unit_resulting_resolution,[],[f340,f3049,f262])).
fof(f3063,plain,(
  sP1('F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f340,f3049,f265])).
fof(f3066,plain,(
  ~sP5('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f3049,f275])).
fof(f3067,plain,(
  ~sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0')),
  inference(unit_resulting_resolution,[],[f3049,f340,f3049,f275])).
fof(f3068,plain,(
  sP5('T_0','T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f351,f347,f344,f3049,f276])).
fof(f3069,plain,(
  sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f351,f362,f344,f3049,f276])).
fof(f3070,plain,(
  sP5('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f364,f347,f361,f3049,f276])).
fof(f3071,plain,(
  sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f364,f362,f361,f3049,f276])).
fof(f3078,plain,(
  ~sP4('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f3049,f279])).
fof(f3079,plain,(
  ~sP4('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0')),
  inference(unit_resulting_resolution,[],[f3049,f340,f3049,f279])).
fof(f3084,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) )),
  inference(unit_resulting_resolution,[],[f3049,f280])).
fof(f3091,plain,(
  ~sP3('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f3049,f283])).
fof(f3092,plain,(
  ~sP3('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0')),
  inference(unit_resulting_resolution,[],[f3049,f340,f3049,f283])).
fof(f3097,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) )),
  inference(unit_resulting_resolution,[],[f3049,f284])).
fof(f3098,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) )),
  inference(unit_resulting_resolution,[],[f3049,f288])).
fof(f3099,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X1)) )),
  inference(unit_resulting_resolution,[],[f3049,f289])).
fof(f3100,plain,(
  ( ! [X0,X1] : (sP2('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f3049,f290])).
fof(f3102,plain,(
  sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f380,f3049,f292])).
fof(f3105,plain,(
  sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f380,f3049,f292])).
fof(f3107,plain,(
  ~sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f376,f3049,f293])).
fof(f3111,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) )),
  inference(unit_resulting_resolution,[],[f340,f3049,f295])).
fof(f3116,plain,(
  ( ! [X0] : (~sP7('F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0)) )),
  inference(unit_resulting_resolution,[],[f340,f3049,f296])).
fof(f3125,plain,(
  ~sP6('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f340,f3049,f302])).
fof(f3130,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) )),
  inference(unit_resulting_resolution,[],[f3049,f303])).
fof(f3131,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X1)) )),
  inference(unit_resulting_resolution,[],[f3049,f304])).
fof(f3136,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) )),
  inference(unit_resulting_resolution,[],[f3049,f3049,f307])).
fof(f3141,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0)) )),
  inference(unit_resulting_resolution,[],[f3049,f3049,f308])).
fof(f3142,plain,(
  ~sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f340,f3049,f309])).
fof(f3143,plain,(
  ( ! [X0] : (~sP9('F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0)) )),
  inference(unit_resulting_resolution,[],[f340,f3049,f313])).
fof(f3170,plain,(
  ~sP8('T_0','F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f340,f196,f3049,f319])).
fof(f3171,plain,(
  ~sP8('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f340,f3049,f3049,f319])).
fof(f3176,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X1)) )),
  inference(unit_resulting_resolution,[],[f3049,f321])).
fof(f3177,plain,(
  sP9('F_0','F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f340,f340,f3049,f323])).
fof(f3184,plain,(
  ( ! [X0] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0)) )),
  inference(unit_resulting_resolution,[],[f3049,f3049,f325])).
fof(f3189,plain,(
  ~sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f3049,f326])).
fof(f3190,plain,(
  ~sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f3049,f340,f3049,f326])).
fof(f3191,plain,(
  ( ! [X0] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f356,f3049,f329])).
fof(f3193,plain,(
  ( ! [X0] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f3049,f329])).
fof(f3194,plain,(
  ( ! [X0] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f368,f3049,f329])).
fof(f3195,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f353,f3049,f330])).
fof(f3196,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f366,f3049,f330])).
fof(f3197,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f354,f3049,f330])).
fof(f3198,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f367,f3049,f330])).
fof(f3199,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) )),
  inference(unit_resulting_resolution,[],[f3049,f443])).
fof(f3200,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0)) )),
  inference(unit_resulting_resolution,[],[f3049,f444])).
fof(f3201,plain,(
  ( ! [X0] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),X0)) )),
  inference(unit_resulting_resolution,[],[f3049,f477])).
fof(f3202,plain,(
  ~sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f340,f3049,f628])).
fof(f3203,plain,(
  ~sP4('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f340,f3049,f655])).
fof(f3206,plain,(
  ~sP3('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f340,f3049,f676])).
fof(f3209,plain,(
  ~sP8('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))),'F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f340,f3049,f816])).
fof(f3212,plain,(
  ~sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f340,f3049,f856])).
fof(f3051,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) | s(bool,X2) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f3044,f196])).
fof(f3044,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) | s(bool,X2) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1065,f521])).
fof(f3049,plain,(
  p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18))))))),
  inference(subsumption_resolution,[],[f3042,f196])).
fof(f3042,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18))))))),
  inference(superposition,[],[f1065,f380])).
fof(f3011,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f366,f513])).
fof(f3013,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f367,f513])).
fof(f2983,plain,(
  ( ! [X0] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f368,f495])).
fof(f2586,plain,(
  ~sP8('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f2570,f318])).
fof(f2570,plain,(
  ~sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f856])).
fof(f2554,plain,(
  sP9('F_0','F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f340,f340,f835])).
fof(f2543,plain,(
  ~sP9('T_0','F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f2533,f314])).
fof(f2533,plain,(
  ~sP8('T_0','F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f340,f196,f816])).
fof(f2123,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK19(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK19(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f2122])).
fof(f2122,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f2115,f203])).
fof(f2115,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f989,f521])).
fof(f2109,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK19(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f989,f521])).
fof(f2108,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f989,f521])).
fof(f2089,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK19(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f988,f203])).
fof(f2088,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK19(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f988,f521])).
fof(f2087,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK19(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f988,f521])).
fof(f2086,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK19(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f988,f521])).
fof(f2085,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK19(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f988,f521])).
fof(f214,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X17] : (s(bool,'F_0') = s(bool,X17) <=> ~p(s(bool,X17)))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f215,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f138])).
fof(f340,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f199,f196])).
fof(f199,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f988,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK19(X0,X1,X2,bool))) )),
  inference(superposition,[],[f270,f203])).
fof(f989,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK19(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK19(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f270,f203])).
fof(f1288,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'T_0') != s(bool,f(s(X127,sK19(X128,X129,bool,X127)))) | s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,X130) = s(bool,g(s(X127,sK19(X128,X129,bool,X127)))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f270,f521])).
fof(f1178,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,X134) != s(bool,f(s(X131,sK19(X132,X133,bool,X131)))) | s(fun(X131,bool),X132) = s(fun(X131,bool),X133) | s(bool,'T_0') = s(bool,X134) | s(bool,'T_0') = s(bool,g(s(X131,sK19(X132,X133,bool,X131))))) )),
  inference(superposition,[],[f270,f521])).
fof(f1025,plain,(
  ( ! [X0] : (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),X0)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))),s(fun(sK11,'type.num.num'),X0))))))) )),
  inference(subsumption_resolution,[],[f1022,f196])).
fof(f1022,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),X0)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))),s(fun(sK11,'type.num.num'),X0))))))) )),
  inference(superposition,[],[f240,f380])).
fof(f1066,plain,(
  ( ! [X1] : (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X1),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK17))))))) )),
  inference(superposition,[],[f238,f1043])).
fof(f1065,plain,(
  ( ! [X0] : (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X0),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),X0),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18))))))) )),
  inference(superposition,[],[f239,f1043])).
fof(f1557,plain,(
  ( ! [X111,X114,X112,X115,X113] : (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X111,bool),X112),s(fun(X111,'type.num.num'),X114)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X111,bool),X112),s(fun(X111,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X111,'type.num.num'),X113),s(fun(X111,'type.num.num'),X114)))))) | s(bool,X115) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X111,bool),X112),s(fun(X111,'type.num.num'),X113))) | s(bool,'T_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f1286,f196])).
fof(f1286,plain,(
  ( ! [X111,X114,X112,X115,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X111,bool),X112),s(fun(X111,'type.num.num'),X114)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X111,bool),X112),s(fun(X111,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X111,'type.num.num'),X113),s(fun(X111,'type.num.num'),X114)))))) | s(bool,X115) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X111,bool),X112),s(fun(X111,'type.num.num'),X113))) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f240,f521])).
fof(f1556,plain,(
  ( ! [X109,X107,X110,X108,X106] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X106,bool),X107),s(fun(X106,'type.num.num'),X108)))) | s(bool,X110) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X106,bool),X107),s(fun(X106,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X106,'type.num.num'),X108),s(fun(X106,'type.num.num'),X109))))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f1285,f196])).
fof(f1285,plain,(
  ( ! [X109,X107,X110,X108,X106] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X106,bool),X107),s(fun(X106,'type.num.num'),X108)))) | s(bool,X110) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X106,bool),X107),s(fun(X106,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X106,'type.num.num'),X108),s(fun(X106,'type.num.num'),X109))))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f238,f521])).
fof(f1555,plain,(
  ( ! [X103,X101,X105,X102,X104] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X101,bool),X102),s(fun(X101,'type.num.num'),X104)))) | s(bool,X105) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X101,bool),X102),s(fun(X101,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X101,'type.num.num'),X103),s(fun(X101,'type.num.num'),X104))))) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1284,f196])).
fof(f1284,plain,(
  ( ! [X103,X101,X105,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X101,bool),X102),s(fun(X101,'type.num.num'),X104)))) | s(bool,X105) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X101,bool),X102),s(fun(X101,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X101,'type.num.num'),X103),s(fun(X101,'type.num.num'),X104))))) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f239,f521])).
fof(f238,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X2))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1)))))) | (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X2)))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X1)))))) & (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X2)))) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X1)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1))))))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1)))))) | (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X2)))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X1)))))) & ((~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X2)))) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X1))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1))))))))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1)))))) <=> (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X2)))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X1))))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1)))))) <=> (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X2)))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X1))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ! [X15,X16,X8] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X7,'type.num.num'),X16),s(fun(X7,'type.num.num'),X15)))))) <=> (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),X16)))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),X15))))))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.bag.BAG_EVERY_UNION')).
fof(f239,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X1))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f240,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X2)))) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),X1)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X3),s(fun(X0,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X0,'type.num.num'),X2),s(fun(X0,'type.num.num'),X1))))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f1598,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(condensation,[],[f1597])).
fof(f1597,plain,(
  ( ! [X94,X95,X96] : (p(s(bool,X95)) | ~sP3('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f1317,f1222])).
fof(f1317,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,X94)) | p(s(bool,X95)) | ~sP3('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f676,f521])).
fof(f1596,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(condensation,[],[f1595])).
fof(f1595,plain,(
  ( ! [X92,X93,X91] : (p(s(bool,X92)) | ~sP4('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X92) | s(bool,X91) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f1316,f1222])).
fof(f1316,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,X91)) | p(s(bool,X92)) | ~sP4('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X92) | s(bool,X91) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f655,f521])).
fof(f1594,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP6('T_0',X1,X2) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1593])).
fof(f1593,plain,(
  ( ! [X90,X88,X87,X89] : (~sP6('T_0',X88,X89) | sP7('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1315,f1222])).
fof(f1315,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP6('T_0',X88,X89) | sP7('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f330,f521])).
fof(f1592,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1591])).
fof(f1591,plain,(
  ( ! [X85,X83,X86,X84] : (~sP8('T_0',X84,X85) | sP9('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1314,f1222])).
fof(f1314,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~sP8('T_0',X84,X85) | sP9('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f329,f521])).
fof(f1590,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1589])).
fof(f1589,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP9('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1313,f1222])).
fof(f1313,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP9('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f326,f521])).
fof(f1588,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1587])).
fof(f1587,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP9('T_0',X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1312,f1222])).
fof(f1312,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | sP9('T_0',X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f325,f521])).
fof(f1586,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1585])).
fof(f1585,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP9('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1311,f1222])).
fof(f1311,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | sP9('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f324,f521])).
fof(f1584,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1583])).
fof(f1583,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | sP9(X69,X68,'T_0') | p(s(bool,X69)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1310,f1222])).
fof(f1310,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | sP9(X69,X68,'T_0') | p(s(bool,X69)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f323,f521])).
fof(f1582,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1581])).
fof(f1581,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP8(X65,X64,'T_0') | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1309,f1222])).
fof(f1309,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP8(X65,X64,'T_0') | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f319,f521])).
fof(f1580,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f1579])).
fof(f1579,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | p(s(bool,X61)) | ~sP7('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1308,f1222])).
fof(f1308,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | p(s(bool,X60)) | p(s(bool,X61)) | ~sP7('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f309,f521])).
fof(f1578,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1577])).
fof(f1577,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP7('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1307,f1222])).
fof(f1307,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | sP7('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f308,f521])).
fof(f1576,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f1575])).
fof(f1575,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | sP7('T_0',X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1306,f1222])).
fof(f1306,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | sP7('T_0',X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f307,f521])).
fof(f1574,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f1573])).
fof(f1573,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | p(s(bool,X49)) | ~sP6('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1305,f1222])).
fof(f1305,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | p(s(bool,X49)) | ~sP6('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f302,f521])).
fof(f1572,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | ~sP5('T_0',X2,X1)) )),
  inference(condensation,[],[f1571])).
fof(f1571,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X44) = s(bool,X45) | ~sP5('T_0',X45,X44) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1304,f1222])).
fof(f1304,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | s(bool,X44) = s(bool,X45) | ~sP5('T_0',X45,X44) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f293,f521])).
fof(f1570,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1569])).
fof(f1569,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1303,f1222])).
fof(f1303,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f283,f521])).
fof(f1568,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1567])).
fof(f1567,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1302,f1222])).
fof(f1302,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f279,f521])).
fof(f1566,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1565,f352])).
fof(f1565,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1564])).
fof(f1564,plain,(
  ( ! [X33,X31,X34,X32] : (sP5(X32,X33,'T_0') | ~sP4(X32,'T_0',X33) | ~sP3(X33,'T_0',X32) | ~sP2('T_0',X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1301,f1222])).
fof(f1301,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X31)) | sP5(X32,X33,'T_0') | ~sP4(X32,'T_0',X33) | ~sP3(X33,'T_0',X32) | ~sP2('T_0',X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f276,f521])).
fof(f1563,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1562])).
fof(f1562,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'T_0',X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1300,f1222])).
fof(f1300,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'T_0',X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f275,f521])).
fof(f1561,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1560])).
fof(f1560,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X25)) | sP1(X25,'T_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f1299,f1222])).
fof(f1299,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(X25,'T_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f265,f521])).
fof(f1559,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1558])).
fof(f1558,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | sP1('T_0',X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f1298,f1222])).
fof(f1298,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | p(s(bool,X22)) | sP1('T_0',X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f262,f521])).
fof(f1289,plain,(
  ( ! [X134,X132,X133,X131] : (s(X133,f(s(bool,'T_0'))) != s(X133,g(s(bool,'T_0'))) | s(fun(bool,X133),X131) = s(fun(bool,X133),X132) | s(bool,X134) = s(bool,sK19(X131,X132,X133,bool)) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f270,f521])).
fof(f1554,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X90)) | ~sP3(X89,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X90) | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f1276,f196])).
fof(f1276,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X90)) | ~sP3(X89,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X90) | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f676,f521])).
fof(f1553,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X87)) | ~sP4(X86,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X87) | s(bool,X86) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f1275,f196])).
fof(f1275,plain,(
  ( ! [X88,X87,X86] : (~p(s(bool,'T_0')) | p(s(bool,X87)) | ~sP4(X86,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X87) | s(bool,X86) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f655,f521])).
fof(f1552,plain,(
  ( ! [X85,X83,X84,X82] : (~sP6(X82,X83,X84) | sP7(X82,X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f1274,f196])).
fof(f1274,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'T_0')) | ~sP6(X82,X83,X84) | sP7(X82,X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f330,f521])).
fof(f1551,plain,(
  ( ! [X80,X78,X81,X79] : (~sP8(X78,X79,X80) | sP9(X78,X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f1273,f196])).
fof(f1273,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,'T_0')) | ~sP8(X78,X79,X80) | sP9(X78,X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f329,f521])).
fof(f1550,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | ~sP9(X74,X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f1272,f196])).
fof(f1272,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | p(s(bool,X76)) | ~sP9(X74,X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f326,f521])).
fof(f1549,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X71)) | sP9(X70,X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f1271,f196])).
fof(f1271,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X71)) | sP9(X70,X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f325,f521])).
fof(f1548,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X67)) | sP9(X66,X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f1270,f196])).
fof(f1270,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | sP9(X66,X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f324,f521])).
fof(f1547,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | sP9(X64,X63,X62) | p(s(bool,X64)) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f1269,f196])).
fof(f1269,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X63)) | sP9(X64,X63,X62) | p(s(bool,X64)) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f323,f521])).
fof(f1546,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X59)) | ~p(s(bool,X60)) | ~sP8(X60,X59,X58) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f1268,f196])).
fof(f1268,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~sP8(X60,X59,X58) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f319,f521])).
fof(f1545,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X54,X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f1267,f196])).
fof(f1267,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X54,X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f309,f521])).
fof(f1544,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | sP7(X50,X51,X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f1266,f196])).
fof(f1266,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | sP7(X50,X51,X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f308,f521])).
fof(f1543,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | sP7(X46,X48,X47) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f1265,f196])).
fof(f1265,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | sP7(X46,X48,X47) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f307,f521])).
fof(f1542,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X43)) | p(s(bool,X44)) | ~sP6(X42,X43,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f1264,f196])).
fof(f1264,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X43)) | p(s(bool,X44)) | ~sP6(X42,X43,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f302,f521])).
fof(f1541,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X39) = s(bool,X40) | ~sP5(X38,X40,X39) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f1263,f196])).
fof(f1263,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | s(bool,X39) = s(bool,X40) | ~sP5(X38,X40,X39) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f293,f521])).
fof(f1540,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~sP3(X35,X34,X36) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f1262,f196])).
fof(f1262,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP3(X35,X34,X36) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f283,f521])).
fof(f1539,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | ~sP4(X31,X30,X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f1261,f196])).
fof(f1261,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP4(X31,X30,X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f279,f521])).
fof(f1538,plain,(
  ( ! [X28,X26,X29,X27] : (sP5(X27,X28,X26) | ~sP4(X27,X26,X28) | ~sP3(X28,X26,X27) | ~sP2(X26,X28,X27) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f1260,f196])).
fof(f1260,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'T_0')) | sP5(X27,X28,X26) | ~sP4(X27,X26,X28) | ~sP3(X28,X26,X27) | ~sP2(X26,X28,X27) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f276,f521])).
fof(f1537,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,X23)) | p(s(bool,X24)) | ~sP5(X23,X22,X24) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f1259,f196])).
fof(f1259,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP5(X23,X22,X24) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f275,f521])).
fof(f1536,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X20)) | sP1(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1258,f196])).
fof(f1258,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,X20)) | sP1(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f265,f521])).
fof(f1535,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X17)) | sP1(X16,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1257,f196])).
fof(f1257,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | sP1(X16,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f262,f521])).
fof(f1534,plain,(
  ( ! [X14,X15,X13] : (sP0(X13,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f1256,f196])).
fof(f1256,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | sP0(X13,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f247,f521])).
fof(f1533,plain,(
  ( ! [X12,X10,X11] : (sP0(X11,X10) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1255,f196])).
fof(f1255,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | sP0(X11,X10) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f246,f521])).
fof(f1532,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1253,f196])).
fof(f1253,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f214,f521])).
fof(f1530,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(condensation,[],[f1529])).
fof(f1529,plain,(
  ( ! [X94,X95,X96] : (p(s(bool,X95)) | ~sP3('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f1248,f1222])).
fof(f1248,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,X94)) | p(s(bool,X95)) | ~sP3('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f676,f521])).
fof(f1527,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(condensation,[],[f1526])).
fof(f1526,plain,(
  ( ! [X92,X93,X91] : (p(s(bool,X92)) | ~sP4('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X92) | s(bool,X91) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f1247,f1222])).
fof(f1247,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,X91)) | p(s(bool,X92)) | ~sP4('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X92) | s(bool,X91) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f655,f521])).
fof(f1524,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP6('T_0',X1,X2) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1523])).
fof(f1523,plain,(
  ( ! [X90,X88,X87,X89] : (~sP6('T_0',X88,X89) | sP7('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1246,f1222])).
fof(f1246,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP6('T_0',X88,X89) | sP7('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f330,f521])).
fof(f1521,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1520])).
fof(f1520,plain,(
  ( ! [X85,X83,X86,X84] : (~sP8('T_0',X84,X85) | sP9('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1245,f1222])).
fof(f1245,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~sP8('T_0',X84,X85) | sP9('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f329,f521])).
fof(f1518,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1517])).
fof(f1517,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP9('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1244,f1222])).
fof(f1244,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP9('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f326,f521])).
fof(f1515,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1514])).
fof(f1514,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP9('T_0',X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1243,f1222])).
fof(f1243,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | sP9('T_0',X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f325,f521])).
fof(f1512,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1511])).
fof(f1511,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP9('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1242,f1222])).
fof(f1242,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | sP9('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f324,f521])).
fof(f1509,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1508])).
fof(f1508,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | sP9(X69,X68,'T_0') | p(s(bool,X69)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1241,f1222])).
fof(f1241,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | sP9(X69,X68,'T_0') | p(s(bool,X69)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f323,f521])).
fof(f1506,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1505])).
fof(f1505,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP8(X65,X64,'T_0') | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1240,f1222])).
fof(f1240,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP8(X65,X64,'T_0') | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f319,f521])).
fof(f1503,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f1502])).
fof(f1502,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | p(s(bool,X61)) | ~sP7('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1239,f1222])).
fof(f1239,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | p(s(bool,X60)) | p(s(bool,X61)) | ~sP7('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f309,f521])).
fof(f1500,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1499])).
fof(f1499,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP7('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1238,f1222])).
fof(f1238,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | sP7('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f308,f521])).
fof(f1497,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f1496])).
fof(f1496,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | sP7('T_0',X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1237,f1222])).
fof(f1237,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | sP7('T_0',X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f307,f521])).
fof(f1494,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f1493])).
fof(f1493,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | p(s(bool,X49)) | ~sP6('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1236,f1222])).
fof(f1236,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | p(s(bool,X49)) | ~sP6('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f302,f521])).
fof(f1491,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | ~sP5('T_0',X2,X1)) )),
  inference(condensation,[],[f1490])).
fof(f1490,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X44) = s(bool,X45) | ~sP5('T_0',X45,X44) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1235,f1222])).
fof(f1235,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | s(bool,X44) = s(bool,X45) | ~sP5('T_0',X45,X44) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f293,f521])).
fof(f1488,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1487])).
fof(f1487,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1234,f1222])).
fof(f1234,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f283,f521])).
fof(f1485,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1484])).
fof(f1484,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1233,f1222])).
fof(f1233,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f279,f521])).
fof(f1482,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1481,f352])).
fof(f1481,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1480])).
fof(f1480,plain,(
  ( ! [X33,X31,X34,X32] : (sP5(X32,X33,'T_0') | ~sP4(X32,'T_0',X33) | ~sP3(X33,'T_0',X32) | ~sP2('T_0',X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1232,f1222])).
fof(f1232,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X31)) | sP5(X32,X33,'T_0') | ~sP4(X32,'T_0',X33) | ~sP3(X33,'T_0',X32) | ~sP2('T_0',X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f276,f521])).
fof(f1478,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1477])).
fof(f1477,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'T_0',X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1231,f1222])).
fof(f1231,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'T_0',X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f275,f521])).
fof(f1475,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1474])).
fof(f1474,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X25)) | sP1(X25,'T_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f1230,f1222])).
fof(f1230,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(X25,'T_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f265,f521])).
fof(f1471,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1470])).
fof(f1470,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | sP1('T_0',X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f1229,f1222])).
fof(f1229,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | p(s(bool,X22)) | sP1('T_0',X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f262,f521])).
fof(f1222,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f196,f521])).
fof(f1221,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f210,f521])).
fof(f1179,plain,(
  ( ! [X138,X136,X135,X137] : (s(X137,f(s(bool,X138))) != s(X137,g(s(bool,X138))) | s(fun(bool,X137),X135) = s(fun(bool,X137),X136) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,sK19(X135,X136,X137,bool))) )),
  inference(superposition,[],[f270,f521])).
fof(f1117,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f521,f521])).
fof(f1087,plain,(
  ~sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1077,f297])).
fof(f1077,plain,(
  ~sP6('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f340,f734])).
fof(f1064,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP3('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(subsumption_resolution,[],[f1058,f196])).
fof(f1058,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~sP3('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(superposition,[],[f676,f380])).
fof(f1049,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.SplayHeap.heap'(sK11),'type.SplayHeap.heap'(sK11)),'const.pair.,_2'(s('type.SplayHeap.heap'(sK11),sK17),s('type.SplayHeap.heap'(sK11),sK18))) != s('type.pair.prod'('type.SplayHeap.heap'(sK11),'type.SplayHeap.heap'(sK11)),'const.pair.,_2'(s('type.SplayHeap.heap'(sK11),X0),s('type.SplayHeap.heap'(sK11),X1))) | s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))) = s(fun(sK11,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),X0))),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),X1)))))) )),
  inference(superposition,[],[f271,f193])).
fof(f1043,plain,(
  s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))) = s(fun(sK11,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK17))),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))),
  inference(unit_resulting_resolution,[],[f193,f271])).
fof(f1041,plain,(
  ~sP5('T_0','F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f1026,f274])).
fof(f1037,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP4('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(subsumption_resolution,[],[f1031,f196])).
fof(f1031,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~sP4('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(superposition,[],[f655,f380])).
fof(f1026,plain,(
  ~sP4('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f655])).
fof(f954,plain,(
  ~sP3('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f365,f608,f344,f340,f278])).
fof(f959,plain,(
  sP5('F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f364,f348,f345,f340,f278])).
fof(f905,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f350,f347,f345,f340,f277])).
fof(f906,plain,(
  sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f363,f362,f345,f340,f277])).
fof(f916,plain,(
  sP5('F_0','F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f365,f349,f346,f340,f277])).
fof(f883,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | ~sP4(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1) | ~sP3(X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(subsumption_resolution,[],[f882,f365])).
fof(f882,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | ~sP4(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1) | ~sP3(X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0) | ~sP2('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0)) )),
  inference(subsumption_resolution,[],[f877,f196])).
fof(f877,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP5(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | ~sP4(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1) | ~sP3(X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0) | ~sP2('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0)) )),
  inference(superposition,[],[f276,f380])).
fof(f864,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f352,f347,f344,f196,f276])).
fof(f865,plain,(
  sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f352,f362,f344,f196,f276])).
fof(f867,plain,(
  sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f352,f362,f361,f196,f276])).
fof(f875,plain,(
  ~sP3('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f352,f346,f606,f196,f276])).
fof(f876,plain,(
  ~sP4('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f352,f349,f836,f196,f276])).
fof(f859,plain,(
  ~sP5('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f628])).
fof(f856,plain,(
  ( ! [X0,X1] : (~sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f849,f196])).
fof(f849,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0)) )),
  inference(superposition,[],[f326,f380])).
fof(f836,plain,(
  ~sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f376,f597])).
fof(f838,plain,(
  ~sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f376,f597])).
fof(f835,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f831,f196])).
fof(f831,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP9(X1,X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f323,f380])).
fof(f820,plain,(
  ~sP9('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f808,f314])).
fof(f816,plain,(
  ( ! [X0,X1] : (~sP8(X1,X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | ~p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f809,f196])).
fof(f809,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP8(X1,X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(superposition,[],[f319,f380])).
fof(f808,plain,(
  ~sP8('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f196,f340,f196,f319])).
fof(f796,plain,(
  sP9('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f358,f196,f340,f318])).
fof(f788,plain,(
  ( ! [X0] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f196,f477])).
fof(f779,plain,(
  ( ! [X0] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f340,f466])).
fof(f770,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f196,f444])).
fof(f769,plain,(
  ( ! [X0,X1] : (~sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f762,f196])).
fof(f762,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0)) )),
  inference(superposition,[],[f309,f380])).
fof(f738,plain,(
  ~sP7('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f726,f297])).
fof(f734,plain,(
  ( ! [X0,X1] : (~sP6('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f727,f196])).
fof(f727,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP6('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1)) )),
  inference(superposition,[],[f302,f380])).
fof(f726,plain,(
  ~sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f340,f196,f302])).
fof(f713,plain,(
  sP7('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f355,f340,f340,f301])).
fof(f695,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f686,f272])).
fof(f686,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f340,f340,f287])).
fof(f680,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f668,f273])).
fof(f676,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP3(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(subsumption_resolution,[],[f669,f196])).
fof(f669,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP3(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(superposition,[],[f283,f380])).
fof(f668,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f196,f196,f283])).
fof(f655,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP4(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(subsumption_resolution,[],[f648,f196])).
fof(f648,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP4(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(superposition,[],[f279,f380])).
fof(f647,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f196,f196,f279])).
fof(f628,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f623,f196])).
fof(f623,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(superposition,[],[f275,f380])).
fof(f606,plain,(
  ~sP5('F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f340,f380,f294])).
fof(f608,plain,(
  ~sP5('F_0','T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f340,f380,f294])).
fof(f597,plain,(
  ( ! [X0,X1] : (~sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f593,f196])).
fof(f593,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | ~sP5('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0)) )),
  inference(superposition,[],[f293,f380])).
fof(f586,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f196,f376,f293])).
fof(f590,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f196,f376,f293])).
fof(f575,plain,(
  sP5('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f196,f380,f292])).
fof(f577,plain,(
  sP5('T_0','T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f196,f380,f292])).
fof(f565,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f340,f376,f291])).
fof(f567,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f376,f291])).
fof(f564,plain,(
  ( ! [X0] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f196,f443])).
fof(f521,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f203,f203])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1) | ~sP6('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1)) )),
  inference(subsumption_resolution,[],[f511,f196])).
fof(f511,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP6('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1) | sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1)) )),
  inference(superposition,[],[f330,f380])).
fof(f508,plain,(
  ( ! [X0] : (sP7('T_0',X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f366,f196,f330])).
fof(f510,plain,(
  ( ! [X0] : (sP7('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f367,f196,f330])).
fof(f495,plain,(
  ( ! [X0,X1] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1) | ~sP8('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1)) )),
  inference(subsumption_resolution,[],[f493,f196])).
fof(f493,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1) | sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1)) )),
  inference(superposition,[],[f329,f380])).
fof(f492,plain,(
  ( ! [X0] : (sP9('T_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f368,f196,f329])).
fof(f481,plain,(
  sP1('F_0','const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f340,f442])).
fof(f477,plain,(
  ( ! [X0,X1] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f475,f196])).
fof(f475,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1)) )),
  inference(superposition,[],[f325,f380])).
fof(f474,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f196,f196,f325])).
fof(f466,plain,(
  ( ! [X0,X1] : (sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f464,f196])).
fof(f464,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP9('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1,X0)) )),
  inference(superposition,[],[f324,f380])).
fof(f463,plain,(
  ( ! [X0] : (sP9('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f340,f196,f324])).
fof(f460,plain,(
  sP1('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f340,f441])).
fof(f459,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f196,f340,f313])).
fof(f445,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f340,f340,f312])).
fof(f444,plain,(
  ( ! [X6,X7] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X6,X7) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f435,f196])).
fof(f435,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X6,X7)) )),
  inference(superposition,[],[f308,f380])).
fof(f443,plain,(
  ( ! [X4,X5] : (sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X5,X4) | ~p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f434,f196])).
fof(f434,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | sP7('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X5,X4)) )),
  inference(superposition,[],[f307,f380])).
fof(f442,plain,(
  ( ! [X3] : (sP1(X3,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f433,f196])).
fof(f433,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(superposition,[],[f265,f380])).
fof(f441,plain,(
  ( ! [X2] : (sP1('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f432,f196])).
fof(f432,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X2)) )),
  inference(superposition,[],[f262,f380])).
fof(f421,plain,(
  ( ! [X0] : (sP7('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f196,f196,f308])).
fof(f413,plain,(
  ( ! [X0] : (sP7('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f196,f196,f307])).
fof(f410,plain,(
  ( ! [X0] : (~sP7('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f196,f340,f296])).
fof(f409,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f196,f340,f295])).
fof(f404,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f340,f341])).
fof(f397,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f340,f196,f265])).
fof(f392,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f196,f262])).
fof(f389,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f196,f196,f261])).
fof(f388,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f340,f340,f260])).
fof(f380,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))),
  inference(unit_resulting_resolution,[],[f194,f211])).
fof(f376,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f340,f210])).
fof(f359,plain,(
  ( ! [X0] : (sP0(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f194,f246])).
fof(f360,plain,(
  ( ! [X0] : (sP0('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f194,f247])).
fof(f361,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f194,f280])).
fof(f362,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f194,f284])).
fof(f363,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f194,f288])).
fof(f364,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f194,f289])).
fof(f365,plain,(
  ( ! [X0,X1] : (sP2('const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f194,f290])).
fof(f366,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f194,f303])).
fof(f367,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f194,f304])).
fof(f368,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f194,f321])).
fof(f358,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f340,f322])).
fof(f357,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f196,f321])).
fof(f356,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f340,f320])).
fof(f355,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f340,f305])).
fof(f354,plain,(
  ( ! [X0,X1] : (sP6(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f196,f304])).
fof(f353,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f196,f303])).
fof(f352,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f196,f290])).
fof(f351,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f196,f289])).
fof(f350,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f196,f288])).
fof(f349,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f340,f286])).
fof(f348,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f340,f285])).
fof(f347,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f196,f284])).
fof(f346,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f340,f282])).
fof(f345,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f340,f281])).
fof(f344,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f196,f280])).
fof(f343,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f196,f247])).
fof(f342,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f196,f246])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP9(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f117,f131,f130])).
fof(f130,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f131,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X19,X20,X11] : ((p(s(bool,X11)) <=> (p(s(bool,X20)) => p(s(bool,X19)))) <=> ((p(s(bool,X11)) | p(s(bool,X20))) & (p(s(bool,X11)) | ~p(s(bool,X19))) & (~p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X11)))))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.sat.sat_27_prove_ltb')).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f189])).
fof(f189,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f130])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f190])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f190])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f190])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f131])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f187])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f315])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f187])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f187])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f88,f128,f127])).
fof(f127,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f128,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f88,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X19,X20,X11] : ((p(s(bool,X11)) <=> (p(s(bool,X20)) | p(s(bool,X19)))) <=> ((p(s(bool,X11)) | ~p(s(bool,X20))) & (p(s(bool,X11)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X11)))))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.sat.sat_25_prove_ltb')).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f181])).
fof(f181,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f127])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f182])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f182])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f128])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f179])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f179])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f298])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f86,f125,f124,f123,f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f123,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f124,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f125,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f86,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X19,X20,X11] : ((p(s(bool,X11)) <=> s(bool,X19) = s(bool,X20)) <=> ((p(s(bool,X11)) | p(s(bool,X20)) | p(s(bool,X19))) & (p(s(bool,X11)) | ~p(s(bool,X19)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X19)) | ~p(s(bool,X11))) & (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X11)))))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,X0) = s(bool,X1) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f176])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f122])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f175])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f175])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f175])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f172])).
fof(f172,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f123])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f173])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f173])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f173])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f169])).
fof(f169,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f124])).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f170])).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f170])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f170])).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f166])).
fof(f166,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f125])).
fof(f273,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f274,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f167])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f271,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.pair.,_2'(s('type.SplayHeap.heap'(X1),X6),s('type.SplayHeap.heap'(X1),X7))) != s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.SplayHeap.partition_4'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s(X1,X4),s('type.SplayHeap.heap'(X1),X5))) | s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X5))) = s(fun(X1,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X6))),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X7)))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.pair.,_2'(s('type.SplayHeap.heap'(X1),X6),s('type.SplayHeap.heap'(X1),X7))) != s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.SplayHeap.partition_4'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s(X1,X4),s('type.SplayHeap.heap'(X1),X5))) | s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X5))) = s(fun(X1,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X6))),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X7))))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.pair.,_2'(s('type.SplayHeap.heap'(X1),X6),s('type.SplayHeap.heap'(X1),X7))) = s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.SplayHeap.partition_4'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s(X1,X4),s('type.SplayHeap.heap'(X1),X5))) => s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X5))) = s(fun(X1,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X6))),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X7))))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : (s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.pair.,_2'(s('type.SplayHeap.heap'(X1),X6),s('type.SplayHeap.heap'(X1),X7))) = s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.SplayHeap.partition_4'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s(X1,X4),s('type.SplayHeap.heap'(X1),X5))) => s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X5))) = s(fun(X1,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X6))),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X7))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X9,X10,X11,X12,X13,X14] : (s('type.pair.prod'('type.SplayHeap.heap'(X7),'type.SplayHeap.heap'(X7)),'const.pair.,_2'(s('type.SplayHeap.heap'(X7),X13),s('type.SplayHeap.heap'(X7),X14))) = s('type.pair.prod'('type.SplayHeap.heap'(X7),'type.SplayHeap.heap'(X7)),'const.SplayHeap.partition_4'(s(fun(X7,X6),X9),s(fun(X6,fun(X6,bool)),X10),s(X7,X11),s('type.SplayHeap.heap'(X7),X12))) => s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X12))) = s(fun(X7,'type.num.num'),'const.bag.BAG_UNION_2'(s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X13))),s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X14))))))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.SplayHeap.SplayHeap_7|49|_prove_ltb')).
fof(f270,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb',aHL_EXT)).
fof(f260,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f161])).
fof(f161,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f261,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f162])).
fof(f262,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f162])).
fof(f265,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f162])).
fof(f341,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f245,f340])).
fof(f245,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f246,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f155])).
fof(f247,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f155])).
fof(f210,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f211,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f203,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb',aHL_BOOL_CASES)).
fof(f196,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb',aHL_TRUTH)).
fof(f193,plain,(
  s('type.pair.prod'('type.SplayHeap.heap'(sK11),'type.SplayHeap.heap'(sK11)),'const.pair.,_2'(s('type.SplayHeap.heap'(sK11),sK17),s('type.SplayHeap.heap'(sK11),sK18))) = s('type.pair.prod'('type.SplayHeap.heap'(sK11),'type.SplayHeap.heap'(sK11)),'const.SplayHeap.partition_4'(s(fun(sK11,sK10),sK13),s(fun(sK10,fun(sK10,bool)),sK14),s(sK11,sK15),s('type.SplayHeap.heap'(sK11),sK16)))),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  s('type.pair.prod'('type.SplayHeap.heap'(sK11),'type.SplayHeap.heap'(sK11)),'const.pair.,_2'(s('type.SplayHeap.heap'(sK11),sK17),s('type.SplayHeap.heap'(sK11),sK18))) = s('type.pair.prod'('type.SplayHeap.heap'(sK11),'type.SplayHeap.heap'(sK11)),'const.SplayHeap.partition_4'(s(fun(sK11,sK10),sK13),s(fun(sK10,fun(sK10,bool)),sK14),s(sK11,sK15),s('type.SplayHeap.heap'(sK11),sK16))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16)))))) & (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK17)))))) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13,sK14,sK15,sK16,sK17,sK18])],[f92])).
fof(f92,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.pair.,_2'(s('type.SplayHeap.heap'(X1),X7),s('type.SplayHeap.heap'(X1),X8))) = s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.SplayHeap.partition_4'(s(fun(X1,X0),X3),s(fun(X0,fun(X0,bool)),X4),s(X1,X5),s('type.SplayHeap.heap'(X1),X6))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X6)))))) & (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X7)))))) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X8))))))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.pair.,_2'(s('type.SplayHeap.heap'(X1),X7),s('type.SplayHeap.heap'(X1),X8))) = s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.SplayHeap.partition_4'(s(fun(X1,X0),X3),s(fun(X0,fun(X0,bool)),X4),s(X1,X5),s('type.SplayHeap.heap'(X1),X6))) & (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X6)))))) & (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X7)))))) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X8)))))))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.pair.,_2'(s('type.SplayHeap.heap'(X1),X7),s('type.SplayHeap.heap'(X1),X8))) = s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.SplayHeap.partition_4'(s(fun(X1,X0),X3),s(fun(X0,fun(X0,bool)),X4),s(X1,X5),s('type.SplayHeap.heap'(X1),X6))) => (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X6)))))) => (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X7)))))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X8)))))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.pair.,_2'(s('type.SplayHeap.heap'(X1),X7),s('type.SplayHeap.heap'(X1),X8))) = s('type.pair.prod'('type.SplayHeap.heap'(X1),'type.SplayHeap.heap'(X1)),'const.SplayHeap.partition_4'(s(fun(X1,X0),X3),s(fun(X0,fun(X0,bool)),X4),s(X1,X5),s('type.SplayHeap.heap'(X1),X6))) => (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X6)))))) => (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X7)))))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X1,bool),X2),s(fun(X1,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X1),X8)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11,X12,X13,X14] : (s('type.pair.prod'('type.SplayHeap.heap'(X7),'type.SplayHeap.heap'(X7)),'const.pair.,_2'(s('type.SplayHeap.heap'(X7),X13),s('type.SplayHeap.heap'(X7),X14))) = s('type.pair.prod'('type.SplayHeap.heap'(X7),'type.SplayHeap.heap'(X7)),'const.SplayHeap.partition_4'(s(fun(X7,X6),X9),s(fun(X6,fun(X6,bool)),X10),s(X7,X11),s('type.SplayHeap.heap'(X7),X12))) => (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X12)))))) => (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X13)))))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X14)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11,X12,X13,X14] : (s('type.pair.prod'('type.SplayHeap.heap'(X7),'type.SplayHeap.heap'(X7)),'const.pair.,_2'(s('type.SplayHeap.heap'(X7),X13),s('type.SplayHeap.heap'(X7),X14))) = s('type.pair.prod'('type.SplayHeap.heap'(X7),'type.SplayHeap.heap'(X7)),'const.SplayHeap.partition_4'(s(fun(X7,X6),X9),s(fun(X6,fun(X6,bool)),X10),s(X7,X11),s('type.SplayHeap.heap'(X7),X12))) => (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X12)))))) => (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X13)))))) & p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X7,bool),X8),s(fun(X7,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(X7),X14)))))))))),
  file('thm.SplayHeap.SplayHeap_73_prove_ltb','thm.SplayHeap.SplayHeap_73_prove_ltb')).
fof(f194,plain,(
  p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK16))))))),
  inference(cnf_transformation,[],[f133])).
fof(f195,plain,(
  ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK18)))))) | ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK11,bool),sK12),s(fun(sK11,'type.num.num'),'const.SplayHeap.heap_to_bag_1'(s('type.SplayHeap.heap'(sK11),sK17))))))),
  inference(cnf_transformation,[],[f133])).
% SZS output end Proof for thm.SplayHeap
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 299
% Passive clauses: 488
% Generated clauses: 3000
% Final active clauses: 256
% Final passive clauses: 166
% Input formulas: 35
% Initial clauses: 136
% 
% Duplicate literals: 147
% Trivial inequalities: 15
% Fw subsumption resolutions: 146
% Fw demodulations: 255
% Bw demodulations: 22
% Condensations: 101
% Global subsumptions: 4
% 
% Simple tautologies: 55
% Equational tautologies: 44
% Forward subsumptions: 1691
% 
% Unit resulting resolution: 1268
% Forward superposition: 384
% Backward superposition: 424
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 2
% 
% Split clauses: 48
% Split components: 101
% SAT solver clauses: 593
% SAT solver unit clauses: 165
% SAT solver binary clauses: 186
% 
% Sat splits: 48
% Sat splitting refutations: 3
% 
% Memory used [KB]: 2430
% Time elapsed: 0.156 s
% ------------------------------
% ------------------------------
% Success in time 1.58 s
