Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f5131,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5101,f1976,f1002,f123,f198,f196,f128,f203,f1911,f1977,f1975,f5120,f122,f197,f195,f4933,f4939,f4937,f4935,f4979,f5060])).
fof(f5060,plain,(
  $false | ($spl0 | $spl3 | $spl164)),
  inference(subsumption_resolution,[],[f5059,f198])).
fof(f5059,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | ($spl0 | $spl164)),
  inference(forward_demodulation,[],[f5027,f1031])).
fof(f1031,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | $spl0),
  inference(resolution,[],[f195,f169])).
fof(f169,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) & (s('type.clos_known.val_approx',X1) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))))))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) <=> s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X11,X49] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X11),s('type.clos_known.val_approx',X49)))) <=> s('type.clos_known.val_approx',X49) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X11),s('type.clos_known.val_approx',X49))))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.clos_knownProof.subapprox_def')).
fof(f5027,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | $spl164),
  inference(superposition,[],[f135,f4935])).
fof(f135,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X23] : s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X23))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X23))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.clos_known.merge_def|split|2')).
fof(f4979,plain,(
  $false | ($spl0 | $spl3 | $spl166)),
  inference(subsumption_resolution,[],[f4978,f198])).
fof(f4978,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | ($spl0 | $spl166)),
  inference(forward_demodulation,[],[f4943,f1031])).
fof(f4943,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | $spl166),
  inference(superposition,[],[f134,f4937])).
fof(f134,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X24] : s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X24))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X24))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.clos_known.merge_def|split|3')).
fof(f4935,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | $spl164),
  inference(cnf_transformation,[],[f4935_D])).
fof(f4935_D,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f4937,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | $spl166),
  inference(cnf_transformation,[],[f4937_D])).
fof(f4937_D,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f4939,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | $spl168),
  inference(cnf_transformation,[],[f4939_D])).
fof(f4939_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f4933,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4910,f198])).
fof(f4910,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | $spl0),
  inference(duplicate_literal_removal,[],[f4882])).
fof(f4882,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | $spl0),
  inference(superposition,[],[f2804,f1031])).
fof(f2804,plain,(
  ( ! [X2] : (s('type.clos_known.val_approx',X2) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X2),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X2) | s('type.clos_known.val_approx',X2) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X2)))) | s('type.clos_known.val_approx',X2) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X2)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X2)) )),
  inference(superposition,[],[f167,f137])).
fof(f137,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK1(X0)),s('type.num.num',sK2(X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK1(X0)),s('type.num.num',sK2(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f54])).
fof(f54,plain,(
  ! [X0] : (? [X3,X4] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X4))) | ? [X2] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))) | ? [X1] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X15] : (? [X16,X17] : s('type.clos_known.val_approx',X15) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X16),s('type.num.num',X17))) | ? [X18] : s('type.clos_known.val_approx',X15) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X18))) | ? [X19] : s('type.clos_known.val_approx',X15) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X19))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X15) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X15))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.clos_known.val_approx_nchotomy')).
fof(f167,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X21,X22] : s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X22),s('type.num.num',X21))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X22),s('type.num.num',X21))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.clos_known.merge_def|split|1')).
fof(f195,plain,(
  p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) | $spl0),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f197,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | $spl2),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f122,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0)) & (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx',sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f89])).
fof(f89,plain,(
  ? [X0] : ((p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0)) & (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx',X0)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ? [X0] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) <~> s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ~! [X0] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) <=> s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) <=> s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) <=> s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X6))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.clos_knownProof.subapprox_Impossible|split|0')).
fof(f5120,plain,(
  $false | ($spl0 | $spl9 | $spl70 | $spl168)),
  inference(subsumption_resolution,[],[f5100,f2004])).
fof(f2004,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | ($spl9 | $spl70)),
  inference(subsumption_resolution,[],[f1985,f203])).
fof(f1985,plain,(
  p(s(bool,'F_0')) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | $spl70),
  inference(superposition,[],[f170,f1975])).
fof(f170,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) | s('type.clos_known.val_approx',X1) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f5100,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | ($spl0 | $spl168)),
  inference(backward_demodulation,[],[f4939,f1031])).
fof(f1975,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | $spl70),
  inference(cnf_transformation,[],[f1975_D])).
fof(f1975_D,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1977,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.Other_0') | $spl72),
  inference(cnf_transformation,[],[f1977_D])).
fof(f1977_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.Other_0') <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1911,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.Other_0') | s(bool,'F_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(superposition,[],[f816,f127])).
fof(f127,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(cnf_transformation,[],[f26])).
fof(f26,axiom,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.clos_known.merge_def|split|4')).
fof(f816,plain,(
  ( ! [X4,X3] : (s('type.clos_known.val_approx',X3) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X4),s('type.clos_known.val_approx',X3))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X4),s('type.clos_known.val_approx',X3)))) )),
  inference(resolution,[],[f169,f143])).
fof(f143,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f203,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f203_D])).
fof(f203_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f128,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0',aHL_FALSITY)).
fof(f196,plain,(
  ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) | $spl1),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f198,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx',sK0) | $spl3),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx',sK0) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f123,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx',sK0) | ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))))),
  inference(cnf_transformation,[],[f90])).
fof(f1002,plain,(
  $false | ($spl1 | $spl2 | $spl9)),
  inference(subsumption_resolution,[],[f1001,f129])).
fof(f129,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X20] : s('type.clos_known.val_approx',X20) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X20)))),
  file('thm.clos_knownProof.subapprox_Impossible|split|0','thm.clos_known.merge_def|split|0')).
fof(f1001,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | ($spl1 | $spl2 | $spl9)),
  inference(subsumption_resolution,[],[f960,f203])).
fof(f960,plain,(
  p(s(bool,'F_0')) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | ($spl1 | $spl2)),
  inference(superposition,[],[f170,f210])).
fof(f210,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | ($spl1 | $spl2)),
  inference(resolution,[],[f143,f205])).
fof(f205,plain,(
  ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f196,f197])).
fof(f1976,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx','const.clos_known.Other_0') | $spl73),
  inference(cnf_transformation,[],[f1976_D])).
fof(f1976_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx','const.clos_known.Other_0') <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f5101,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') != s('type.clos_known.val_approx','const.clos_known.Other_0') | ($spl3 | $spl168)),
  inference(superposition,[],[f198,f4939])).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 155
% Passive clauses: 630
% Generated clauses: 4894
% Final active clauses: 122
% Final passive clauses: 344
% Input formulas: 43
% Initial clauses: 71
% 
% Duplicate literals: 373
% Trivial inequalities: 94
% Fw subsumption resolutions: 488
% Fw demodulations: 519
% Bw demodulations: 18
% 
% Simple tautologies: 99
% Equational tautologies: 33
% Forward subsumptions: 2366
% 
% Binary resolution: 59
% Forward superposition: 1772
% Backward superposition: 1292
% Self superposition: 72
% Equality factoring: 38
% Equality resolution: 17
% 
% Split clauses: 288
% Split components: 590
% SAT solver clauses: 377
% SAT solver unit clauses: 17
% SAT solver binary clauses: 63
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 288
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 372
% TWLsolver calls for satisfiability: 71
% 
% Memory used [KB]: 3582
% Time elapsed: 0.118 s
% ------------------------------
% ------------------------------
% Success in time 0.14 s
