Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_numberProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_numberProof
% SZS output start Proof for thm.clos_numberProof
fof(f1134,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1133,f199,f317])).
fof(f317,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f317_D])).
fof(f317_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f199,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_numberProof.renumber_code_locs_imp_inc|split|1',aHL_FALSITY)).
fof(f1133,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f1132,f317])).
fof(f1132,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f1125,f333])).
fof(f333,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9),s('type.num.num',sK7)))),
  inference(resolution,[],[f212,f192])).
fof(f192,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9),s('type.num.num',sK7))))),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',sK9),s('type.closLang.exp',sK10))) = s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.pair.,_2'(s('type.num.num',sK7),s('type.closLang.exp',sK8))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9),s('type.num.num',sK7))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9,sK10])],[f97])).
fof(f97,plain,(
  ? [X0,X1,X2,X3] : (s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',X2),s('type.closLang.exp',X3))) = s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.pair.,_2'(s('type.num.num',X0),s('type.closLang.exp',X1))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3] : (s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',X2),s('type.closLang.exp',X3))) = s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.pair.,_2'(s('type.num.num',X0),s('type.closLang.exp',X1))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : (s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',X8),s('type.closLang.exp',X9))) = s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.pair.,_2'(s('type.num.num',X6),s('type.closLang.exp',X7))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : (s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',X8),s('type.closLang.exp',X9))) = s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.pair.,_2'(s('type.num.num',X6),s('type.closLang.exp',X7))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6)))))),
  file('thm.clos_numberProof.renumber_code_locs_imp_inc|split|1','thm.clos_numberProof.renumber_code_locs_imp_inc|split|1')).
fof(f212,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.clos_numberProof.renumber_code_locs_imp_inc|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f1125,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9),s('type.num.num',sK7))))),
  inference(superposition,[],[f248,f920])).
fof(f920,plain,(
  s('type.num.num',sK7) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',sK9),s('type.closLang.exp',sK10)))))),
  inference(superposition,[],[f270,f191])).
fof(f191,plain,(
  s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',sK9),s('type.closLang.exp',sK10))) = s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.pair.,_2'(s('type.num.num',sK7),s('type.closLang.exp',sK8)))),
  inference(cnf_transformation,[],[f132])).
fof(f270,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X17,X11] : ! [X12,X13] : s(X11,X12) = s(X11,'const.pair.FST_1'(s('type.pair.prod'(X11,X17),'const.pair.,_2'(s(X11,X12),s(X17,X13)))))),
  file('thm.clos_numberProof.renumber_code_locs_imp_inc|split|1','thm.pair.FST')).
fof(f248,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',X0),s('type.closLang.exp',X1))))))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',X0),s('type.closLang.exp',X1))))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X8,X9] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num','type.closLang.exp'),'const.clos_number.renumber_code_locs_2'(s('type.num.num',X8),s('type.closLang.exp',X9))))))))),
  file('thm.clos_numberProof.renumber_code_locs_imp_inc|split|1','thm.clos_numberProof.renumber_code_locs_inc|split|1')).
% SZS output end Proof for thm.clos_numberProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 126
% Passive clauses: 316
% Generated clauses: 936
% Final active clauses: 125
% Final passive clauses: 151
% Input formulas: 38
% Initial clauses: 121
% 
% Duplicate literals: 29
% Trivial inequalities: 1
% Fw subsumption resolutions: 52
% Fw demodulations: 51
% Bw demodulations: 1
% 
% Simple tautologies: 94
% Equational tautologies: 10
% Forward subsumptions: 372
% 
% Binary resolution: 463
% Factoring: 8
% Forward superposition: 85
% Backward superposition: 103
% Self superposition: 5
% Equality factoring: 2
% Equality resolution: 2
% 
% Split clauses: 34
% Split components: 69
% SAT solver clauses: 54
% SAT solver unit clauses: 5
% SAT solver binary clauses: 25
% SAT solver learnt clauses: 1
% 
% Sat splits: 34
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 53
% TWLsolver calls for satisfiability: 34
% 
% Memory used [KB]: 1151
% Time elapsed: 0.059 s
% ------------------------------
% ------------------------------
% Success in time 0.064 s
