Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f4578,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4577,f4488,f4517,f4515,f4513,f4511,f4554,f4529,f4569])).
fof(f4569,plain,(
  $false | $spl164),
  inference(subsumption_resolution,[],[f4564,f97])).
fof(f97,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(cnf_transformation,[],[f35])).
fof(f35,axiom,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.merge_Other|split|1','thm.clos_known.merge_def|split|20')).
fof(f4564,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | $spl164),
  inference(backward_demodulation,[],[f4517,f94])).
fof(f94,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f69])).
fof(f69,plain,(
  ? [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ~! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.merge_Other|split|1','thm.clos_knownProof.merge_Other|split|1')).
fof(f4529,plain,(
  $false | $spl160),
  inference(subsumption_resolution,[],[f4518,f94])).
fof(f4518,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | $spl160),
  inference(superposition,[],[f103,f4513])).
fof(f103,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X32] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X32))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.merge_Other|split|1','thm.clos_known.merge_def|split|13')).
fof(f4554,plain,(
  $false | $spl162),
  inference(subsumption_resolution,[],[f4540,f94])).
fof(f4540,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK0),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | $spl162),
  inference(superposition,[],[f102,f4515])).
fof(f102,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X38] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X38))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.merge_Other|split|1','thm.clos_known.merge_def|split|16')).
fof(f4511,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) | $spl158),
  inference(cnf_transformation,[],[f4511_D])).
fof(f4511_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f4513,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | $spl160),
  inference(cnf_transformation,[],[f4513_D])).
fof(f4513_D,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f4515,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | $spl162),
  inference(cnf_transformation,[],[f4515_D])).
fof(f4515_D,plain,(
  s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f4517,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | $spl164),
  inference(cnf_transformation,[],[f4517_D])).
fof(f4517_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f4488,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0)),
  inference(trivial_inequality_removal,[],[f4484])).
fof(f4484,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.Other_0') | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK0) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(sK0)))) | s('type.clos_known.val_approx',sK0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(sK0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK0)),
  inference(superposition,[],[f94,f2003])).
fof(f2003,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0)) )),
  inference(superposition,[],[f126,f108])).
fof(f108,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK1(X0)),s('type.num.num',sK2(X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK1(X0)),s('type.num.num',sK2(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK3(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK4(X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f46])).
fof(f46,plain,(
  ! [X0] : (? [X3,X4] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X4))) | ? [X2] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))) | ? [X1] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : (? [X11,X12] : s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X11),s('type.num.num',X12))) | ? [X13] : s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X13))) | ? [X14] : s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X10) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X10))),
  file('thm.clos_knownProof.merge_Other|split|1','thm.clos_known.val_approx_nchotomy')).
fof(f126,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X28,X29] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X29),s('type.num.num',X28))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.merge_Other|split|1','thm.clos_known.merge_def|split|10')).
fof(f4577,plain,(
  $false | $spl158),
  inference(subsumption_resolution,[],[f4576,f100])).
fof(f100,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X15] : s('type.clos_known.val_approx',X15) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X15)))),
  file('thm.clos_knownProof.merge_Other|split|1','thm.clos_known.merge_def|split|0')).
fof(f4576,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | $spl158),
  inference(forward_demodulation,[],[f94,f4511])).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 113
% Passive clauses: 411
% Generated clauses: 4350
% Final active clauses: 101
% Final passive clauses: 239
% Input formulas: 35
% Initial clauses: 53
% 
% Duplicate literals: 608
% Trivial inequalities: 39
% Fw subsumption resolutions: 368
% Fw demodulations: 38
% Bw demodulations: 6
% 
% Simple tautologies: 71
% Equational tautologies: 27
% Forward subsumptions: 2379
% 
% Binary resolution: 28
% Forward superposition: 1514
% Backward superposition: 1443
% Self superposition: 120
% Equality factoring: 52
% Equality resolution: 9
% 
% Split clauses: 355
% Split components: 728
% SAT solver clauses: 453
% SAT solver unit clauses: 8
% SAT solver binary clauses: 15
% SAT solver learnt clauses: 1
% 
% Sat splits: 355
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 442
% TWLsolver calls for satisfiability: 50
% 
% Memory used [KB]: 3326
% Time elapsed: 0.220 s
% ------------------------------
% ------------------------------
% Success in time 0.24 s
