Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relationProps
% (93345)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 149
% Passive clauses: 2138
% Generated clauses: 9888
% Final active clauses: 118
% Final passive clauses: 1801
% Input formulas: 40
% Initial clauses: 65
% 
% Duplicate literals: 254
% Trivial inequalities: 202
% Fw subsumption resolutions: 562
% Fw demodulations: 619
% Bw demodulations: 23
% 
% Simple tautologies: 293
% Equational tautologies: 21
% Forward subsumptions: 5721
% 
% Binary resolution: 190
% Forward superposition: 4735
% Backward superposition: 3343
% Self superposition: 209
% Equality factoring: 44
% Equality resolution: 28
% 
% Split clauses: 105
% Split components: 211
% SAT solver clauses: 238
% SAT solver unit clauses: 6
% SAT solver binary clauses: 59
% 
% Sat splits: 105
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 235
% TWLsolver calls for satisfiability: 55
% 
% Memory used [KB]: 18933
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_relationProps
% (93366)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 74
% Passive clauses: 635
% Generated clauses: 3511
% Final active clauses: 69
% Final passive clauses: 554
% Input formulas: 40
% Initial clauses: 65
% 
% Duplicate literals: 10
% Trivial inequalities: 53
% Fw subsumption resolutions: 65
% Fw demodulations: 444
% Bw demodulations: 11
% Condensations: 17
% 
% Simple tautologies: 31
% Equational tautologies: 2
% Forward subsumptions: 1359
% 
% Binary resolution: 32
% Forward superposition: 1146
% Backward superposition: 878
% Self superposition: 87
% Equality factoring: 2
% Equality resolution: 26
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 908
% SAT solver unit clauses: 80
% SAT solver binary clauses: 38
% 
% Sat splits: 10
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 908
% TWLsolver calls for satisfiability: 322
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_relationProps
% (93405)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 414
% Passive clauses: 13230
% Generated clauses: 81479
% Final active clauses: 67
% Final passive clauses: 306
% Input formulas: 40
% Initial clauses: 65
% Discarded non-redundant clauses: 3329
% 
% Duplicate literals: 1742
% Trivial inequalities: 592
% Fw subsumption resolutions: 4208
% Fw demodulations: 13304
% Bw demodulations: 8803
% Condensations: 1641
% Global subsumptions: 87
% 
% Simple tautologies: 555
% Equational tautologies: 583
% Forward subsumptions: 44416
% Fw demodulations to eq. taut.: 69
% Bw demodulations to eq. taut.: 374
% 
% Unit resulting resolution: 4864
% Forward superposition: 20597
% Backward superposition: 14668
% Self superposition: 391
% Equality factoring: 75
% Equality resolution: 62
% 
% Split clauses: 178
% Split components: 364
% SAT solver clauses: 13827
% SAT solver unit clauses: 978
% SAT solver binary clauses: 3457
% 
% Sat splits: 178
% Sat splitting refutations: 152
% 
% Memory used [KB]: 90702
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 30 next slice time: 30
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.clos_relationProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relationProps
% SZS output start Proof for thm.clos_relationProps
fof(f2231,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f139,f143,f99,f145,f98,f147,f97,f149,f151,f152,f140,f156,f157,f158,f159,f160,f161,f163,f164,f177,f178,f179,f180,f181,f183,f184,f185,f188,f189,f190,f191,f192,f193,f195,f196,f203,f204,f205,f206,f207,f209,f210,f211,f215,f216,f217,f218,f219,f221,f222,f223,f237,f259,f261,f263,f260,f264,f256,f255,f253,f252,f251,f250,f249,f248,f285,f278,f292,f293,f294,f295,f296,f297,f301,f302,f303,f304,f305,f307,f308,f309,f314,f316,f317,f319,f321,f322,f325,f331,f333,f334,f335,f338,f339,f344,f346,f347,f348,f351,f352,f354,f356,f357,f359,f361,f362,f368,f370,f371,f372,f375,f376,f380,f382,f383,f385,f387,f388,f402,f403,f404,f405,f406,f407,f400,f399,f397,f395,f394,f392,f413,f415,f416,f417,f420,f421,f374,f369,f437,f429,f438,f439,f440,f434,f441,f442,f436,f435,f432,f431,f430,f428,f455,f447,f456,f457,f458,f452,f459,f460,f454,f453,f450,f449,f448,f446,f466,f468,f469,f471,f473,f474,f478,f480,f481,f483,f485,f486,f504,f506,f507,f509,f511,f512,f390,f516,f518,f519,f521,f523,f524,f527,f529,f530,f532,f534,f535,f198,f541,f543,f544,f545,f548,f549,f552,f554,f555,f557,f559,f560,f562,f574,f575,f576,f577,f578,f579,f582,f584,f585,f586,f589,f590,f592,f605,f606,f607,f609,f610,f612,f614,f615,f634,f636,f637,f639,f641,f642,f658,f660,f661,f663,f665,f666,f638,f635,f670,f675,f676,f677,f678,f679,f681,f187,f682,f683,f186,f687,f688,f689,f690,f691,f693,f169,f225,f694,f695,f700,f702,f703,f704,f707,f708,f726,f728,f729,f731,f733,f734,f739,f741,f742,f744,f746,f747,f526,f763,f765,f766,f768,f770,f771,f775,f777,f778,f780,f782,f783,f265,f798,f799,f800,f801,f802,f803,f168,f796,f795,f793,f791,f790,f788,f809,f811,f812,f814,f816,f817,f736,f821,f823,f824,f826,f828,f829,f833,f835,f836,f838,f840,f841,f847,f849,f850,f851,f854,f855,f224,f859,f861,f862,f863,f866,f867,f872,f874,f875,f877,f879,f880,f773,f884,f886,f887,f889,f891,f892,f898,f900,f901,f903,f905,f906,f912,f914,f915,f916,f919,f920,f353,f923,f925,f926,f928,f930,f931,f945,f946,f947,f948,f949,f950,f943,f942,f940,f938,f937,f935,f955,f957,f958,f960,f962,f963,f968,f970,f971,f973,f975,f976,f668,f980,f982,f983,f985,f987,f988,f843,f994,f996,f997,f998,f1001,f1002,f1006,f1008,f1009,f1011,f1013,f1014,f1021,f1023,f1024,f1025,f1028,f1029,f461,f1055,f1056,f1057,f1058,f1059,f1054,f1053,f1050,f1049,f1048,f1046,f1064,f1066,f1067,f1069,f1071,f1072,f894,f1076,f1078,f1079,f1081,f1083,f1084,f1086,f1090,f1092,f1093,f1095,f1097,f1098,f1102,f1104,f1105,f1107,f1109,f1110,f1151,f1153,f1154,f1156,f1158,f1159,f933,f1163,f1165,f1166,f1168,f1170,f1171,f1174,f1176,f1177,f1179,f1181,f1182,f1186,f1188,f1189,f1191,f1193,f1194,f1199,f1201,f1202,f1204,f1206,f1207,f1225,f1227,f1228,f1229,f1232,f1233,f591,f1236,f1238,f1239,f1241,f1243,f1244,f1060,f1248,f146,f141,f1252,f1254,f1255,f1257,f1259,f1260,f166,f1274,f1275,f1276,f1277,f1278,f1279,f423,f1281,f1273,f1272,f1269,f1268,f1267,f1265,f340,f1282,f1284,f1286,f1287,f1289,f1291,f1292,f488,f1298,f1300,f1301,f1302,f1305,f1306,f1307,f1312,f1314,f1315,f1316,f1319,f1320,f422,f1324,f1326,f1327,f1329,f1331,f1332,f1173,f1351,f1353,f1354,f1356,f1358,f1359,f1209,f1362,f1364,f1365,f1367,f1369,f1370,f965,f1374,f1376,f1377,f1379,f1381,f1382,f1384,f1387,f1389,f1390,f1392,f1394,f1395,f1397,f1402,f1404,f1405,f1407,f1409,f1410,f1246,f1415,f1417,f1418,f1420,f1422,f1423,f1016,f1429,f1431,f1432,f1433,f1436,f1437,f1438,f1442,f1444,f1445,f1447,f1449,f1450,f1452,f1457,f1464,f1466,f1467,f1468,f1471,f1472,f1475,f1481,f1483,f1484,f1486,f1488,f1489,f1491,f1494,f1496,f1497,f1499,f1501,f1502,f1504,f1506,f1508,f1509,f1511,f1513,f1514,f1516,f1521,f1523,f1524,f1525,f1528,f1529,f1532,f1536,f1538,f1539,f1541,f1543,f1544,f1548,f1566,f1568,f1569,f1571,f1573,f1574,f1578,f1580,f1581,f1583,f1585,f1586,f1588,f1592,f1594,f1595,f1596,f1599,f1600,f1603,f1624,f1620,f1619,f1617,f1615,f1614,f1612,f1627,f1628,f1630,f1631,f1633,f1635,f1636,f364,f1687,f1689,f1690,f1692,f1694,f1695,f1576,f1711,f1713,f1714,f1716,f1718,f1719,f1723,f1728,f1730,f1731,f1733,f1735,f1736,f1740,f1752,f1754,f1755,f1757,f1759,f1760,f1763,f1767,f1769,f1770,f1772,f1774,f1775,f1779,f1790,f1792,f1793,f1795,f1797,f1798,f1800,f1806,f1808,f1809,f1810,f1813,f1814,f1831,f1820,f1822,f1823,f1824,f1827,f1828,f1832,f1849,f1851,f1852,f1854,f1856,f1857,f1861,f1866,f1868,f1869,f1871,f1873,f1874,f1878,f1883,f1885,f1886,f1888,f1890,f1891,f1895,f1900,f1902,f1903,f1905,f1907,f1908,f1912,f1927,f1929,f1930,f1931,f1934,f1935,f1950,f1939,f1941,f1942,f1943,f1946,f1947,f1951,f1953,f1955,f1956,f1958,f1960,f1961,f1963,f1966,f1968,f1969,f1971,f1973,f1974,f1976,f1981,f1983,f1984,f1986,f1988,f1989,f1993,f1996,f1998,f1999,f2001,f2003,f2004,f2008,f2035,f2037,f2038,f2040,f2042,f2043,f2045,f2049,f2051,f2052,f2054,f2056,f2057,f2059,f2062,f2064,f2065,f2067,f2069,f2070,f2072,f2082,f2084,f2085,f2087,f2089,f2090,f2094,f2097,f2099,f2100,f2102,f2104,f2105,f2109,f2112,f2114,f2115,f2117,f2119,f2120,f2122,f2125,f2127,f2128,f2130,f2132,f2133,f2135,f2140,f2142,f2143,f2145,f2147,f2148,f2155,f2161,f2163,f2164,f2166,f2168,f2169,f2171,f2176,f2178,f2179,f167,f2180,f171,f2183,f170,f2184,f2187,f172,f2226,f2227,f2229,f2230])).
fof(f2230,plain,(
  $false | ($spl1 | $spl10)),
  inference(subsumption_resolution,[],[f2210,f1414])).
fof(f1414,plain,(
  ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | $spl1),
  inference(resolution,[],[f1251,f103])).
fof(f103,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) | ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X5,X4,X3,X2,X1,X0)),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X5,X4,X3,X2,X1,X0)),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) | ! [X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X7),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X7),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))))) & ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(X5,X4,X3,X2,X1,X0)),s('type.num.num',X5)))) & ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X5,X4,X3,X2,X1,X0)),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X5,X4,X3,X2,X1,X0)),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) | ! [X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X7),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X7),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))))) & (? [X6] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X5)))) & ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))))),
  inference(rectify,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) | ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))))) & (? [X6] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X5)))) & ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) <=> ! [X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) <=> ! [X6] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X5)))) => p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) <=> ! [X6] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X5)))) => p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X6),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X15] : ! [X23,X16,X13,X12,X9] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X9),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X15)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X16),s('type.closSem.state'(X15),X12))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X15)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X23),s('type.closSem.state'(X15),X13)))))) <=> ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9)))) => p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X15)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp',X16),s('type.closSem.state'(X15),X12))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X15)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp',X23),s('type.closSem.state'(X15),X13))))))))),
  file('thm.clos_relationProps.exp_rel_thm','thm.clos_relation.exec_rel_rw')).
fof(f1251,plain,(
  ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | $spl1),
  inference(resolution,[],[f141,f110])).
fof(f110,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X2,X1,X0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK11(X2,X1,X0)))),s('type.closSem.state'(X0),sK13(X2,X1,X0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK12(X2,X1,X0)))),s('type.closSem.state'(X0),sK14(X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | ! [X8,X9,X10,X11,X12] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X8),s('type.closSem.state'(X0),X11),s('type.closSem.state'(X0),X12)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X8))),s('type.list.list'('type.closSem.v'),X9),s('type.list.list'('type.closSem.v'),X10)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X8),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(X0),X11))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(X0),X12)))))))) & ((p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X2,X1,X0)),s('type.closSem.state'(X0),sK13(X2,X1,X0)),s('type.closSem.state'(X0),sK14(X2,X1,X0))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',sK10(X2,X1,X0)))),s('type.list.list'('type.closSem.v'),sK11(X2,X1,X0)),s('type.list.list'('type.closSem.v'),sK12(X2,X1,X0))))) & ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X2,X1,X0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK11(X2,X1,X0)))),s('type.closSem.state'(X0),sK13(X2,X1,X0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK12(X2,X1,X0)))),s('type.closSem.state'(X0),sK14(X2,X1,X0)))))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13,sK14])],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | ! [X8,X9,X10,X11,X12] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X8),s('type.closSem.state'(X0),X11),s('type.closSem.state'(X0),X12)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X8))),s('type.list.list'('type.closSem.v'),X9),s('type.list.list'('type.closSem.v'),X10)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X8),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(X0),X11))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(X0),X12)))))))) & (? [X3,X4,X5,X6,X7] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) & ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(X0),X7))))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2))))))),
  inference(rectify,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | ! [X3,X4,X5,X6,X7] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(X0),X7)))))))) & (? [X3,X4,X5,X6,X7] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) & ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(X0),X7))))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2))))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) <=> ! [X3,X4,X5,X6,X7] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(X0),X7))))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) <=> ! [X3,X4,X5,X6,X7] : ((~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(X0),X7))))))))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) <=> ! [X3,X4,X5,X6,X7] : ((p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5))))) => p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(X0),X7))))))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) <=> ! [X3,X4,X5,X6,X7] : ((p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5))))) => p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(X0),X7))))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X7),s('type.list.list'('type.closLang.exp'),X8)))) <=> ! [X9,X10,X11,X12,X13] : ((p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(X6),X12),s('type.closSem.state'(X6),X13)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11))))) => p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X9),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X6)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(X6),X12))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X6)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X8),s('type.list.list'('type.closSem.v'),X11))),s('type.closSem.state'(X6),X13))))))))),
  file('thm.clos_relationProps.exp_rel_thm','thm.clos_relation.exp_rel_def')).
fof(f2210,plain,(
  p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ($spl1 | $spl10)),
  inference(resolution,[],[f2153,f1251])).
fof(f2153,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X4),X0))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X4),X1))))))) ) | ($spl1 | $spl10)),
  inference(subsumption_resolution,[],[f2137,f141])).
fof(f2137,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X4),X0))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X4),X1)))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))) ) | ($spl1 | $spl10)),
  inference(resolution,[],[f1788,f1250])).
fof(f1250,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) | $spl1),
  inference(resolution,[],[f141,f109])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',sK10(X2,X1,X0)))),s('type.list.list'('type.closSem.v'),sK11(X2,X1,X0)),s('type.list.list'('type.closSem.v'),sK12(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f1788,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X7,X8,sK0)))),s('type.list.list'('type.closSem.v'),X9),s('type.list.list'('type.closSem.v'),X10)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X7,X8,sK0),X11,X12,X13,X14,X15)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(sK0),sK13(X7,X8,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X7,X8,sK0),X11,X12,X13,X14,X15)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(sK0),sK14(X7,X8,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X7,X8,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X15)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X13),s('type.closSem.state'(X15),X11))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X15)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X14),s('type.closSem.state'(X15),X12)))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X8),s('type.list.list'('type.closLang.exp'),X7))))) ) | $spl10),
  inference(resolution,[],[f807,f108])).
fof(f108,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X2,X1,X0)),s('type.closSem.state'(X0),sK13(X2,X1,X0)),s('type.closSem.state'(X0),sK14(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f807,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X0),s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),X2)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X0))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X0,X5,X6,X7,X8,X9)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),X1))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X0,X5,X6,X7,X8,X9)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(sK0),X2)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X9)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X7),s('type.closSem.state'(X9),X5))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X9)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X8),s('type.closSem.state'(X9),X6))))))) ) | $spl10),
  inference(resolution,[],[f152,f102])).
fof(f102,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(X5,X4,X3,X2,X1,X0)),s('type.num.num',X5))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f2229,plain,(
  $false | ($spl1 | $spl10)),
  inference(subsumption_resolution,[],[f2228,f141])).
fof(f2228,plain,(
  p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) | ($spl1 | $spl10)),
  inference(subsumption_resolution,[],[f2209,f1414])).
fof(f2209,plain,(
  p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) | ($spl1 | $spl10)),
  inference(resolution,[],[f2153,f110])).
fof(f2227,plain,(
  $false | ($spl1 | $spl10)),
  inference(subsumption_resolution,[],[f2225,f1251])).
fof(f2225,plain,(
  p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ($spl1 | $spl10)),
  inference(duplicate_literal_removal,[],[f2197])).
fof(f2197,plain,(
  p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ($spl1 | $spl10)),
  inference(resolution,[],[f2153,f103])).
fof(f2226,plain,(
  $false | ($spl1 | $spl10)),
  inference(subsumption_resolution,[],[f2196,f1251])).
fof(f2196,plain,(
  p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ($spl1 | $spl10)),
  inference(resolution,[],[f2153,f1414])).
fof(f172,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl24),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2187,plain,(
  ( ! [X10,X13,X11] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closLang.exp'),X11)))) | s(bool,X13) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | p(s(bool,X13))) ) | $spl170),
  inference(cnf_transformation,[],[f2187_D])).
fof(f2187_D,plain,(
  ( ! [X10,X13,X11] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closLang.exp'),X11)))) | s(bool,X13) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | p(s(bool,X13))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f2184,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X50),s('type.list.list'('type.closLang.exp'),X51)))) | p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X50),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X51),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4 | $spl106)),
  inference(resolution,[],[f1684,f137])).
fof(f137,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f121])).
fof(f121,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X23,X17,X24] : ((s(bool,X16) = s(bool,X23) & (p(s(bool,X23)) => s(bool,X17) = s(bool,X24))) => ((p(s(bool,X16)) => p(s(bool,X17))) <=> (p(s(bool,X23)) => p(s(bool,X24)))))),
  file('thm.clos_relationProps.exp_rel_thm','thm.bool.IMP_CONG')).
fof(f1684,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X1))))) ) | ($spl3 | $spl4 | $spl106)),
  inference(resolution,[],[f1652,f1197])).
fof(f1197,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)))) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f102])).
fof(f954,plain,(
  ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ($spl3 | $spl4)),
  inference(resolution,[],[f633,f145])).
fof(f633,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X0)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))))) ) | $spl3),
  inference(resolution,[],[f143,f101])).
fof(f101,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X7),s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X7),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X5)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X5),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X0),X4))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f1652,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK3)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X2),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X2),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X1))))) ) | $spl106),
  inference(resolution,[],[f1553,f101])).
fof(f1553,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X1))))) ) | $spl106),
  inference(equality_resolution,[],[f1532])).
fof(f170,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) | $spl20),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2183,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X45),s('type.list.list'('type.closLang.exp'),X46)))) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X45),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X46),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X48) != s(bool,X49)) ) | ($spl3 | $spl4 | $spl106)),
  inference(resolution,[],[f1684,f124])).
fof(f124,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f171,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl22),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2180,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X30),s('type.list.list'('type.closLang.exp'),X31)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X30),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4 | $spl106)),
  inference(resolution,[],[f1684,f123])).
fof(f123,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f167,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl14),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2179,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closLang.exp'),X26)))) | p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X26),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4 | $spl106)),
  inference(resolution,[],[f1684,f122])).
fof(f122,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f2178,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closLang.exp'),X21)))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X21),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X23) != s(bool,X24)) ) | ($spl3 | $spl4 | $spl106)),
  inference(resolution,[],[f1684,f120])).
fof(f120,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f2176,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closLang.exp'),X11)))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4 | $spl106)),
  inference(resolution,[],[f1684,f118])).
fof(f118,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f2171,plain,(
  ( ! [X10,X13,X11,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK10(sK2,sK1,sK0)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X11,X10,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X11,X10,sK0)),s('type.list.list'('type.closSem.v'),sK12(X11,X10,sK0)))))) ) | $spl168),
  inference(cnf_transformation,[],[f2171_D])).
fof(f2171_D,plain,(
  ( ! [X10,X13,X11,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK10(sK2,sK1,sK0)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X11,X10,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X11,X10,sK0)),s('type.list.list'('type.closSem.v'),sK12(X11,X10,sK0)))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f2169,plain,(
  ( ! [X61,X59,X57,X60,X58] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X57),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X58),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X57),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X59),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X59,X58,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X59,X58,sK0)),s('type.list.list'('type.closSem.v'),sK12(X59,X58,sK0))))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X57),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1924,f137])).
fof(f1924,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK10(sK2,sK1,sK0))))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X2),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X2),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X0,X1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X0,X1,sK0)),s('type.list.list'('type.closSem.v'),sK12(X0,X1,sK0)))))) ) | $spl1),
  inference(resolution,[],[f1765,f101])).
fof(f1765,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X3),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X3,X2,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X3,X2,sK0)),s('type.list.list'('type.closSem.v'),sK12(X3,X2,sK0)))))) ) | $spl1),
  inference(resolution,[],[f1762,f109])).
fof(f1762,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closLang.exp'),X3)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X3),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) | $spl1),
  inference(subsumption_resolution,[],[f1751,f141])).
fof(f1751,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X3),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closLang.exp'),X3)))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))) ) | $spl1),
  inference(resolution,[],[f1309,f108])).
fof(f1309,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),X3)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X1))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X3)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X2))))) ) | $spl1),
  inference(resolution,[],[f1250,f107])).
fof(f107,plain,(
  ( ! [X2,X0,X12,X10,X8,X1,X11,X9] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X8),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(X0),X11))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(X0),X12)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X8))),s('type.list.list'('type.closSem.v'),X9),s('type.list.list'('type.closSem.v'),X10)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X8),s('type.closSem.state'(X0),X11),s('type.closSem.state'(X0),X12)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f2168,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X51),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X52),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X51),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X53,X52,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X53,X52,sK0)),s('type.list.list'('type.closSem.v'),sK12(X53,X52,sK0))))) | ~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num',sK10(sK2,sK1,sK0)))) | s(bool,X54) != s(bool,X56)) ) | $spl1),
  inference(resolution,[],[f1924,f124])).
fof(f2166,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X39),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X39),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X41),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X41,X40,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X41,X40,sK0)),s('type.list.list'('type.closSem.v'),sK12(X41,X40,sK0))))) | p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X39),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1924,f123])).
fof(f2164,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X27),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X27),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X29),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X29,X28,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X29,X28,sK0)),s('type.list.list'('type.closSem.v'),sK12(X29,X28,sK0))))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1924,f122])).
fof(f2163,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X21),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X22),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X21),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X23,X22,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X23,X22,sK0)),s('type.list.list'('type.closSem.v'),sK12(X23,X22,sK0))))) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X21),s('type.num.num',sK10(sK2,sK1,sK0)))) | s(bool,X24) != s(bool,X26)) ) | $spl1),
  inference(resolution,[],[f1924,f120])).
fof(f2161,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X11,X10,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X11,X10,sK0)),s('type.list.list'('type.closSem.v'),sK12(X11,X10,sK0))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1924,f118])).
fof(f2155,plain,(
  ( ! [X26,X24,X23,X21,X19,X29,X27,X25,X22,X20] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X19,X20,sK0),X21,X22,X23,X24,X25)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X26))),s('type.closSem.state'(sK0),sK13(X19,X20,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X19,X20,sK0),X21,X22,X23,X24,X25)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X27))),s('type.closSem.state'(sK0),sK14(X19,X20,sK0))))))) | s(bool,X29) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X19,X20,sK0)))),s('type.list.list'('type.closSem.v'),X26),s('type.list.list'('type.closSem.v'),X27))) | ~p(s(bool,X29)) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closLang.exp'),X19)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X19,X20,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X25)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X23),s('type.closSem.state'(X25),X21))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X25)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X25),X22))))))) ) | $spl166),
  inference(cnf_transformation,[],[f2155_D])).
fof(f2155_D,plain,(
  ( ! [X26,X24,X23,X21,X19,X29,X27,X25,X22,X20] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X19,X20,sK0),X21,X22,X23,X24,X25)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X26))),s('type.closSem.state'(sK0),sK13(X19,X20,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X19,X20,sK0),X21,X22,X23,X24,X25)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X27))),s('type.closSem.state'(sK0),sK14(X19,X20,sK0))))))) | s(bool,X29) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X19,X20,sK0)))),s('type.list.list'('type.closSem.v'),X26),s('type.list.list'('type.closSem.v'),X27))) | ~p(s(bool,X29)) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closLang.exp'),X19)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X19,X20,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X25)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X23),s('type.closSem.state'(X25),X21))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X25)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X25),X22))))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f2148,plain,(
  ( ! [X125,X123,X121,X118,X116,X124,X122,X120,X119,X117,X115] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X115,X116,sK0),X117,X118,X119,X120,X121)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X122))),s('type.closSem.state'(sK0),sK13(X115,X116,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X115,X116,sK0),X117,X118,X119,X120,X121)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X123))),s('type.closSem.state'(sK0),sK14(X115,X116,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X115,X116,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X121)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X119),s('type.closSem.state'(X121),X117))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X121)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X120),s('type.closSem.state'(X121),X118)))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X116),s('type.list.list'('type.closLang.exp'),X115)))) | p(s(bool,X124)) | ~p(s(bool,X125)) | s(bool,X125) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X115,X116,sK0)))),s('type.list.list'('type.closSem.v'),X122),s('type.list.list'('type.closSem.v'),X123)))) ) | $spl10),
  inference(resolution,[],[f1788,f137])).
fof(f2147,plain,(
  ( ! [X103,X111,X109,X107,X105,X114,X112,X110,X108,X106,X104,X113] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X103,X104,sK0),X105,X106,X107,X108,X109)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X110))),s('type.closSem.state'(sK0),sK13(X103,X104,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X103,X104,sK0),X105,X106,X107,X108,X109)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X111))),s('type.closSem.state'(sK0),sK14(X103,X104,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X103,X104,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X109)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X107),s('type.closSem.state'(X109),X105))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X109)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X108),s('type.closSem.state'(X109),X106)))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X104),s('type.list.list'('type.closLang.exp'),X103)))) | ~p(s(bool,X112)) | ~p(s(bool,X113)) | s(bool,X113) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X103,X104,sK0)))),s('type.list.list'('type.closSem.v'),X110),s('type.list.list'('type.closSem.v'),X111))) | s(bool,X112) != s(bool,X114)) ) | $spl10),
  inference(resolution,[],[f1788,f124])).
fof(f2145,plain,(
  ( ! [X80,X90,X88,X87,X85,X83,X81,X89,X79,X86,X84,X82] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X79,X80,sK0),X81,X82,X83,X84,X85)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X86))),s('type.closSem.state'(sK0),sK13(X79,X80,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X79,X80,sK0),X81,X82,X83,X84,X85)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X87))),s('type.closSem.state'(sK0),sK14(X79,X80,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X79,X80,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X85)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X83),s('type.closSem.state'(X85),X81))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X85)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X84),s('type.closSem.state'(X85),X82)))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closLang.exp'),X79)))) | p(s(bool,X88)) | ~p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X89) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X79,X80,sK0)))),s('type.list.list'('type.closSem.v'),X86),s('type.list.list'('type.closSem.v'),X87)))) ) | $spl10),
  inference(resolution,[],[f1788,f123])).
fof(f2143,plain,(
  ( ! [X61,X59,X57,X66,X64,X62,X60,X58,X56,X65,X55,X63] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X55,X56,sK0),X57,X58,X59,X60,X61)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X62))),s('type.closSem.state'(sK0),sK13(X55,X56,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X55,X56,sK0),X57,X58,X59,X60,X61)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X63))),s('type.closSem.state'(sK0),sK14(X55,X56,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X55,X56,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X61)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X59),s('type.closSem.state'(X61),X57))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X61)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X60),s('type.closSem.state'(X61),X58)))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closLang.exp'),X55)))) | p(s(bool,X64)) | ~p(s(bool,X65)) | s(bool,X64) != s(bool,X66) | s(bool,X65) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X55,X56,sK0)))),s('type.list.list'('type.closSem.v'),X62),s('type.list.list'('type.closSem.v'),X63)))) ) | $spl10),
  inference(resolution,[],[f1788,f122])).
fof(f2142,plain,(
  ( ! [X47,X45,X43,X54,X52,X50,X48,X46,X44,X53,X51,X49] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X43,X44,sK0),X45,X46,X47,X48,X49)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X50))),s('type.closSem.state'(sK0),sK13(X43,X44,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X43,X44,sK0),X45,X46,X47,X48,X49)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X51))),s('type.closSem.state'(sK0),sK14(X43,X44,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X43,X44,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X49)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X47),s('type.closSem.state'(X49),X45))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X49)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X48),s('type.closSem.state'(X49),X46)))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X44),s('type.list.list'('type.closLang.exp'),X43)))) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X43,X44,sK0)))),s('type.list.list'('type.closSem.v'),X50),s('type.list.list'('type.closSem.v'),X51))) | s(bool,X52) != s(bool,X54)) ) | $spl10),
  inference(resolution,[],[f1788,f120])).
fof(f2140,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X19,X29,X27,X25,X22,X20] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X19,X20,sK0),X21,X22,X23,X24,X25)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X26))),s('type.closSem.state'(sK0),sK13(X19,X20,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(X19,X20,sK0),X21,X22,X23,X24,X25)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X27))),s('type.closSem.state'(sK0),sK14(X19,X20,sK0))))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(X19,X20,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X25)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X23),s('type.closSem.state'(X25),X21))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X25)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X25),X22)))))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closLang.exp'),X19)))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X19,X20,sK0)))),s('type.list.list'('type.closSem.v'),X26),s('type.list.list'('type.closSem.v'),X27)))) ) | $spl10),
  inference(resolution,[],[f1788,f118])).
fof(f2135,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X16,X15,sK0)),s('type.closSem.state'(sK0),sK13(X16,X15,sK0)),s('type.closSem.state'(sK0),sK14(X16,X15,sK0))),X12)))) | ~p(s(bool,X18)) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) | $spl164),
  inference(cnf_transformation,[],[f2135_D])).
fof(f2135_D,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X16,X15,sK0)),s('type.closSem.state'(sK0),sK13(X16,X15,sK0)),s('type.closSem.state'(sK0),sK14(X16,X15,sK0))),X12)))) | ~p(s(bool,X18)) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f2133,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK16(X77,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X80,X79,sK0)),s('type.closSem.state'(sK0),sK13(X80,X79,sK0)),s('type.closSem.state'(sK0),sK14(X80,X79,sK0))),X76))))) ) | $spl1),
  inference(resolution,[],[f1819,f137])).
fof(f1819,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,happ(s(fun(X24,bool),X25),s(X24,sK16(X25,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X23,X22,sK0)),s('type.closSem.state'(sK0),sK13(X23,X22,sK0)),s('type.closSem.state'(sK0),sK14(X23,X22,sK0))),X24))))) | p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X22),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) | $spl1),
  inference(resolution,[],[f1764,f116])).
fof(f116,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0)))))) | ~p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(X2,X1,X0)))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((? [X6] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ~p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : (! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) => ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) => ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X15] : ! [X18,X19] : (! [X16] : (p(s(bool,X18)) => p(s(bool,happ(s(fun(X15,bool),X19),s(X15,X16))))) <=> (p(s(bool,X18)) => ! [X16] : p(s(bool,happ(s(fun(X15,bool),X19),s(X15,X16))))))),
  file('thm.clos_relationProps.exp_rel_thm','thm.bool.RIGHT_FORALL_IMP_THM')).
fof(f1764,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X1,X0,sK0)),s('type.closSem.state'(sK0),sK13(X1,X0,sK0)),s('type.closSem.state'(sK0),sK14(X1,X0,sK0)))))) ) | $spl1),
  inference(resolution,[],[f1762,f108])).
fof(f2132,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X72),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(X68,bool),X69),s(X68,sK16(X69,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X72,X71,sK0)),s('type.closSem.state'(sK0),sK13(X72,X71,sK0)),s('type.closSem.state'(sK0),sK14(X72,X71,sK0))),X68)))) | s(bool,X73) != s(bool,X75)) ) | $spl1),
  inference(resolution,[],[f1819,f124])).
fof(f2130,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(X52,bool),X53),s(X52,sK16(X53,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X56,X55,sK0)),s('type.closSem.state'(sK0),sK13(X56,X55,sK0)),s('type.closSem.state'(sK0),sK14(X56,X55,sK0))),X52))))) ) | $spl1),
  inference(resolution,[],[f1819,f123])).
fof(f2128,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X37),s(X36,sK16(X37,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X40,X39,sK0)),s('type.closSem.state'(sK0),sK13(X40,X39,sK0)),s('type.closSem.state'(sK0),sK14(X40,X39,sK0))),X36))))) ) | $spl1),
  inference(resolution,[],[f1819,f122])).
fof(f2127,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X28,bool),X29),s(X28,sK16(X29,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X32,X31,sK0)),s('type.closSem.state'(sK0),sK13(X32,X31,sK0)),s('type.closSem.state'(sK0),sK14(X32,X31,sK0))),X28)))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1819,f120])).
fof(f2125,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X16,X15,sK0)),s('type.closSem.state'(sK0),sK13(X16,X15,sK0)),s('type.closSem.state'(sK0),sK14(X16,X15,sK0))),X12))))) ) | $spl1),
  inference(resolution,[],[f1819,f118])).
fof(f2122,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X16,X15,sK0)),s('type.closSem.state'(sK0),sK13(X16,X15,sK0)),s('type.closSem.state'(sK0),sK14(X16,X15,sK0))),X12)))) | ~p(s(bool,X18)) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) | $spl162),
  inference(cnf_transformation,[],[f2122_D])).
fof(f2122_D,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X16,X15,sK0)),s('type.closSem.state'(sK0),sK13(X16,X15,sK0)),s('type.closSem.state'(sK0),sK14(X16,X15,sK0))),X12)))) | ~p(s(bool,X18)) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f2120,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK15(X77,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X80,X79,sK0)),s('type.closSem.state'(sK0),sK13(X80,X79,sK0)),s('type.closSem.state'(sK0),sK14(X80,X79,sK0))),X76))))) ) | $spl1),
  inference(resolution,[],[f1818,f137])).
fof(f1818,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X20,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X18,X17,sK0)),s('type.closSem.state'(sK0),sK13(X18,X17,sK0)),s('type.closSem.state'(sK0),sK14(X18,X17,sK0))),X19))))) | p(s(bool,happ(s(fun(X19,bool),X20),s(X19,X21)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X17),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X18),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) | $spl1),
  inference(resolution,[],[f1764,f114])).
fof(f114,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f2119,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X72),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(X68,bool),X69),s(X68,sK15(X69,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X72,X71,sK0)),s('type.closSem.state'(sK0),sK13(X72,X71,sK0)),s('type.closSem.state'(sK0),sK14(X72,X71,sK0))),X68)))) | s(bool,X73) != s(bool,X75)) ) | $spl1),
  inference(resolution,[],[f1818,f124])).
fof(f2117,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(X52,bool),X53),s(X52,sK15(X53,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X56,X55,sK0)),s('type.closSem.state'(sK0),sK13(X56,X55,sK0)),s('type.closSem.state'(sK0),sK14(X56,X55,sK0))),X52))))) ) | $spl1),
  inference(resolution,[],[f1818,f123])).
fof(f2115,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X37),s(X36,sK15(X37,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X40,X39,sK0)),s('type.closSem.state'(sK0),sK13(X40,X39,sK0)),s('type.closSem.state'(sK0),sK14(X40,X39,sK0))),X36))))) ) | $spl1),
  inference(resolution,[],[f1818,f122])).
fof(f2114,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X28,bool),X29),s(X28,sK15(X29,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X32,X31,sK0)),s('type.closSem.state'(sK0),sK13(X32,X31,sK0)),s('type.closSem.state'(sK0),sK14(X32,X31,sK0))),X28)))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1818,f120])).
fof(f2112,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X16,X15,sK0)),s('type.closSem.state'(sK0),sK13(X16,X15,sK0)),s('type.closSem.state'(sK0),sK14(X16,X15,sK0))),X12))))) ) | $spl1),
  inference(resolution,[],[f1818,f118])).
fof(f2109,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X15,X16,sK0)),s('type.closSem.state'(sK0),sK13(X15,X16,sK0)),s('type.closSem.state'(sK0),sK14(X15,X16,sK0)))))) ) | $spl160),
  inference(cnf_transformation,[],[f2109_D])).
fof(f2109_D,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X15,X16,sK0)),s('type.closSem.state'(sK0),sK13(X15,X16,sK0)),s('type.closSem.state'(sK0),sK14(X15,X16,sK0)))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f2105,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X79,X80,sK0)),s('type.closSem.state'(sK0),sK13(X79,X80,sK0)),s('type.closSem.state'(sK0),sK14(X79,X80,sK0))))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK16(X77,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X76))))) ) | $spl1),
  inference(resolution,[],[f1805,f137])).
fof(f1805,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X10,bool),X11),s(X10,sK16(X11,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X9),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X8),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X10))))) | p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X8,X9,sK0)),s('type.closSem.state'(sK0),sK13(X8,X9,sK0)),s('type.closSem.state'(sK0),sK14(X8,X9,sK0)))))) ) | $spl1),
  inference(resolution,[],[f1764,f116])).
fof(f2104,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X71,X72,sK0)),s('type.closSem.state'(sK0),sK13(X71,X72,sK0)),s('type.closSem.state'(sK0),sK14(X71,X72,sK0))))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(X68,bool),X69),s(X68,sK16(X69,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X72),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X68)))) | s(bool,X73) != s(bool,X75)) ) | $spl1),
  inference(resolution,[],[f1805,f124])).
fof(f2102,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X55,X56,sK0)),s('type.closSem.state'(sK0),sK13(X55,X56,sK0)),s('type.closSem.state'(sK0),sK14(X55,X56,sK0))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(X52,bool),X53),s(X52,sK16(X53,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X52))))) ) | $spl1),
  inference(resolution,[],[f1805,f123])).
fof(f2100,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X39,X40,sK0)),s('type.closSem.state'(sK0),sK13(X39,X40,sK0)),s('type.closSem.state'(sK0),sK14(X39,X40,sK0))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X37),s(X36,sK16(X37,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X36))))) ) | $spl1),
  inference(resolution,[],[f1805,f122])).
fof(f2099,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X31,X32,sK0)),s('type.closSem.state'(sK0),sK13(X31,X32,sK0)),s('type.closSem.state'(sK0),sK14(X31,X32,sK0))))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X28,bool),X29),s(X28,sK16(X29,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X28)))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1805,f120])).
fof(f2097,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X15,X16,sK0)),s('type.closSem.state'(sK0),sK13(X15,X16,sK0)),s('type.closSem.state'(sK0),sK14(X15,X16,sK0))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X12))))) ) | $spl1),
  inference(resolution,[],[f1805,f118])).
fof(f2094,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X15,X16,sK0)),s('type.closSem.state'(sK0),sK13(X15,X16,sK0)),s('type.closSem.state'(sK0),sK14(X15,X16,sK0)))))) ) | $spl158),
  inference(cnf_transformation,[],[f2094_D])).
fof(f2094_D,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X15,X16,sK0)),s('type.closSem.state'(sK0),sK13(X15,X16,sK0)),s('type.closSem.state'(sK0),sK14(X15,X16,sK0)))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f2090,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X79,X80,sK0)),s('type.closSem.state'(sK0),sK13(X79,X80,sK0)),s('type.closSem.state'(sK0),sK14(X79,X80,sK0))))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK15(X77,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X76))))) ) | $spl1),
  inference(resolution,[],[f1804,f137])).
fof(f1804,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK15(X6,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X3),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X5))))) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X3,X4,sK0)),s('type.closSem.state'(sK0),sK13(X3,X4,sK0)),s('type.closSem.state'(sK0),sK14(X3,X4,sK0)))))) ) | $spl1),
  inference(resolution,[],[f1764,f114])).
fof(f2089,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X71,X72,sK0)),s('type.closSem.state'(sK0),sK13(X71,X72,sK0)),s('type.closSem.state'(sK0),sK14(X71,X72,sK0))))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(X68,bool),X69),s(X68,sK15(X69,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X72),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X68)))) | s(bool,X73) != s(bool,X75)) ) | $spl1),
  inference(resolution,[],[f1804,f124])).
fof(f2087,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X55,X56,sK0)),s('type.closSem.state'(sK0),sK13(X55,X56,sK0)),s('type.closSem.state'(sK0),sK14(X55,X56,sK0))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(X52,bool),X53),s(X52,sK15(X53,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X52))))) ) | $spl1),
  inference(resolution,[],[f1804,f123])).
fof(f2085,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X39,X40,sK0)),s('type.closSem.state'(sK0),sK13(X39,X40,sK0)),s('type.closSem.state'(sK0),sK14(X39,X40,sK0))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X37),s(X36,sK15(X37,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X36))))) ) | $spl1),
  inference(resolution,[],[f1804,f122])).
fof(f2084,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X31,X32,sK0)),s('type.closSem.state'(sK0),sK13(X31,X32,sK0)),s('type.closSem.state'(sK0),sK14(X31,X32,sK0))))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X28,bool),X29),s(X28,sK15(X29,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X28)))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1804,f120])).
fof(f2082,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X15,X16,sK0)),s('type.closSem.state'(sK0),sK13(X15,X16,sK0)),s('type.closSem.state'(sK0),sK14(X15,X16,sK0))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))),X12))))) ) | $spl1),
  inference(resolution,[],[f1804,f118])).
fof(f2072,plain,(
  ( ! [X2,X0] : (s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X0))) | ~p(s(bool,X2)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) | $spl156),
  inference(cnf_transformation,[],[f2072_D])).
fof(f2072_D,plain,(
  ( ! [X2,X0] : (s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X0))) | ~p(s(bool,X2)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f2070,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X32),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X32)))) ) | $spl1),
  inference(resolution,[],[f2034,f137])).
fof(f2034,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X0)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))))) ) | $spl1),
  inference(resolution,[],[f1414,f101])).
fof(f2069,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X28),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X28))) | s(bool,X29) != s(bool,X31)) ) | $spl1),
  inference(resolution,[],[f2034,f124])).
fof(f2067,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X20),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X20)))) ) | $spl1),
  inference(resolution,[],[f2034,f123])).
fof(f2065,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X12),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X12)))) ) | $spl1),
  inference(resolution,[],[f2034,f122])).
fof(f2064,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X8),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X8))) | s(bool,X9) != s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f2034,f120])).
fof(f2062,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',X0)))) ) | $spl1),
  inference(resolution,[],[f2034,f118])).
fof(f2059,plain,(
  ( ! [X10,X13,X11,X9] : (s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK10(sK2,sK1,sK0)))) | ~p(s(bool,X13)) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X11,X10,sK0)),s('type.closSem.state'(sK0),sK13(X11,X10,sK0)),s('type.closSem.state'(sK0),sK14(X11,X10,sK0)))))) ) | $spl154),
  inference(cnf_transformation,[],[f2059_D])).
fof(f2059_D,plain,(
  ( ! [X10,X13,X11,X9] : (s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK10(sK2,sK1,sK0)))) | ~p(s(bool,X13)) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X11,X10,sK0)),s('type.closSem.state'(sK0),sK13(X11,X10,sK0)),s('type.closSem.state'(sK0),sK14(X11,X10,sK0)))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f2057,plain,(
  ( ! [X61,X59,X57,X60,X58] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X57),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X58),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X57),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X59),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X59,X58,sK0)),s('type.closSem.state'(sK0),sK13(X59,X58,sK0)),s('type.closSem.state'(sK0),sK14(X59,X58,sK0))))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X57),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1803,f137])).
fof(f1803,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK10(sK2,sK1,sK0))))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X2),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X2),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X0,X1,sK0)),s('type.closSem.state'(sK0),sK13(X0,X1,sK0)),s('type.closSem.state'(sK0),sK14(X0,X1,sK0)))))) ) | $spl1),
  inference(resolution,[],[f1764,f101])).
fof(f2056,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X51),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X52),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X51),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X53,X52,sK0)),s('type.closSem.state'(sK0),sK13(X53,X52,sK0)),s('type.closSem.state'(sK0),sK14(X53,X52,sK0))))) | ~p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num',sK10(sK2,sK1,sK0)))) | s(bool,X54) != s(bool,X56)) ) | $spl1),
  inference(resolution,[],[f1803,f124])).
fof(f2054,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X39),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X39),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X41),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X41,X40,sK0)),s('type.closSem.state'(sK0),sK13(X41,X40,sK0)),s('type.closSem.state'(sK0),sK14(X41,X40,sK0))))) | p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X39),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1803,f123])).
fof(f2052,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X27),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X27),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X29),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X29,X28,sK0)),s('type.closSem.state'(sK0),sK13(X29,X28,sK0)),s('type.closSem.state'(sK0),sK14(X29,X28,sK0))))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1803,f122])).
fof(f2051,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X21),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X22),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X21),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X23,X22,sK0)),s('type.closSem.state'(sK0),sK13(X23,X22,sK0)),s('type.closSem.state'(sK0),sK14(X23,X22,sK0))))) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X21),s('type.num.num',sK10(sK2,sK1,sK0)))) | s(bool,X24) != s(bool,X26)) ) | $spl1),
  inference(resolution,[],[f1803,f120])).
fof(f2049,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X11,X10,sK0)),s('type.closSem.state'(sK0),sK13(X11,X10,sK0)),s('type.closSem.state'(sK0),sK14(X11,X10,sK0))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1803,f118])).
fof(f2045,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl152),
  inference(cnf_transformation,[],[f2045_D])).
fof(f2045_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f2043,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1414,f137])).
fof(f2042,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f1414,f124])).
fof(f2040,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1414,f123])).
fof(f2038,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1414,f122])).
fof(f2037,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X7) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f1414,f120])).
fof(f2035,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1414,f118])).
fof(f2008,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X6))),s('type.closSem.state'(sK0),sK7)))))) | s(bool,X8) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X5),s('type.list.list'('type.closSem.v'),X6))) | ~p(s(bool,X8)) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X4),X0))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X4),X1))))))) ) | $spl150),
  inference(cnf_transformation,[],[f2008_D])).
fof(f2008_D,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X6))),s('type.closSem.state'(sK0),sK7)))))) | s(bool,X8) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X5),s('type.list.list'('type.closSem.v'),X6))) | ~p(s(bool,X8)) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X4),X0))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X4),X1))))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f2004,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X80,X81,X82,X83,X84)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X85))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X80,X81,X82,X83,X84)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X86))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X84)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X82),s('type.closSem.state'(X84),X80))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X84)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X83),s('type.closSem.state'(X84),X81)))))) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X88) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X85),s('type.list.list'('type.closSem.v'),X86)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f1789,f137])).
fof(f1789,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X16),s('type.list.list'('type.closSem.v'),X17)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X18,X19,X20,X21,X22)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X16))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X18,X19,X20,X21,X22)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X17))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X22)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X20),s('type.closSem.state'(X22),X18))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X22)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X21),s('type.closSem.state'(X22),X19))))))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f807,f149])).
fof(f2003,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X70,X71,X72,X73,X74)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X75))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X70,X71,X72,X73,X74)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X76))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X74)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X72),s('type.closSem.state'(X74),X70))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X74)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X73),s('type.closSem.state'(X74),X71)))))) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X78) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X75),s('type.list.list'('type.closSem.v'),X76))) | s(bool,X77) != s(bool,X79)) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f1789,f124])).
fof(f2001,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X50,X51,X52,X53,X54)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X55))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X50,X51,X52,X53,X54)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X56))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X54)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X52),s('type.closSem.state'(X54),X50))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X54)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X53),s('type.closSem.state'(X54),X51)))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X55),s('type.list.list'('type.closSem.v'),X56)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f1789,f123])).
fof(f1999,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X30,X31,X32,X33,X34)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X35))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X30,X31,X32,X33,X34)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X36))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X34)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X32),s('type.closSem.state'(X34),X30))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X34)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X33),s('type.closSem.state'(X34),X31)))))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X35),s('type.list.list'('type.closSem.v'),X36)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f1789,f122])).
fof(f1998,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X20,X21,X22,X23,X24)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X25))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X20,X21,X22,X23,X24)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X26))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X24)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X24),X20))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X24)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X23),s('type.closSem.state'(X24),X21)))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X25),s('type.list.list'('type.closSem.v'),X26))) | s(bool,X27) != s(bool,X29)) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f1789,f120])).
fof(f1996,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,X0,X1,X2,X3,X4)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X6))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X2),s('type.closSem.state'(X4),X0))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X4)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X4),X1)))))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X5),s('type.list.list'('type.closSem.v'),X6)))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f1789,f118])).
fof(f1993,plain,(
  ( ! [X6,X8,X7,X11,X9] : (s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X7),s('type.list.list'('type.closSem.v'),X9))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),X7))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X8),s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X11)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closLang.exp'),X8))))) ) | $spl148),
  inference(cnf_transformation,[],[f1993_D])).
fof(f1993_D,plain,(
  ( ! [X6,X8,X7,X11,X9] : (s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X7),s('type.list.list'('type.closSem.v'),X9))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),X7))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X8),s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X11)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closLang.exp'),X8))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f1989,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X62),s('type.list.list'('type.closSem.v'),X63))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X64),s('type.list.list'('type.closSem.v'),X65))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X62),s('type.list.list'('type.closLang.exp'),X64)))) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X63),s('type.list.list'('type.closSem.v'),X65)))) ) | $spl42),
  inference(resolution,[],[f1031,f137])).
fof(f1031,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X1),s('type.list.list'('type.closSem.v'),X3)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),X1))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X2))))) ) | $spl42),
  inference(resolution,[],[f377,f107])).
fof(f377,plain,(
  p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) | $spl42),
  inference(equality_resolution,[],[f353])).
fof(f1988,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),X56))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X57),s('type.list.list'('type.closSem.v'),X58))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closLang.exp'),X57)))) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X56),s('type.list.list'('type.closSem.v'),X58))) | s(bool,X59) != s(bool,X61)) ) | $spl42),
  inference(resolution,[],[f1031,f124])).
fof(f1986,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X41),s('type.list.list'('type.closSem.v'),X42))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X43),s('type.list.list'('type.closSem.v'),X44))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X41),s('type.list.list'('type.closLang.exp'),X43)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X42),s('type.list.list'('type.closSem.v'),X44)))) ) | $spl42),
  inference(resolution,[],[f1031,f123])).
fof(f1984,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X32] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X27),s('type.list.list'('type.closSem.v'),X28))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X29),s('type.list.list'('type.closSem.v'),X30))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X27),s('type.list.list'('type.closLang.exp'),X29)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X28),s('type.list.list'('type.closSem.v'),X30)))) ) | $spl42),
  inference(resolution,[],[f1031,f122])).
fof(f1983,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closSem.v'),X21))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X22),s('type.list.list'('type.closSem.v'),X23))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closLang.exp'),X22)))) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X21),s('type.list.list'('type.closSem.v'),X23))) | s(bool,X24) != s(bool,X26)) ) | $spl42),
  inference(resolution,[],[f1031,f120])).
fof(f1981,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),X7))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X8),s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closLang.exp'),X8)))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),X7),s('type.list.list'('type.closSem.v'),X9)))) ) | $spl42),
  inference(resolution,[],[f1031,f118])).
fof(f1976,plain,(
  ( ! [X2,X0] : (s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X0))) | ~p(s(bool,X2)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))))) ) | $spl146),
  inference(cnf_transformation,[],[f1976_D])).
fof(f1976_D,plain,(
  ( ! [X2,X0] : (s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X0))) | ~p(s(bool,X2)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f1974,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X32),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X32)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1952,f137])).
fof(f1952,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X0)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1350,f101])).
fof(f1350,plain,(
  ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f103])).
fof(f1973,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X28),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X28))) | s(bool,X29) != s(bool,X31)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1952,f124])).
fof(f1971,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X20),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X20)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1952,f123])).
fof(f1969,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X12),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X12)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1952,f122])).
fof(f1968,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X8),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X8))) | s(bool,X9) != s(bool,X11)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1952,f120])).
fof(f1966,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',X0)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1952,f118])).
fof(f1963,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,X2))) ) | $spl144),
  inference(cnf_transformation,[],[f1963_D])).
fof(f1963_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,X2))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f1961,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1350,f137])).
fof(f1960,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X22) != s(bool,X24)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1350,f124])).
fof(f1958,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1350,f123])).
fof(f1956,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1350,f122])).
fof(f1955,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X7) != s(bool,X9)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1350,f120])).
fof(f1953,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1350,f118])).
fof(f1951,plain,(
  ( ! [X19,X17,X16] : (s(bool,X19) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X17,X16,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X17,X16,sK0)),s('type.list.list'('type.closSem.v'),sK12(X17,X16,sK0)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X17),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X19))) ) | $spl142),
  inference(cnf_transformation,[],[f1951_D])).
fof(f1951_D,plain,(
  ( ! [X19,X17,X16] : (s(bool,X19) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X17,X16,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X17,X16,sK0)),s('type.list.list'('type.closSem.v'),sK12(X17,X16,sK0)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X17),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X19))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f1947,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X57),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X59) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X57,X56,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X57,X56,sK0)),s('type.list.list'('type.closSem.v'),sK12(X57,X56,sK0))))) ) | $spl1),
  inference(resolution,[],[f1765,f137])).
fof(f1946,plain,(
  ( ! [X54,X52,X55,X53,X51] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X51),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X52),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X52,X51,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X52,X51,sK0)),s('type.list.list'('type.closSem.v'),sK12(X52,X51,sK0)))) | s(bool,X54) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f1765,f124])).
fof(f1943,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X36),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X37),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X37,X36,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X37,X36,sK0)),s('type.list.list'('type.closSem.v'),sK12(X37,X36,sK0))))) ) | $spl1),
  inference(resolution,[],[f1765,f123])).
fof(f1942,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X32,X31,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X32,X31,sK0)),s('type.list.list'('type.closSem.v'),sK12(X32,X31,sK0))))) ) | $spl1),
  inference(resolution,[],[f1765,f122])).
fof(f1941,plain,(
  ( ! [X30,X28,X26,X29,X27] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X26),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X27),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X27,X26,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X27,X26,sK0)),s('type.list.list'('type.closSem.v'),sK12(X27,X26,sK0)))) | s(bool,X29) != s(bool,X30)) ) | $spl1),
  inference(resolution,[],[f1765,f120])).
fof(f1939,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X17),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X17,X16,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X17,X16,sK0)),s('type.list.list'('type.closSem.v'),sK12(X17,X16,sK0))))) ) | $spl1),
  inference(resolution,[],[f1765,f118])).
fof(f1950,plain,(
  ( ! [X14,X13,X16] : (s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X13,X14,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X13,X14,sK0)),s('type.list.list'('type.closSem.v'),sK12(X13,X14,sK0))))) | p(s(bool,X16))) ) | $spl140),
  inference(cnf_transformation,[],[f1950_D])).
fof(f1950_D,plain,(
  ( ! [X14,X13,X16] : (s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X13,X14,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X13,X14,sK0)),s('type.list.list'('type.closSem.v'),sK12(X13,X14,sK0))))) | p(s(bool,X16))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f1935,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X53,X54,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X53,X54,sK0)),s('type.list.list'('type.closSem.v'),sK12(X53,X54,sK0))))) | p(s(bool,X55)) | p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X54),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1765,f137])).
fof(f1934,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X48,X49,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X48,X49,sK0)),s('type.list.list'('type.closSem.v'),sK12(X48,X49,sK0))))) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X49),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X48),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X51) != s(bool,X52)) ) | $spl1),
  inference(resolution,[],[f1765,f124])).
fof(f1931,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X33,X34,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X33,X34,sK0)),s('type.list.list'('type.closSem.v'),sK12(X33,X34,sK0))))) | p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X37) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X34),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X33),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1765,f123])).
fof(f1930,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X28,X29,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X28,X29,sK0)),s('type.list.list'('type.closSem.v'),sK12(X28,X29,sK0))))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X29),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1765,f122])).
fof(f1929,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X23,X24,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X23,X24,sK0)),s('type.list.list'('type.closSem.v'),sK12(X23,X24,sK0))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X24),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X26) != s(bool,X27)) ) | $spl1),
  inference(resolution,[],[f1765,f120])).
fof(f1927,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(X13,X14,sK0)))),s('type.list.list'('type.closSem.v'),sK11(X13,X14,sK0)),s('type.list.list'('type.closSem.v'),sK12(X13,X14,sK0))))) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1765,f118])).
fof(f1912,plain,(
  ( ! [X24,X23,X21,X19,X27,X25,X22,X20,X18] : (s(bool,X27) != s(bool,happ(s(fun(X18,bool),X19),s(X18,sK16(X19,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25)))),X18)))) | p(s(bool,happ(s(fun(X18,bool),X19),s(X18,X20)))) | ~p(s(bool,X27)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25))))))) ) | $spl138),
  inference(cnf_transformation,[],[f1912_D])).
fof(f1912_D,plain,(
  ( ! [X24,X23,X21,X19,X27,X25,X22,X20,X18] : (s(bool,X27) != s(bool,happ(s(fun(X18,bool),X19),s(X18,sK16(X19,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25)))),X18)))) | p(s(bool,happ(s(fun(X18,bool),X19),s(X18,X20)))) | ~p(s(bool,X27)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25))))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f1908,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X115,X113] : (p(s(bool,happ(s(fun(X106,bool),X107),s(X106,X108)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X109)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X110),s('type.closSem.state'(X109),X111))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X109)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X112),s('type.closSem.state'(X109),X113)))))) | p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X115) != s(bool,happ(s(fun(X106,bool),X107),s(X106,sK16(X107,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X109)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X110),s('type.closSem.state'(X109),X111))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X109)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X112),s('type.closSem.state'(X109),X113)))),X106))))) ) | $spl4),
  inference(resolution,[],[f1463,f137])).
fof(f1463,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK16(X14,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X8)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X9),s('type.closSem.state'(X8),X10))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X8)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X11),s('type.closSem.state'(X8),X12)))),X13))))) | p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X8)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X9),s('type.closSem.state'(X8),X10))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X8)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X11),s('type.closSem.state'(X8),X12))))))) ) | $spl4),
  inference(resolution,[],[f672,f116])).
fof(f672,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X0),X4)))))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X0),X4))))))) ) | $spl4),
  inference(resolution,[],[f145,f101])).
fof(f1907,plain,(
  ( ! [X103,X101,X99,X97,X105,X95,X102,X100,X98,X96,X104] : (p(s(bool,happ(s(fun(X95,bool),X96),s(X95,X97)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X98)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X99),s('type.closSem.state'(X98),X100))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X98)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X101),s('type.closSem.state'(X98),X102)))))) | ~p(s(bool,X103)) | ~p(s(bool,X104)) | s(bool,X104) != s(bool,happ(s(fun(X95,bool),X96),s(X95,sK16(X96,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X98)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X99),s('type.closSem.state'(X98),X100))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X98)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X101),s('type.closSem.state'(X98),X102)))),X95)))) | s(bool,X103) != s(bool,X105)) ) | $spl4),
  inference(resolution,[],[f1463,f124])).
fof(f1905,plain,(
  ( ! [X80,X78,X76,X74,X83,X81,X79,X77,X75,X73,X82] : (p(s(bool,happ(s(fun(X73,bool),X74),s(X73,X75)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X76)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X77),s('type.closSem.state'(X76),X78))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X76)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X79),s('type.closSem.state'(X76),X80)))))) | p(s(bool,X81)) | ~p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,happ(s(fun(X73,bool),X74),s(X73,sK16(X74,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X76)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X77),s('type.closSem.state'(X76),X78))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X76)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X79),s('type.closSem.state'(X76),X80)))),X73))))) ) | $spl4),
  inference(resolution,[],[f1463,f123])).
fof(f1903,plain,(
  ( ! [X61,X59,X57,X54,X52,X60,X58,X56,X55,X53,X51] : (p(s(bool,happ(s(fun(X51,bool),X52),s(X51,X53)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X54)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X55),s('type.closSem.state'(X54),X56))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X54)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X57),s('type.closSem.state'(X54),X58)))))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(X51,bool),X52),s(X51,sK16(X52,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X54)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X55),s('type.closSem.state'(X54),X56))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X54)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X57),s('type.closSem.state'(X54),X58)))),X51))))) ) | $spl4),
  inference(resolution,[],[f1463,f122])).
fof(f1902,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X40,X49] : (p(s(bool,happ(s(fun(X40,bool),X41),s(X40,X42)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X43)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X44),s('type.closSem.state'(X43),X45))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X43)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X46),s('type.closSem.state'(X43),X47)))))) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,happ(s(fun(X40,bool),X41),s(X40,sK16(X41,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X43)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X44),s('type.closSem.state'(X43),X45))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X43)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X46),s('type.closSem.state'(X43),X47)))),X40)))) | s(bool,X48) != s(bool,X50)) ) | $spl4),
  inference(resolution,[],[f1463,f120])).
fof(f1900,plain,(
  ( ! [X28,X26,X24,X23,X21,X19,X27,X25,X22,X20,X18] : (p(s(bool,happ(s(fun(X18,bool),X19),s(X18,X20)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25)))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,happ(s(fun(X18,bool),X19),s(X18,sK16(X19,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25)))),X18))))) ) | $spl4),
  inference(resolution,[],[f1463,f118])).
fof(f1895,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16))))) ) | $spl136),
  inference(cnf_transformation,[],[f1895_D])).
fof(f1895_D,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f1891,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closLang.exp'),X80)))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK16(X77,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X76))))) ) | $spl114),
  inference(resolution,[],[f1670,f137])).
fof(f1670,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X5),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X5),s('type.list.list'('type.closLang.exp'),X6))))) ) | $spl114),
  inference(resolution,[],[f1608,f116])).
fof(f1608,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X1))))) ) | $spl114),
  inference(equality_resolution,[],[f1603])).
fof(f1890,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closLang.exp'),X72)))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(X68,bool),X69),s(X68,sK16(X69,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X72),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X68)))) | s(bool,X73) != s(bool,X75)) ) | $spl114),
  inference(resolution,[],[f1670,f124])).
fof(f1888,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closLang.exp'),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(X52,bool),X53),s(X52,sK16(X53,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X52))))) ) | $spl114),
  inference(resolution,[],[f1670,f123])).
fof(f1886,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closLang.exp'),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X37),s(X36,sK16(X37,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X36))))) ) | $spl114),
  inference(resolution,[],[f1670,f122])).
fof(f1885,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closLang.exp'),X32)))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X28,bool),X29),s(X28,sK16(X29,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X28)))) | s(bool,X33) != s(bool,X35)) ) | $spl114),
  inference(resolution,[],[f1670,f120])).
fof(f1883,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12))))) ) | $spl114),
  inference(resolution,[],[f1670,f118])).
fof(f1878,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16))))) ) | $spl134),
  inference(cnf_transformation,[],[f1878_D])).
fof(f1878_D,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f1874,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closLang.exp'),X80)))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK15(X77,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X76))))) ) | $spl114),
  inference(resolution,[],[f1669,f137])).
fof(f1669,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK15(X3,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X2))))) | p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X1))))) ) | $spl114),
  inference(resolution,[],[f1608,f114])).
fof(f1873,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closLang.exp'),X72)))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(X68,bool),X69),s(X68,sK15(X69,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X72),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X68)))) | s(bool,X73) != s(bool,X75)) ) | $spl114),
  inference(resolution,[],[f1669,f124])).
fof(f1871,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closLang.exp'),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(X52,bool),X53),s(X52,sK15(X53,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X52))))) ) | $spl114),
  inference(resolution,[],[f1669,f123])).
fof(f1869,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closLang.exp'),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X37),s(X36,sK15(X37,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X36))))) ) | $spl114),
  inference(resolution,[],[f1669,f122])).
fof(f1868,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closLang.exp'),X32)))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X28,bool),X29),s(X28,sK15(X29,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X28)))) | s(bool,X33) != s(bool,X35)) ) | $spl114),
  inference(resolution,[],[f1669,f120])).
fof(f1866,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12))))) ) | $spl114),
  inference(resolution,[],[f1669,f118])).
fof(f1861,plain,(
  ( ! [X24,X23,X21,X19,X27,X25,X22,X20,X18] : (s(bool,X27) != s(bool,happ(s(fun(X18,bool),X19),s(X18,sK15(X19,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25)))),X18)))) | p(s(bool,happ(s(fun(X18,bool),X19),s(X18,X20)))) | ~p(s(bool,X27)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25))))))) ) | $spl132),
  inference(cnf_transformation,[],[f1861_D])).
fof(f1861_D,plain,(
  ( ! [X24,X23,X21,X19,X27,X25,X22,X20,X18] : (s(bool,X27) != s(bool,happ(s(fun(X18,bool),X19),s(X18,sK15(X19,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25)))),X18)))) | p(s(bool,happ(s(fun(X18,bool),X19),s(X18,X20)))) | ~p(s(bool,X27)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25))))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f1857,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X115,X113] : (p(s(bool,happ(s(fun(X106,bool),X107),s(X106,X108)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X109)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X110),s('type.closSem.state'(X109),X111))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X109)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X112),s('type.closSem.state'(X109),X113)))))) | p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X115) != s(bool,happ(s(fun(X106,bool),X107),s(X106,sK15(X107,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X109)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X110),s('type.closSem.state'(X109),X111))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X109)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X112),s('type.closSem.state'(X109),X113)))),X106))))) ) | $spl4),
  inference(resolution,[],[f1462,f137])).
fof(f1462,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK15(X6,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X0),X4)))),X5))))) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X1),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X3),s('type.closSem.state'(X0),X4))))))) ) | $spl4),
  inference(resolution,[],[f672,f114])).
fof(f1856,plain,(
  ( ! [X103,X101,X99,X97,X105,X95,X102,X100,X98,X96,X104] : (p(s(bool,happ(s(fun(X95,bool),X96),s(X95,X97)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X98)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X99),s('type.closSem.state'(X98),X100))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X98)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X101),s('type.closSem.state'(X98),X102)))))) | ~p(s(bool,X103)) | ~p(s(bool,X104)) | s(bool,X104) != s(bool,happ(s(fun(X95,bool),X96),s(X95,sK15(X96,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X98)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X99),s('type.closSem.state'(X98),X100))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X98)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X101),s('type.closSem.state'(X98),X102)))),X95)))) | s(bool,X103) != s(bool,X105)) ) | $spl4),
  inference(resolution,[],[f1462,f124])).
fof(f1854,plain,(
  ( ! [X80,X78,X76,X74,X83,X81,X79,X77,X75,X73,X82] : (p(s(bool,happ(s(fun(X73,bool),X74),s(X73,X75)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X76)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X77),s('type.closSem.state'(X76),X78))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X76)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X79),s('type.closSem.state'(X76),X80)))))) | p(s(bool,X81)) | ~p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,happ(s(fun(X73,bool),X74),s(X73,sK15(X74,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X76)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X77),s('type.closSem.state'(X76),X78))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X76)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X79),s('type.closSem.state'(X76),X80)))),X73))))) ) | $spl4),
  inference(resolution,[],[f1462,f123])).
fof(f1852,plain,(
  ( ! [X61,X59,X57,X54,X52,X60,X58,X56,X55,X53,X51] : (p(s(bool,happ(s(fun(X51,bool),X52),s(X51,X53)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X54)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X55),s('type.closSem.state'(X54),X56))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X54)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X57),s('type.closSem.state'(X54),X58)))))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(X51,bool),X52),s(X51,sK15(X52,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X54)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X55),s('type.closSem.state'(X54),X56))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X54)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X57),s('type.closSem.state'(X54),X58)))),X51))))) ) | $spl4),
  inference(resolution,[],[f1462,f122])).
fof(f1851,plain,(
  ( ! [X47,X45,X43,X41,X50,X48,X46,X44,X42,X40,X49] : (p(s(bool,happ(s(fun(X40,bool),X41),s(X40,X42)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X43)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X44),s('type.closSem.state'(X43),X45))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X43)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X46),s('type.closSem.state'(X43),X47)))))) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,happ(s(fun(X40,bool),X41),s(X40,sK15(X41,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X43)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X44),s('type.closSem.state'(X43),X45))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X43)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X46),s('type.closSem.state'(X43),X47)))),X40)))) | s(bool,X48) != s(bool,X50)) ) | $spl4),
  inference(resolution,[],[f1462,f120])).
fof(f1849,plain,(
  ( ! [X28,X26,X24,X23,X21,X19,X27,X25,X22,X20,X18] : (p(s(bool,happ(s(fun(X18,bool),X19),s(X18,X20)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X21)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25)))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,happ(s(fun(X18,bool),X19),s(X18,sK15(X19,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X22),s('type.closSem.state'(X21),X23))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X24),s('type.closSem.state'(X21),X25)))),X18))))) ) | $spl4),
  inference(resolution,[],[f1462,f118])).
fof(f1832,plain,(
  ( ! [X30,X28,X27] : (s(bool,X30) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X28,X27,sK0)),s('type.closSem.state'(sK0),sK13(X28,X27,sK0)),s('type.closSem.state'(sK0),sK14(X28,X27,sK0)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X27),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X30))) ) | $spl130),
  inference(cnf_transformation,[],[f1832_D])).
fof(f1832_D,plain,(
  ( ! [X30,X28,X27] : (s(bool,X30) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X28,X27,sK0)),s('type.closSem.state'(sK0),sK13(X28,X27,sK0)),s('type.closSem.state'(sK0),sK14(X28,X27,sK0)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X27),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X30))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f1828,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X67),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X68),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X70) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X68,X67,sK0)),s('type.closSem.state'(sK0),sK13(X68,X67,sK0)),s('type.closSem.state'(sK0),sK14(X68,X67,sK0))))) ) | $spl1),
  inference(resolution,[],[f1764,f137])).
fof(f1827,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X62),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X63),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X64)) | ~p(s(bool,X65)) | s(bool,X64) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X63,X62,sK0)),s('type.closSem.state'(sK0),sK13(X63,X62,sK0)),s('type.closSem.state'(sK0),sK14(X63,X62,sK0)))) | s(bool,X65) != s(bool,X66)) ) | $spl1),
  inference(resolution,[],[f1764,f124])).
fof(f1824,plain,(
  ( ! [X47,X50,X48,X51,X49] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X47),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X48),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X48,X47,sK0)),s('type.closSem.state'(sK0),sK13(X48,X47,sK0)),s('type.closSem.state'(sK0),sK14(X48,X47,sK0))))) ) | $spl1),
  inference(resolution,[],[f1764,f123])).
fof(f1823,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X42),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X43),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X43,X42,sK0)),s('type.closSem.state'(sK0),sK13(X43,X42,sK0)),s('type.closSem.state'(sK0),sK14(X43,X42,sK0))))) ) | $spl1),
  inference(resolution,[],[f1764,f122])).
fof(f1822,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X37),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X38),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X38,X37,sK0)),s('type.closSem.state'(sK0),sK13(X38,X37,sK0)),s('type.closSem.state'(sK0),sK14(X38,X37,sK0)))) | s(bool,X40) != s(bool,X41)) ) | $spl1),
  inference(resolution,[],[f1764,f120])).
fof(f1820,plain,(
  ( ! [X30,X28,X31,X29,X27] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X27),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X28,X27,sK0)),s('type.closSem.state'(sK0),sK13(X28,X27,sK0)),s('type.closSem.state'(sK0),sK14(X28,X27,sK0))))) ) | $spl1),
  inference(resolution,[],[f1764,f118])).
fof(f1831,plain,(
  ( ! [X14,X13,X16] : (s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X13,X14,sK0)),s('type.closSem.state'(sK0),sK13(X13,X14,sK0)),s('type.closSem.state'(sK0),sK14(X13,X14,sK0))))) | p(s(bool,X16))) ) | $spl128),
  inference(cnf_transformation,[],[f1831_D])).
fof(f1831_D,plain,(
  ( ! [X14,X13,X16] : (s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X13,X14,sK0)),s('type.closSem.state'(sK0),sK13(X13,X14,sK0)),s('type.closSem.state'(sK0),sK14(X13,X14,sK0))))) | p(s(bool,X16))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1814,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X53,X54,sK0)),s('type.closSem.state'(sK0),sK13(X53,X54,sK0)),s('type.closSem.state'(sK0),sK14(X53,X54,sK0))))) | p(s(bool,X55)) | p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X54),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1764,f137])).
fof(f1813,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X48,X49,sK0)),s('type.closSem.state'(sK0),sK13(X48,X49,sK0)),s('type.closSem.state'(sK0),sK14(X48,X49,sK0))))) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X49),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X48),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X51) != s(bool,X52)) ) | $spl1),
  inference(resolution,[],[f1764,f124])).
fof(f1810,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X33,X34,sK0)),s('type.closSem.state'(sK0),sK13(X33,X34,sK0)),s('type.closSem.state'(sK0),sK14(X33,X34,sK0))))) | p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X37) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X34),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X33),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1764,f123])).
fof(f1809,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X28,X29,sK0)),s('type.closSem.state'(sK0),sK13(X28,X29,sK0)),s('type.closSem.state'(sK0),sK14(X28,X29,sK0))))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X29),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1764,f122])).
fof(f1808,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X23,X24,sK0)),s('type.closSem.state'(sK0),sK13(X23,X24,sK0)),s('type.closSem.state'(sK0),sK14(X23,X24,sK0))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X24),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X26) != s(bool,X27)) ) | $spl1),
  inference(resolution,[],[f1764,f120])).
fof(f1806,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X13,X14,sK0)),s('type.closSem.state'(sK0),sK13(X13,X14,sK0)),s('type.closSem.state'(sK0),sK14(X13,X14,sK0))))) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1764,f118])).
fof(f1800,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X34,X32] : (s(bool,X34) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X23),s('type.closSem.state'(sK0),X31),s('type.closSem.state'(sK0),X32))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X23))),s('type.list.list'('type.closSem.v'),X24),s('type.list.list'('type.closSem.v'),X25)))) | ~p(s(bool,X34)) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X23),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X30)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X28),s('type.closSem.state'(X30),X26))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X30)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X29),s('type.closSem.state'(X30),X27)))))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X23,X26,X27,X28,X29,X30)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X24))),s('type.closSem.state'(sK0),X31))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X23,X26,X27,X28,X29,X30)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X25))),s('type.closSem.state'(sK0),X32))))))) ) | $spl126),
  inference(cnf_transformation,[],[f1800_D])).
fof(f1800_D,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X34,X32] : (s(bool,X34) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X23),s('type.closSem.state'(sK0),X31),s('type.closSem.state'(sK0),X32))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X23))),s('type.list.list'('type.closSem.v'),X24),s('type.list.list'('type.closSem.v'),X25)))) | ~p(s(bool,X34)) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X23),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X30)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X28),s('type.closSem.state'(X30),X26))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X30)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X29),s('type.closSem.state'(X30),X27)))))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X23,X26,X27,X28,X29,X30)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X24))),s('type.closSem.state'(sK0),X31))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X23,X26,X27,X28,X29,X30)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X25))),s('type.closSem.state'(sK0),X32))))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1798,plain,(
  ( ! [X127,X134,X132,X130,X128,X138,X136,X135,X133,X131,X129,X137] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X127))),s('type.list.list'('type.closSem.v'),X128),s('type.list.list'('type.closSem.v'),X129)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X127,X130,X131,X132,X133,X134)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X128))),s('type.closSem.state'(sK0),X135))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X127,X130,X131,X132,X133,X134)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X129))),s('type.closSem.state'(sK0),X136)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X127),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X134)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X132),s('type.closSem.state'(X134),X130))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X134)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X133),s('type.closSem.state'(X134),X131)))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X138) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X127),s('type.closSem.state'(sK0),X135),s('type.closSem.state'(sK0),X136)))) ) | $spl10),
  inference(resolution,[],[f807,f137])).
fof(f1797,plain,(
  ( ! [X125,X123,X121,X118,X116,X114,X126,X124,X122,X120,X119,X117,X115] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X114))),s('type.list.list'('type.closSem.v'),X115),s('type.list.list'('type.closSem.v'),X116)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X114,X117,X118,X119,X120,X121)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X115))),s('type.closSem.state'(sK0),X122))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X114,X117,X118,X119,X120,X121)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X116))),s('type.closSem.state'(sK0),X123)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X114),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X121)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X119),s('type.closSem.state'(X121),X117))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X121)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X120),s('type.closSem.state'(X121),X118)))))) | ~p(s(bool,X124)) | ~p(s(bool,X125)) | s(bool,X125) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X114),s('type.closSem.state'(sK0),X122),s('type.closSem.state'(sK0),X123))) | s(bool,X124) != s(bool,X126)) ) | $spl10),
  inference(resolution,[],[f807,f124])).
fof(f1795,plain,(
  ( ! [X94,X92,X90,X88,X99,X97,X95,X93,X91,X89,X100,X98,X96] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X88))),s('type.list.list'('type.closSem.v'),X89),s('type.list.list'('type.closSem.v'),X90)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X88,X91,X92,X93,X94,X95)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X89))),s('type.closSem.state'(sK0),X96))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X88,X91,X92,X93,X94,X95)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X90))),s('type.closSem.state'(sK0),X97)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X88),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X95)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X93),s('type.closSem.state'(X95),X91))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X95)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X94),s('type.closSem.state'(X95),X92)))))) | p(s(bool,X98)) | ~p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X99) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X88),s('type.closSem.state'(sK0),X96),s('type.closSem.state'(sK0),X97)))) ) | $spl10),
  inference(resolution,[],[f807,f123])).
fof(f1793,plain,(
  ( ! [X70,X68,X66,X64,X74,X72,X62,X71,X69,X67,X65,X73,X63] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X62))),s('type.list.list'('type.closSem.v'),X63),s('type.list.list'('type.closSem.v'),X64)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X62,X65,X66,X67,X68,X69)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X63))),s('type.closSem.state'(sK0),X70))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X62,X65,X66,X67,X68,X69)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X64))),s('type.closSem.state'(sK0),X71)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X62),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X69)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X67),s('type.closSem.state'(X69),X65))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X69)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X68),s('type.closSem.state'(X69),X66)))))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X72) != s(bool,X74) | s(bool,X73) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X62),s('type.closSem.state'(sK0),X70),s('type.closSem.state'(sK0),X71)))) ) | $spl10),
  inference(resolution,[],[f807,f122])).
fof(f1792,plain,(
  ( ! [X61,X59,X57,X54,X52,X50,X60,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X49))),s('type.list.list'('type.closSem.v'),X50),s('type.list.list'('type.closSem.v'),X51)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X49,X52,X53,X54,X55,X56)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X50))),s('type.closSem.state'(sK0),X57))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X49,X52,X53,X54,X55,X56)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X51))),s('type.closSem.state'(sK0),X58)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X49),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X56)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X54),s('type.closSem.state'(X56),X52))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X56)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X55),s('type.closSem.state'(X56),X53)))))) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X49),s('type.closSem.state'(sK0),X57),s('type.closSem.state'(sK0),X58))) | s(bool,X59) != s(bool,X61)) ) | $spl10),
  inference(resolution,[],[f807,f120])).
fof(f1790,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X34,X32] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X23))),s('type.list.list'('type.closSem.v'),X24),s('type.list.list'('type.closSem.v'),X25)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X23,X26,X27,X28,X29,X30)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X24))),s('type.closSem.state'(sK0),X31))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK9(X23,X26,X27,X28,X29,X30)),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X25))),s('type.closSem.state'(sK0),X32)))))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X23),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X30)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X28),s('type.closSem.state'(X30),X26))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X30)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X29),s('type.closSem.state'(X30),X27)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X23),s('type.closSem.state'(sK0),X31),s('type.closSem.state'(sK0),X32)))) ) | $spl10),
  inference(resolution,[],[f807,f118])).
fof(f1779,plain,(
  ( ! [X6,X7,X9] : (s(bool,X9) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closLang.exp'),X7))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X9))) ) | $spl124),
  inference(cnf_transformation,[],[f1779_D])).
fof(f1779_D,plain,(
  ( ! [X6,X7,X9] : (s(bool,X9) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closLang.exp'),X7))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X9))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1775,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X46),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X47),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X46),s('type.list.list'('type.closLang.exp'),X47)))) ) | $spl1),
  inference(resolution,[],[f1762,f137])).
fof(f1774,plain,(
  ( ! [X45,X43,X41,X44,X42] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X41),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X42),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X41),s('type.list.list'('type.closLang.exp'),X42))) | s(bool,X43) != s(bool,X45)) ) | $spl1),
  inference(resolution,[],[f1762,f124])).
fof(f1772,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closLang.exp'),X32)))) ) | $spl1),
  inference(resolution,[],[f1762,f123])).
fof(f1770,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X21),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X22),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X21),s('type.list.list'('type.closLang.exp'),X22)))) ) | $spl1),
  inference(resolution,[],[f1762,f122])).
fof(f1769,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X17),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closLang.exp'),X17))) | s(bool,X18) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f1762,f120])).
fof(f1767,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closLang.exp'),X7)))) ) | $spl1),
  inference(resolution,[],[f1762,f118])).
fof(f1763,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,X9) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X5),s('type.closSem.state'(sK0),X7))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X5))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X7)))))) | ~p(s(bool,X9)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closLang.exp'),X6))))) ) | $spl122),
  inference(cnf_transformation,[],[f1763_D])).
fof(f1763_D,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,X9) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X5),s('type.closSem.state'(sK0),X7))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X5))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X7)))))) | ~p(s(bool,X9)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closLang.exp'),X6))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1760,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X60),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X61))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X62),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X63)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X60),s('type.list.list'('type.closLang.exp'),X62)))) | p(s(bool,X64)) | ~p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X61),s('type.closSem.state'(sK0),X63)))) ) | $spl1),
  inference(resolution,[],[f1309,f137])).
fof(f1759,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X54))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X56)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closLang.exp'),X55)))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X54),s('type.closSem.state'(sK0),X56))) | s(bool,X57) != s(bool,X59)) ) | $spl1),
  inference(resolution,[],[f1309,f124])).
fof(f1757,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X40))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X41),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X42)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closLang.exp'),X41)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X40),s('type.closSem.state'(sK0),X42)))) ) | $spl1),
  inference(resolution,[],[f1309,f123])).
fof(f1755,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X26))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X27),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X28)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closLang.exp'),X27)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X26),s('type.closSem.state'(sK0),X28)))) ) | $spl1),
  inference(resolution,[],[f1309,f122])).
fof(f1754,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X18),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X19))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X21)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X18),s('type.list.list'('type.closLang.exp'),X20)))) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X19),s('type.closSem.state'(sK0),X21))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f1309,f120])).
fof(f1752,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X5))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X6),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),X7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closLang.exp'),X6)))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),X5),s('type.closSem.state'(sK0),X7)))) ) | $spl1),
  inference(resolution,[],[f1309,f118])).
fof(f1740,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16))))) ) | $spl120),
  inference(cnf_transformation,[],[f1740_D])).
fof(f1740_D,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1736,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closLang.exp'),X80)))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK16(X77,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X76))))) ) | $spl106),
  inference(resolution,[],[f1654,f137])).
fof(f1654,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X10,bool),X11),s(X10,sK16(X11,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X8),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X9),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X10))))) | p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X8),s('type.list.list'('type.closLang.exp'),X9))))) ) | $spl106),
  inference(resolution,[],[f1553,f116])).
fof(f1735,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closLang.exp'),X72)))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(X68,bool),X69),s(X68,sK16(X69,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X72),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X68)))) | s(bool,X73) != s(bool,X75)) ) | $spl106),
  inference(resolution,[],[f1654,f124])).
fof(f1733,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closLang.exp'),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(X52,bool),X53),s(X52,sK16(X53,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X52))))) ) | $spl106),
  inference(resolution,[],[f1654,f123])).
fof(f1731,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closLang.exp'),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X37),s(X36,sK16(X37,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X36))))) ) | $spl106),
  inference(resolution,[],[f1654,f122])).
fof(f1730,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closLang.exp'),X32)))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X28,bool),X29),s(X28,sK16(X29,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X28)))) | s(bool,X33) != s(bool,X35)) ) | $spl106),
  inference(resolution,[],[f1654,f120])).
fof(f1728,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12))))) ) | $spl106),
  inference(resolution,[],[f1654,f118])).
fof(f1723,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16))))) ) | $spl118),
  inference(cnf_transformation,[],[f1723_D])).
fof(f1723_D,plain,(
  ( ! [X14,X12,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12)))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,X18)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f1719,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closLang.exp'),X80)))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X76,bool),X77),s(X76,sK15(X77,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X79),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X80),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X76))))) ) | $spl106),
  inference(resolution,[],[f1653,f137])).
fof(f1653,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK15(X6,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X3),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X5))))) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X3),s('type.list.list'('type.closLang.exp'),X4))))) ) | $spl106),
  inference(resolution,[],[f1553,f114])).
fof(f1718,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closLang.exp'),X72)))) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun(X68,bool),X69),s(X68,sK15(X69,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X71),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X72),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X68)))) | s(bool,X73) != s(bool,X75)) ) | $spl106),
  inference(resolution,[],[f1653,f124])).
fof(f1716,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closLang.exp'),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,happ(s(fun(X52,bool),X53),s(X52,sK15(X53,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X56),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X52))))) ) | $spl106),
  inference(resolution,[],[f1653,f123])).
fof(f1714,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closLang.exp'),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,happ(s(fun(X36,bool),X37),s(X36,sK15(X37,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X40),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X36))))) ) | $spl106),
  inference(resolution,[],[f1653,f122])).
fof(f1713,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closLang.exp'),X32)))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X28,bool),X29),s(X28,sK15(X29,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X32),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X28)))) | s(bool,X33) != s(bool,X35)) ) | $spl106),
  inference(resolution,[],[f1653,f120])).
fof(f1711,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closLang.exp'),X16)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X15),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))),X12))))) ) | $spl106),
  inference(resolution,[],[f1653,f118])).
fof(f1576,plain,(
  ( ! [X12,X15,X13,X11] : (s(bool,X15) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',sK3))) | ~p(s(bool,X15)) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X11),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X12),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X11),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X12),s('type.list.list'('type.closLang.exp'),X13))))) ) | $spl110),
  inference(cnf_transformation,[],[f1576_D])).
fof(f1576_D,plain,(
  ( ! [X12,X15,X13,X11] : (s(bool,X15) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',sK3))) | ~p(s(bool,X15)) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X11),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X12),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X11),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X12),s('type.list.list'('type.closLang.exp'),X13))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1695,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X59),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X60),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X59),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X61),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X60),s('type.list.list'('type.closLang.exp'),X61)))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X59),s('type.num.num',sK3)))) ) | $spl106),
  inference(resolution,[],[f1652,f137])).
fof(f1694,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X53),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X54),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X53),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X54),s('type.list.list'('type.closLang.exp'),X55)))) | ~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X57) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num',sK3))) | s(bool,X56) != s(bool,X58)) ) | $spl106),
  inference(resolution,[],[f1652,f124])).
fof(f1692,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X41),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X42),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X41),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X43),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X42),s('type.list.list'('type.closLang.exp'),X43)))) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X41),s('type.num.num',sK3)))) ) | $spl106),
  inference(resolution,[],[f1652,f123])).
fof(f1690,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X29),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X30),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X29),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X30),s('type.list.list'('type.closLang.exp'),X31)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X29),s('type.num.num',sK3)))) ) | $spl106),
  inference(resolution,[],[f1652,f122])).
fof(f1689,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X23),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X24),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X23),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X24),s('type.list.list'('type.closLang.exp'),X25)))) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num',sK3))) | s(bool,X26) != s(bool,X28)) ) | $spl106),
  inference(resolution,[],[f1652,f120])).
fof(f1687,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X11),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X12),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X11),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X12),s('type.list.list'('type.closLang.exp'),X13)))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',sK3)))) ) | $spl106),
  inference(resolution,[],[f1652,f118])).
fof(f364,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | ~p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | ~p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1636,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f137])).
fof(f1635,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f146,f124])).
fof(f1633,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f123])).
fof(f1631,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f122])).
fof(f1630,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f146,f120])).
fof(f1628,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f118])).
fof(f1627,plain,(
  ( ! [X43,X41,X40] : (s(bool,X43) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X40),s('type.list.list'('type.closSem.v'),X41))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X40))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X41))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X43))) ) | $spl116),
  inference(cnf_transformation,[],[f1627_D])).
fof(f1627_D,plain,(
  ( ! [X43,X41,X40] : (s(bool,X43) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X40),s('type.list.list'('type.closSem.v'),X41))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X40))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X41))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X43))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f1612,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X0))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X1))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X0),s('type.list.list'('type.closSem.v'),X1)))) ) | ($spl4 | $spl8 | $spl10)),
  inference(resolution,[],[f1577,f118])).
fof(f1577,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X0),s('type.list.list'('type.closSem.v'),X1)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X0))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X1))),s('type.closSem.state'(sK0),sK7))))))) ) | ($spl4 | $spl8 | $spl10)),
  inference(resolution,[],[f725,f149])).
fof(f725,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X10),s('type.closSem.state'(sK0),X11)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X12),s('type.list.list'('type.closSem.v'),X13)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X12))),s('type.closSem.state'(sK0),X10))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X13))),s('type.closSem.state'(sK0),X11))))))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f152,f145])).
fof(f1614,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X11))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11))) | s(bool,X12) != s(bool,X14)) ) | ($spl4 | $spl8 | $spl10)),
  inference(resolution,[],[f1577,f120])).
fof(f1615,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X15))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X16))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X15),s('type.list.list'('type.closSem.v'),X16)))) ) | ($spl4 | $spl8 | $spl10)),
  inference(resolution,[],[f1577,f122])).
fof(f1617,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X25))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X26))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X25),s('type.list.list'('type.closSem.v'),X26)))) ) | ($spl4 | $spl8 | $spl10)),
  inference(resolution,[],[f1577,f123])).
fof(f1619,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X35))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X36))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X35),s('type.list.list'('type.closSem.v'),X36))) | s(bool,X37) != s(bool,X39)) ) | ($spl4 | $spl8 | $spl10)),
  inference(resolution,[],[f1577,f124])).
fof(f1620,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X40))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X41))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X43) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X40),s('type.list.list'('type.closSem.v'),X41)))) ) | ($spl4 | $spl8 | $spl10)),
  inference(resolution,[],[f1577,f137])).
fof(f1624,plain,(
  $false | ($spl3 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f1611,f143])).
fof(f1611,plain,(
  p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ($spl4 | $spl6 | $spl8 | $spl10)),
  inference(resolution,[],[f1577,f147])).
fof(f1603,plain,(
  ( ! [X10,X13,X11] : (s(bool,X13) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closLang.exp'),X11)))) | p(s(bool,X13))) ) | $spl114),
  inference(cnf_transformation,[],[f1603_D])).
fof(f1603_D,plain,(
  ( ! [X10,X13,X11] : (s(bool,X13) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closLang.exp'),X11)))) | p(s(bool,X13))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f1600,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X50),s('type.list.list'('type.closLang.exp'),X51)))) | p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X50),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X51),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl4 | $spl6 | $spl8)),
  inference(resolution,[],[f1565,f137])).
fof(f1565,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X9),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X9),s('type.list.list'('type.closLang.exp'),X10))))) ) | ($spl4 | $spl6 | $spl8)),
  inference(resolution,[],[f1518,f145])).
fof(f1518,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK3)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X2),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X2),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X1))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1505,f101])).
fof(f1505,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X1))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1262,f149])).
fof(f1262,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X1),s('type.closSem.state'(sK0),X3)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X1))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X3)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X2))))) ) | $spl6),
  inference(resolution,[],[f147,f107])).
fof(f1599,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X45),s('type.list.list'('type.closLang.exp'),X46)))) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X45),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X46),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X48) != s(bool,X49)) ) | ($spl4 | $spl6 | $spl8)),
  inference(resolution,[],[f1565,f124])).
fof(f1596,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X30),s('type.list.list'('type.closLang.exp'),X31)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X30),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl4 | $spl6 | $spl8)),
  inference(resolution,[],[f1565,f123])).
fof(f1595,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closLang.exp'),X26)))) | p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X26),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl4 | $spl6 | $spl8)),
  inference(resolution,[],[f1565,f122])).
fof(f1594,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closLang.exp'),X21)))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X20),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X21),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X23) != s(bool,X24)) ) | ($spl4 | $spl6 | $spl8)),
  inference(resolution,[],[f1565,f120])).
fof(f1592,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closLang.exp'),X11)))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl4 | $spl6 | $spl8)),
  inference(resolution,[],[f1565,f118])).
fof(f1588,plain,(
  ( ! [X4,X2,X7,X5,X3] : (s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X4),s('type.closSem.state'(sK0),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3)))) | ~p(s(bool,X7)) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(sK0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),X5))))))) ) | $spl112),
  inference(cnf_transformation,[],[f1588_D])).
fof(f1588_D,plain,(
  ( ! [X4,X2,X7,X5,X3] : (s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X4),s('type.closSem.state'(sK0),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3)))) | ~p(s(bool,X7)) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(sK0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),X5))))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1586,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X58),s('type.list.list'('type.closSem.v'),X59)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X58))),s('type.closSem.state'(sK0),X60))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X59))),s('type.closSem.state'(sK0),X61)))))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X60),s('type.closSem.state'(sK0),X61)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f725,f137])).
fof(f1585,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X51),s('type.list.list'('type.closSem.v'),X52)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X51))),s('type.closSem.state'(sK0),X53))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X52))),s('type.closSem.state'(sK0),X54)))))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X53),s('type.closSem.state'(sK0),X54))) | s(bool,X55) != s(bool,X57)) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f725,f124])).
fof(f1583,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X37),s('type.list.list'('type.closSem.v'),X38)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X37))),s('type.closSem.state'(sK0),X39))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X38))),s('type.closSem.state'(sK0),X40)))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X39),s('type.closSem.state'(sK0),X40)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f725,f123])).
fof(f1581,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X23),s('type.list.list'('type.closSem.v'),X24)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X23))),s('type.closSem.state'(sK0),X25))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X24))),s('type.closSem.state'(sK0),X26)))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X25),s('type.closSem.state'(sK0),X26)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f725,f122])).
fof(f1580,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X16),s('type.list.list'('type.closSem.v'),X17)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X16))),s('type.closSem.state'(sK0),X18))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X17))),s('type.closSem.state'(sK0),X19)))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X18),s('type.closSem.state'(sK0),X19))) | s(bool,X20) != s(bool,X22)) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f725,f120])).
fof(f1578,plain,(
  ( ! [X6,X4,X2,X8,X7,X5,X3] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(sK0),X4))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),X5)))))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X4),s('type.closSem.state'(sK0),X5)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f725,f118])).
fof(f1574,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X59),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X60),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X59),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X61),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X60),s('type.list.list'('type.closLang.exp'),X61)))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X59),s('type.num.num',sK3)))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1518,f137])).
fof(f1573,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X53),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X54),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X53),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X55),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X54),s('type.list.list'('type.closLang.exp'),X55)))) | ~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X57) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num',sK3))) | s(bool,X56) != s(bool,X58)) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1518,f124])).
fof(f1571,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X41),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X42),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X41),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X43),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X42),s('type.list.list'('type.closLang.exp'),X43)))) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X41),s('type.num.num',sK3)))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1518,f123])).
fof(f1569,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X29),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X30),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X29),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X31),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X30),s('type.list.list'('type.closLang.exp'),X31)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X29),s('type.num.num',sK3)))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1518,f122])).
fof(f1568,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X23),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X24),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X23),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X24),s('type.list.list'('type.closLang.exp'),X25)))) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num',sK3))) | s(bool,X26) != s(bool,X28)) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1518,f120])).
fof(f1566,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X11),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X12),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X11),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X12),s('type.list.list'('type.closLang.exp'),X13)))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',sK3)))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1518,f118])).
fof(f1548,plain,(
  ( ! [X4,X2,X7,X5,X3] : (s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X5))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X7)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closLang.exp'),X4))))) ) | $spl108),
  inference(cnf_transformation,[],[f1548_D])).
fof(f1548_D,plain,(
  ( ! [X4,X2,X7,X5,X3] : (s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X5))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X7)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closLang.exp'),X4))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1544,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X58),s('type.list.list'('type.closSem.v'),X59))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X60),s('type.list.list'('type.closSem.v'),X61))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X58),s('type.list.list'('type.closLang.exp'),X60)))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X59),s('type.list.list'('type.closSem.v'),X61)))) ) | $spl8),
  inference(resolution,[],[f844,f137])).
fof(f844,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X1),s('type.list.list'('type.closSem.v'),X3)))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closSem.v'),X1))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X0),s('type.list.list'('type.closLang.exp'),X2))))) ) | $spl8),
  inference(resolution,[],[f149,f107])).
fof(f1543,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X51),s('type.list.list'('type.closSem.v'),X52))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closSem.v'),X54))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X51),s('type.list.list'('type.closLang.exp'),X53)))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X52),s('type.list.list'('type.closSem.v'),X54))) | s(bool,X55) != s(bool,X57)) ) | $spl8),
  inference(resolution,[],[f844,f124])).
fof(f1541,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X37),s('type.list.list'('type.closSem.v'),X38))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),X40))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X37),s('type.list.list'('type.closLang.exp'),X39)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X38),s('type.list.list'('type.closSem.v'),X40)))) ) | $spl8),
  inference(resolution,[],[f844,f123])).
fof(f1539,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closSem.v'),X24))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closSem.v'),X26))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closLang.exp'),X25)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X24),s('type.list.list'('type.closSem.v'),X26)))) ) | $spl8),
  inference(resolution,[],[f844,f122])).
fof(f1538,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),X17))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X18),s('type.list.list'('type.closSem.v'),X19))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closLang.exp'),X18)))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X17),s('type.list.list'('type.closSem.v'),X19))) | s(bool,X20) != s(bool,X22)) ) | $spl8),
  inference(resolution,[],[f844,f120])).
fof(f1536,plain,(
  ( ! [X6,X4,X2,X8,X7,X5,X3] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),X5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closLang.exp'),X4)))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X5)))) ) | $spl8),
  inference(resolution,[],[f844,f118])).
fof(f1532,plain,(
  ( ! [X14,X13,X16] : (s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closLang.exp'),X14)))) | p(s(bool,X16))) ) | $spl106),
  inference(cnf_transformation,[],[f1532_D])).
fof(f1532_D,plain,(
  ( ! [X14,X13,X16] : (s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closLang.exp'),X14)))) | p(s(bool,X16))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1529,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closLang.exp'),X54)))) | p(s(bool,X55)) | p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X54),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1505,f137])).
fof(f1528,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X48),s('type.list.list'('type.closLang.exp'),X49)))) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X48),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X49),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X51) != s(bool,X52)) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1505,f124])).
fof(f1525,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X33),s('type.list.list'('type.closLang.exp'),X34)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X37) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X33),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X34),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1505,f123])).
fof(f1524,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closLang.exp'),X29)))) | p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X28),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X29),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1505,f122])).
fof(f1523,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closLang.exp'),X24)))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X24),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X26) != s(bool,X27)) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1505,f120])).
fof(f1521,plain,(
  ( ! [X14,X17,X15,X13,X16] : (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closLang.exp'),X14)))) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X13),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X14),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f1505,f118])).
fof(f1516,plain,(
  ( ! [X4,X2,X7,X5,X3] : (s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X3),s('type.closSem.state'(sK0),X5))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X3))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X5)))))) | ~p(s(bool,X7)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closLang.exp'),X4))))) ) | $spl104),
  inference(cnf_transformation,[],[f1516_D])).
fof(f1516_D,plain,(
  ( ! [X4,X2,X7,X5,X3] : (s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X3),s('type.closSem.state'(sK0),X5))) | p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X3))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X5)))))) | ~p(s(bool,X7)) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closLang.exp'),X4))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1514,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X58),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X59))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X60),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X61)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X58),s('type.list.list'('type.closLang.exp'),X60)))) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X59),s('type.closSem.state'(sK0),X61)))) ) | $spl6),
  inference(resolution,[],[f1262,f137])).
fof(f1513,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X51),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X52))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X53),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X54)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X51),s('type.list.list'('type.closLang.exp'),X53)))) | ~p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X56) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X52),s('type.closSem.state'(sK0),X54))) | s(bool,X55) != s(bool,X57)) ) | $spl6),
  inference(resolution,[],[f1262,f124])).
fof(f1511,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X37),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X38))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X39),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X40)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X37),s('type.list.list'('type.closLang.exp'),X39)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X38),s('type.closSem.state'(sK0),X40)))) ) | $spl6),
  inference(resolution,[],[f1262,f123])).
fof(f1509,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X24))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X25),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X26)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X23),s('type.list.list'('type.closLang.exp'),X25)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X24),s('type.closSem.state'(sK0),X26)))) ) | $spl6),
  inference(resolution,[],[f1262,f122])).
fof(f1508,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X17))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X18),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X19)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X16),s('type.list.list'('type.closLang.exp'),X18)))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X17),s('type.closSem.state'(sK0),X19))) | s(bool,X20) != s(bool,X22)) ) | $spl6),
  inference(resolution,[],[f1262,f120])).
fof(f1506,plain,(
  ( ! [X6,X4,X2,X8,X7,X5,X3] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),X3))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X4),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),X5)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closLang.exp'),X4)))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),X3),s('type.closSem.state'(sK0),X5)))) ) | $spl6),
  inference(resolution,[],[f1262,f118])).
fof(f1504,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl102),
  inference(cnf_transformation,[],[f1504_D])).
fof(f1504_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1502,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X56))))) ) | $spl1),
  inference(resolution,[],[f1428,f137])).
fof(f1428,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X12))))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))))) ) | $spl1),
  inference(resolution,[],[f1413,f116])).
fof(f1413,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))))) | $spl1),
  inference(resolution,[],[f1251,f102])).
fof(f1501,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f1428,f124])).
fof(f1499,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X38))))) ) | $spl1),
  inference(resolution,[],[f1428,f123])).
fof(f1497,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X26))))) ) | $spl1),
  inference(resolution,[],[f1428,f122])).
fof(f1496,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f1428,f120])).
fof(f1494,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X8))))) ) | $spl1),
  inference(resolution,[],[f1428,f118])).
fof(f1491,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl100),
  inference(cnf_transformation,[],[f1491_D])).
fof(f1491_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1489,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X56))))) ) | $spl1),
  inference(resolution,[],[f1427,f137])).
fof(f1427,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK15(X10,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X9))))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))))) ) | $spl1),
  inference(resolution,[],[f1413,f114])).
fof(f1488,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f1427,f124])).
fof(f1486,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X38))))) ) | $spl1),
  inference(resolution,[],[f1427,f123])).
fof(f1484,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X26))))) ) | $spl1),
  inference(resolution,[],[f1427,f122])).
fof(f1483,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f1427,f120])).
fof(f1481,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))),X8))))) ) | $spl1),
  inference(resolution,[],[f1427,f118])).
fof(f1475,plain,(
  ( ! [X19,X17,X22,X20,X18,X16] : (s(bool,X22) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X17),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X19),s('type.closSem.state'(X16),X20))))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X16)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X17),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X16)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X19),s('type.closSem.state'(X16),X20)))))) | p(s(bool,X22))) ) | $spl98),
  inference(cnf_transformation,[],[f1475_D])).
fof(f1475_D,plain,(
  ( ! [X19,X17,X22,X20,X18,X16] : (s(bool,X22) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X17),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X19),s('type.closSem.state'(X16),X20))))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X16)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X17),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X16)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X19),s('type.closSem.state'(X16),X20)))))) | p(s(bool,X22))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1472,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X80)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X81),s('type.closSem.state'(X80),X82))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X80)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X83),s('type.closSem.state'(X80),X84)))))) | p(s(bool,X85)) | p(s(bool,X86)) | s(bool,X86) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X80)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X81),s('type.closSem.state'(X80),X82))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X80)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X83),s('type.closSem.state'(X80),X84)))))) ) | $spl4),
  inference(resolution,[],[f672,f137])).
fof(f1471,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X72)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X73),s('type.closSem.state'(X72),X74))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X72)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X75),s('type.closSem.state'(X72),X76)))))) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X77) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X72)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X73),s('type.closSem.state'(X72),X74))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X72)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X75),s('type.closSem.state'(X72),X76))))) | s(bool,X78) != s(bool,X79)) ) | $spl4),
  inference(resolution,[],[f672,f124])).
fof(f1468,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X48)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X49),s('type.closSem.state'(X48),X50))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X48)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X51),s('type.closSem.state'(X48),X52)))))) | p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X55) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X48)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X49),s('type.closSem.state'(X48),X50))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X48)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X51),s('type.closSem.state'(X48),X52)))))) ) | $spl4),
  inference(resolution,[],[f672,f123])).
fof(f1467,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X40)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X41),s('type.closSem.state'(X40),X42))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X40)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X43),s('type.closSem.state'(X40),X44)))))) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X40)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X41),s('type.closSem.state'(X40),X42))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X40)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X43),s('type.closSem.state'(X40),X44)))))) ) | $spl4),
  inference(resolution,[],[f672,f122])).
fof(f1466,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X32)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X33),s('type.closSem.state'(X32),X34))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X32)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X35),s('type.closSem.state'(X32),X36)))))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X32)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X33),s('type.closSem.state'(X32),X34))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X32)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X35),s('type.closSem.state'(X32),X36))))) | s(bool,X38) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f672,f120])).
fof(f1464,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X16)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X17),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X16)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X19),s('type.closSem.state'(X16),X20)))))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X17),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp',X19),s('type.closSem.state'(X16),X20)))))) ) | $spl4),
  inference(resolution,[],[f672,f118])).
fof(f1457,plain,(
  $false | ($spl54 | $spl60)),
  inference(equality_resolution,[],[f1456])).
fof(f1456,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl54 | $spl60)),
  inference(subsumption_resolution,[],[f591,f488])).
fof(f1452,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl96),
  inference(cnf_transformation,[],[f1452_D])).
fof(f1452_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1450,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X56))))) ) | $spl1),
  inference(resolution,[],[f1311,f137])).
fof(f1311,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))))) ) | $spl1),
  inference(resolution,[],[f1250,f116])).
fof(f1449,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f1311,f124])).
fof(f1447,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X38))))) ) | $spl1),
  inference(resolution,[],[f1311,f123])).
fof(f1445,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X26))))) ) | $spl1),
  inference(resolution,[],[f1311,f122])).
fof(f1444,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f1311,f120])).
fof(f1442,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X8))))) ) | $spl1),
  inference(resolution,[],[f1311,f118])).
fof(f1438,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0)))) | p(s(bool,X16))) ) | $spl94),
  inference(cnf_transformation,[],[f1438_D])).
fof(f1438_D,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0)))) | p(s(bool,X16))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1437,plain,(
  ( ! [X39,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1413,f137])).
fof(f1436,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0)))) | s(bool,X37) != s(bool,X38)) ) | $spl1),
  inference(resolution,[],[f1413,f124])).
fof(f1433,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1413,f123])).
fof(f1432,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1413,f122])).
fof(f1431,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0)))) | s(bool,X22) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f1413,f120])).
fof(f1429,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK10(sK2,sK1,sK0),sK13(sK2,sK1,sK0),sK14(sK2,sK1,sK0),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0))),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),sK0)),s('type.num.num',sK10(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1413,f118])).
fof(f1016,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | ~p(s(bool,X1))) ) | $spl78),
  inference(cnf_transformation,[],[f1016_D])).
fof(f1016_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | ~p(s(bool,X1))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1423,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1251,f137])).
fof(f1422,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f1251,f124])).
fof(f1420,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1251,f123])).
fof(f1418,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1251,f122])).
fof(f1417,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f1251,f120])).
fof(f1415,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f1251,f118])).
fof(f1246,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl86),
  inference(cnf_transformation,[],[f1246_D])).
fof(f1246_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1410,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X56))))) ) | $spl1),
  inference(resolution,[],[f1310,f137])).
fof(f1310,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl1),
  inference(resolution,[],[f1250,f114])).
fof(f1409,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f1310,f124])).
fof(f1407,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X38))))) ) | $spl1),
  inference(resolution,[],[f1310,f123])).
fof(f1405,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X26))))) ) | $spl1),
  inference(resolution,[],[f1310,f122])).
fof(f1404,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f1310,f120])).
fof(f1402,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X8))))) ) | $spl1),
  inference(resolution,[],[f1310,f118])).
fof(f1397,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl92),
  inference(cnf_transformation,[],[f1397_D])).
fof(f1397_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1395,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X56))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1297,f137])).
fof(f1297,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK16(X13,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X12))))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1197,f116])).
fof(f1394,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X50)))) | s(bool,X53) != s(bool,X55)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1297,f124])).
fof(f1392,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X38))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1297,f123])).
fof(f1390,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X26))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1297,f122])).
fof(f1389,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X20)))) | s(bool,X23) != s(bool,X25)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1297,f120])).
fof(f1387,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X8))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1297,f118])).
fof(f1384,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl90),
  inference(cnf_transformation,[],[f1384_D])).
fof(f1384_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1382,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X56))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1296,f137])).
fof(f1296,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK15(X10,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X9))))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1197,f114])).
fof(f1381,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X50)))) | s(bool,X53) != s(bool,X55)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1296,f124])).
fof(f1379,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X38))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1296,f123])).
fof(f1377,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X26))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1296,f122])).
fof(f1376,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X20)))) | s(bool,X23) != s(bool,X25)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1296,f120])).
fof(f1374,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)),X8))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1296,f118])).
fof(f965,plain,(
  ( ! [X2,X0] : (s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X0))) | ~p(s(bool,X2)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))))) ) | $spl76),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  ( ! [X2,X0] : (s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X0))) | ~p(s(bool,X2)) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1370,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X32),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X32)))) ) | $spl3),
  inference(resolution,[],[f1283,f137])).
fof(f1283,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X0)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))))) ) | $spl3),
  inference(resolution,[],[f143,f101])).
fof(f1369,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X28),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X28))) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f1283,f124])).
fof(f1367,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X20),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X20)))) ) | $spl3),
  inference(resolution,[],[f1283,f123])).
fof(f1365,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X12),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X12)))) ) | $spl3),
  inference(resolution,[],[f1283,f122])).
fof(f1364,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X8),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X8))) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1283,f120])).
fof(f1362,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X0)))) ) | $spl3),
  inference(resolution,[],[f1283,f118])).
fof(f1209,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,X1))) ) | $spl84),
  inference(cnf_transformation,[],[f1209_D])).
fof(f1209_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,X1))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1359,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f137])).
fof(f1358,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X21) != s(bool,X23)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f124])).
fof(f1356,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f123])).
fof(f1354,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f122])).
fof(f1353,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X6) != s(bool,X8)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f120])).
fof(f1351,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f118])).
fof(f1173,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl82),
  inference(cnf_transformation,[],[f1173_D])).
fof(f1173_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1332,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X56))))) ) | $spl6),
  inference(resolution,[],[f1264,f137])).
fof(f1264,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))))) ) | $spl6),
  inference(resolution,[],[f147,f116])).
fof(f1331,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl6),
  inference(resolution,[],[f1264,f124])).
fof(f1329,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X38))))) ) | $spl6),
  inference(resolution,[],[f1264,f123])).
fof(f1327,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X26))))) ) | $spl6),
  inference(resolution,[],[f1264,f122])).
fof(f1326,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl6),
  inference(resolution,[],[f1264,f120])).
fof(f1324,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X8))))) ) | $spl6),
  inference(resolution,[],[f1264,f118])).
fof(f422,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))) | p(s(bool,X11))) ) | $spl48),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))) | p(s(bool,X11))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1320,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1250,f137])).
fof(f1319,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))) | s(bool,X32) != s(bool,X33)) ) | $spl1),
  inference(resolution,[],[f1250,f124])).
fof(f1316,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1250,f123])).
fof(f1315,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1250,f122])).
fof(f1314,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))) | s(bool,X17) != s(bool,X18)) ) | $spl1),
  inference(resolution,[],[f1250,f120])).
fof(f1312,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f1250,f118])).
fof(f1307,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3))) | p(s(bool,X16))) ) | $spl88),
  inference(cnf_transformation,[],[f1307_D])).
fof(f1307_D,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3))) | p(s(bool,X16))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1306,plain,(
  ( ! [X39,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1197,f137])).
fof(f1305,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3))) | s(bool,X37) != s(bool,X38)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1197,f124])).
fof(f1302,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1197,f123])).
fof(f1301,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1197,f122])).
fof(f1300,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3))) | s(bool,X22) != s(bool,X23)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1197,f120])).
fof(f1298,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(sK3,sK6,sK7,'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4)),'const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5)),sK0)),s('type.num.num',sK3)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f1197,f118])).
fof(f488,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,X2))) ) | $spl54),
  inference(cnf_transformation,[],[f488_D])).
fof(f488_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | ~p(s(bool,X2))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1292,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f137])).
fof(f1291,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X22) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f143,f124])).
fof(f1289,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f123])).
fof(f1287,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f122])).
fof(f1286,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X7) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f143,f120])).
fof(f1284,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f118])).
fof(f1282,plain,(
  $false | $spl50),
  inference(equality_resolution,[],[f423])).
fof(f340,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | p(s(bool,X11))) ) | $spl40),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | p(s(bool,X11))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1265,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f118])).
fof(f1267,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X17) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f147,f120])).
fof(f1268,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f122])).
fof(f1269,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f123])).
fof(f1272,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X32) != s(bool,X33)) ) | $spl6),
  inference(resolution,[],[f147,f124])).
fof(f1273,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f137])).
fof(f1281,plain,(
  $false | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f952,f147])).
fof(f952,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) | $spl44),
  inference(equality_resolution,[],[f364])).
fof(f423,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl50),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1279,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1273,f364])).
fof(f1278,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X32) != s(bool,X33)) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1272,f364])).
fof(f1277,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1269,f364])).
fof(f1276,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1268,f364])).
fof(f1275,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X17) != s(bool,X18)) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1267,f364])).
fof(f1274,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1265,f364])).
fof(f166,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | ~p(s(bool,X1))) ) | $spl12),
  inference(cnf_transformation,[],[f166_D])).
fof(f166_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | ~p(s(bool,X1))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1260,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f137])).
fof(f1259,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f141,f124])).
fof(f1257,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f123])).
fof(f1255,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f122])).
fof(f1254,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f141,f120])).
fof(f1252,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f118])).
fof(f141,plain,(
  ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f146,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) | $spl7),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1248,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) | ($spl3 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f1196,f149])).
fof(f1196,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f107])).
fof(f1060,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1054,f364])).
fof(f1244,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X56))))) ) | $spl48),
  inference(resolution,[],[f750,f137])).
fof(f750,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl48),
  inference(resolution,[],[f631,f114])).
fof(f631,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) | $spl48),
  inference(equality_resolution,[],[f422])).
fof(f1243,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl48),
  inference(resolution,[],[f750,f124])).
fof(f1241,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X38))))) ) | $spl48),
  inference(resolution,[],[f750,f123])).
fof(f1239,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X26))))) ) | $spl48),
  inference(resolution,[],[f750,f122])).
fof(f1238,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl48),
  inference(resolution,[],[f750,f120])).
fof(f1236,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))),X8))))) ) | $spl48),
  inference(resolution,[],[f750,f118])).
fof(f591,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl60),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1233,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f137])).
fof(f142,plain,(
  p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | $spl2),
  inference(cnf_transformation,[],[f142_D])).
fof(f142_D,plain,(
  p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1232,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X28) != s(bool,X29)) ) | $spl2),
  inference(resolution,[],[f142,f124])).
fof(f1229,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f123])).
fof(f1228,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f122])).
fof(f1227,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f142,f120])).
fof(f1225,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f118])).
fof(f1207,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f137])).
fof(f1206,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X21) != s(bool,X23)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f124])).
fof(f1204,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f123])).
fof(f1202,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f122])).
fof(f1201,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X6) != s(bool,X8)) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f120])).
fof(f1199,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK3),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f954,f118])).
fof(f1194,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X32),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X32)))) ) | $spl3),
  inference(resolution,[],[f1101,f137])).
fof(f1101,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X0)))) | ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))))) ) | $spl3),
  inference(resolution,[],[f143,f101])).
fof(f1193,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X28),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X28))) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f1101,f124])).
fof(f1191,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X20),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X20)))) ) | $spl3),
  inference(resolution,[],[f1101,f123])).
fof(f1189,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X12),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X12)))) ) | $spl3),
  inference(resolution,[],[f1101,f122])).
fof(f1188,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X8),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X8))) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f1101,f120])).
fof(f1186,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X0)))) ) | $spl3),
  inference(resolution,[],[f1101,f118])).
fof(f1182,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f137])).
fof(f1181,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f146,f124])).
fof(f1179,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f123])).
fof(f1177,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f122])).
fof(f1176,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f146,f120])).
fof(f1174,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f118])).
fof(f1171,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X56))))) ) | $spl6),
  inference(resolution,[],[f1045,f137])).
fof(f1045,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))))) ) | $spl6),
  inference(resolution,[],[f147,f116])).
fof(f1170,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl6),
  inference(resolution,[],[f1045,f124])).
fof(f1168,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X38))))) ) | $spl6),
  inference(resolution,[],[f1045,f123])).
fof(f1166,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X26))))) ) | $spl6),
  inference(resolution,[],[f1045,f122])).
fof(f1165,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl6),
  inference(resolution,[],[f1045,f120])).
fof(f1163,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X8))))) ) | $spl6),
  inference(resolution,[],[f1045,f118])).
fof(f933,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl74),
  inference(cnf_transformation,[],[f933_D])).
fof(f933_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1159,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X56))))) ) | $spl6),
  inference(resolution,[],[f1044,f137])).
fof(f1044,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X4))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl6),
  inference(resolution,[],[f147,f114])).
fof(f1158,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl6),
  inference(resolution,[],[f1044,f124])).
fof(f1156,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X38))))) ) | $spl6),
  inference(resolution,[],[f1044,f123])).
fof(f1154,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X26))))) ) | $spl6),
  inference(resolution,[],[f1044,f122])).
fof(f1153,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl6),
  inference(resolution,[],[f1044,f120])).
fof(f1151,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X8))))) ) | $spl6),
  inference(resolution,[],[f1044,f118])).
fof(f1110,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f137])).
fof(f1109,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X22) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f143,f124])).
fof(f1107,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f123])).
fof(f1105,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f122])).
fof(f1104,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X7) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f143,f120])).
fof(f1102,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f118])).
fof(f1098,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f137])).
fof(f1097,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f141,f124])).
fof(f1095,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f123])).
fof(f1093,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f122])).
fof(f1092,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f141,f120])).
fof(f1090,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f118])).
fof(f1086,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl80),
  inference(cnf_transformation,[],[f1086_D])).
fof(f1086_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1084,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X56))))) ) | $spl0),
  inference(resolution,[],[f1020,f137])).
fof(f1020,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X8))))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10))))) ) | $spl0),
  inference(resolution,[],[f140,f116])).
fof(f1083,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl0),
  inference(resolution,[],[f1020,f124])).
fof(f1081,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X38))))) ) | $spl0),
  inference(resolution,[],[f1020,f123])).
fof(f1079,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X26))))) ) | $spl0),
  inference(resolution,[],[f1020,f122])).
fof(f1078,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f1020,f120])).
fof(f1076,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X8))))) ) | $spl0),
  inference(resolution,[],[f1020,f118])).
fof(f894,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl72),
  inference(cnf_transformation,[],[f894_D])).
fof(f894_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1072,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X56))))) ) | $spl0),
  inference(resolution,[],[f1019,f137])).
fof(f1019,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK15(X6,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X5))))) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))))) ) | $spl0),
  inference(resolution,[],[f140,f114])).
fof(f1071,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl0),
  inference(resolution,[],[f1019,f124])).
fof(f1069,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X38))))) ) | $spl0),
  inference(resolution,[],[f1019,f123])).
fof(f1067,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X26))))) ) | $spl0),
  inference(resolution,[],[f1019,f122])).
fof(f1066,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f1019,f120])).
fof(f1064,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X8))))) ) | $spl0),
  inference(resolution,[],[f1019,f118])).
fof(f1046,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f118])).
fof(f1048,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X17) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f147,f120])).
fof(f1049,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f122])).
fof(f1050,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f123])).
fof(f1053,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X32) != s(bool,X33)) ) | $spl6),
  inference(resolution,[],[f147,f124])).
fof(f1054,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f137])).
fof(f1059,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X32) != s(bool,X33)) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1053,f364])).
fof(f1058,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1050,f364])).
fof(f1057,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1049,f364])).
fof(f1056,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X17) != s(bool,X18)) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1048,f364])).
fof(f1055,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl6 | $spl44)),
  inference(subsumption_resolution,[],[f1046,f364])).
fof(f461,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | p(s(bool,X36))) ) | $spl52),
  inference(cnf_transformation,[],[f461_D])).
fof(f461_D,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | p(s(bool,X36))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1029,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f137])).
fof(f1028,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X33) != s(bool,X34)) ) | $spl0),
  inference(resolution,[],[f140,f124])).
fof(f1025,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f123])).
fof(f1024,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f122])).
fof(f1023,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X18) != s(bool,X19)) ) | $spl0),
  inference(resolution,[],[f140,f120])).
fof(f1021,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f118])).
fof(f1014,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f897,f137])).
fof(f897,plain,(
  ~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) | $spl1),
  inference(resolution,[],[f141,f110])).
fof(f1013,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f897,f124])).
fof(f1011,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f897,f123])).
fof(f1009,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f897,f122])).
fof(f1008,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f897,f120])).
fof(f1006,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))))) ) | $spl1),
  inference(resolution,[],[f897,f118])).
fof(f1002,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f137])).
fof(f1001,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X28) != s(bool,X29)) ) | $spl2),
  inference(resolution,[],[f142,f124])).
fof(f998,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f123])).
fof(f997,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f122])).
fof(f996,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f142,f120])).
fof(f994,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f118])).
fof(f843,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl70),
  inference(cnf_transformation,[],[f843_D])).
fof(f843_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f988,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X56))))) ) | $spl1),
  inference(resolution,[],[f911,f137])).
fof(f911,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))))) ) | $spl1),
  inference(resolution,[],[f895,f116])).
fof(f895,plain,(
  p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) | $spl1),
  inference(resolution,[],[f141,f108])).
fof(f987,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f911,f124])).
fof(f985,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X38))))) ) | $spl1),
  inference(resolution,[],[f911,f123])).
fof(f983,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X26))))) ) | $spl1),
  inference(resolution,[],[f911,f122])).
fof(f982,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f911,f120])).
fof(f980,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8))))) ) | $spl1),
  inference(resolution,[],[f911,f118])).
fof(f668,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl62),
  inference(cnf_transformation,[],[f668_D])).
fof(f668_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f976,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X56))))) ) | $spl1),
  inference(resolution,[],[f910,f137])).
fof(f910,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl1),
  inference(resolution,[],[f895,f114])).
fof(f975,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f910,f124])).
fof(f973,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X38))))) ) | $spl1),
  inference(resolution,[],[f910,f123])).
fof(f971,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X26))))) ) | $spl1),
  inference(resolution,[],[f910,f122])).
fof(f970,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f910,f120])).
fof(f968,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8))))) ) | $spl1),
  inference(resolution,[],[f910,f118])).
fof(f963,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X32),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X32)))) ) | $spl3),
  inference(resolution,[],[f633,f137])).
fof(f962,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X28),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X28))) | s(bool,X29) != s(bool,X31)) ) | $spl3),
  inference(resolution,[],[f633,f124])).
fof(f960,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X20),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X20)))) ) | $spl3),
  inference(resolution,[],[f633,f123])).
fof(f958,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X12),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X12)))) ) | $spl3),
  inference(resolution,[],[f633,f122])).
fof(f957,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X8),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X8))) | s(bool,X9) != s(bool,X11)) ) | $spl3),
  inference(resolution,[],[f633,f120])).
fof(f955,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X0),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(sK0)),'const.pair.,_2'(s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',X0)))) ) | $spl3),
  inference(resolution,[],[f633,f118])).
fof(f935,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f118])).
fof(f937,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f146,f120])).
fof(f938,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f122])).
fof(f940,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f123])).
fof(f942,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f146,f124])).
fof(f943,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f137])).
fof(f950,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl7 | $spl40)),
  inference(subsumption_resolution,[],[f943,f340])).
fof(f949,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X21) != s(bool,X23)) ) | ($spl7 | $spl40)),
  inference(subsumption_resolution,[],[f942,f340])).
fof(f948,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl7 | $spl40)),
  inference(subsumption_resolution,[],[f940,f340])).
fof(f947,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl7 | $spl40)),
  inference(subsumption_resolution,[],[f938,f340])).
fof(f946,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X6) != s(bool,X8)) ) | ($spl7 | $spl40)),
  inference(subsumption_resolution,[],[f937,f340])).
fof(f945,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | ($spl7 | $spl40)),
  inference(subsumption_resolution,[],[f935,f340])).
fof(f931,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X56))))) ) | $spl6),
  inference(resolution,[],[f366,f137])).
fof(f366,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X4))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl6),
  inference(resolution,[],[f147,f114])).
fof(f930,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl6),
  inference(resolution,[],[f366,f124])).
fof(f928,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X38))))) ) | $spl6),
  inference(resolution,[],[f366,f123])).
fof(f926,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X26))))) ) | $spl6),
  inference(resolution,[],[f366,f122])).
fof(f925,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl6),
  inference(resolution,[],[f366,f120])).
fof(f923,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)),X8))))) ) | $spl6),
  inference(resolution,[],[f366,f118])).
fof(f353,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))) | p(s(bool,X11))) ) | $spl42),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))) | p(s(bool,X11))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f920,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f895,f137])).
fof(f919,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))) | s(bool,X32) != s(bool,X33)) ) | $spl1),
  inference(resolution,[],[f895,f124])).
fof(f916,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f895,f123])).
fof(f915,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f895,f122])).
fof(f914,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))) | s(bool,X17) != s(bool,X18)) ) | $spl1),
  inference(resolution,[],[f895,f120])).
fof(f912,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f895,f118])).
fof(f906,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f137])).
fof(f905,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f141,f124])).
fof(f903,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f123])).
fof(f901,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f122])).
fof(f900,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f141,f120])).
fof(f898,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f118])).
fof(f892,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X56))))) ) | $spl0),
  inference(resolution,[],[f857,f137])).
fof(f857,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK15(X6,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X5))))) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))))) ) | $spl0),
  inference(resolution,[],[f140,f114])).
fof(f891,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl0),
  inference(resolution,[],[f857,f124])).
fof(f889,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X38))))) ) | $spl0),
  inference(resolution,[],[f857,f123])).
fof(f887,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X26))))) ) | $spl0),
  inference(resolution,[],[f857,f122])).
fof(f886,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f857,f120])).
fof(f884,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)),X8))))) ) | $spl0),
  inference(resolution,[],[f857,f118])).
fof(f773,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl68),
  inference(cnf_transformation,[],[f773_D])).
fof(f773_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f880,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X56))))) ) | $spl8),
  inference(resolution,[],[f846,f137])).
fof(f846,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))))) ) | $spl8),
  inference(resolution,[],[f149,f116])).
fof(f879,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl8),
  inference(resolution,[],[f846,f124])).
fof(f877,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X38))))) ) | $spl8),
  inference(resolution,[],[f846,f123])).
fof(f875,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X26))))) ) | $spl8),
  inference(resolution,[],[f846,f122])).
fof(f874,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl8),
  inference(resolution,[],[f846,f120])).
fof(f872,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X8))))) ) | $spl8),
  inference(resolution,[],[f846,f118])).
fof(f867,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f137])).
fof(f866,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X33) != s(bool,X34)) ) | $spl0),
  inference(resolution,[],[f140,f124])).
fof(f863,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f123])).
fof(f862,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f122])).
fof(f861,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X18) != s(bool,X19)) ) | $spl0),
  inference(resolution,[],[f140,f120])).
fof(f859,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f118])).
fof(f224,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | p(s(bool,X11))) ) | $spl32),
  inference(cnf_transformation,[],[f224_D])).
fof(f224_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | p(s(bool,X11))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f855,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f137])).
fof(f854,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X32) != s(bool,X33)) ) | $spl8),
  inference(resolution,[],[f149,f124])).
fof(f851,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f123])).
fof(f850,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f122])).
fof(f849,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X17) != s(bool,X18)) ) | $spl8),
  inference(resolution,[],[f149,f120])).
fof(f847,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f118])).
fof(f841,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X56))))) ) | $spl42),
  inference(resolution,[],[f714,f137])).
fof(f714,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))))) ) | $spl42),
  inference(resolution,[],[f377,f116])).
fof(f840,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl42),
  inference(resolution,[],[f714,f124])).
fof(f838,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X38))))) ) | $spl42),
  inference(resolution,[],[f714,f123])).
fof(f836,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X26))))) ) | $spl42),
  inference(resolution,[],[f714,f122])).
fof(f835,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl42),
  inference(resolution,[],[f714,f120])).
fof(f833,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8))))) ) | $spl42),
  inference(resolution,[],[f714,f118])).
fof(f829,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X56))))) ) | $spl42),
  inference(resolution,[],[f713,f137])).
fof(f713,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl42),
  inference(resolution,[],[f377,f114])).
fof(f828,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl42),
  inference(resolution,[],[f713,f124])).
fof(f826,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X38))))) ) | $spl42),
  inference(resolution,[],[f713,f123])).
fof(f824,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X26))))) ) | $spl42),
  inference(resolution,[],[f713,f122])).
fof(f823,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl42),
  inference(resolution,[],[f713,f120])).
fof(f821,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8))))) ) | $spl42),
  inference(resolution,[],[f713,f118])).
fof(f736,plain,(
  ( ! [X14,X21,X19,X17,X15,X18,X16] : (s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X14))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X14),s('type.closSem.state'(sK0),X15),s('type.closSem.state'(sK0),X16)))) | ~p(s(bool,X21)) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X19),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X17))),s('type.closSem.state'(sK0),X15))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X19),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X18))),s('type.closSem.state'(sK0),X16)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X14))),s('type.list.list'('type.closSem.v'),X17),s('type.list.list'('type.closSem.v'),X18))))) ) | $spl66),
  inference(cnf_transformation,[],[f736_D])).
fof(f736_D,plain,(
  ( ! [X14,X21,X19,X17,X15,X18,X16] : (s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X14))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X14),s('type.closSem.state'(sK0),X15),s('type.closSem.state'(sK0),X16)))) | ~p(s(bool,X21)) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X19),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X17))),s('type.closSem.state'(sK0),X15))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X19),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X18))),s('type.closSem.state'(sK0),X16)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X14))),s('type.list.list'('type.closSem.v'),X17),s('type.list.list'('type.closSem.v'),X18))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f817,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89,X86] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X86),s('type.closSem.state'(sK0),X87),s('type.closSem.state'(sK0),X88)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X86))),s('type.list.list'('type.closSem.v'),X89),s('type.list.list'('type.closSem.v'),X90)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X91),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X89))),s('type.closSem.state'(sK0),X87))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X91),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X90))),s('type.closSem.state'(sK0),X88)))))) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X91),s('type.num.num',X86)))) ) | $spl10),
  inference(resolution,[],[f152,f137])).
fof(f816,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X77,X84,X82] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X77),s('type.closSem.state'(sK0),X78),s('type.closSem.state'(sK0),X79)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X77))),s('type.list.list'('type.closSem.v'),X80),s('type.list.list'('type.closSem.v'),X81)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X82),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X80))),s('type.closSem.state'(sK0),X78))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X82),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X81))),s('type.closSem.state'(sK0),X79)))))) | ~p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X82),s('type.num.num',X77))) | s(bool,X83) != s(bool,X85)) ) | $spl10),
  inference(resolution,[],[f152,f124])).
fof(f814,plain,(
  ( ! [X61,X59,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X59),s('type.closSem.state'(sK0),X60),s('type.closSem.state'(sK0),X61)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X59))),s('type.list.list'('type.closSem.v'),X62),s('type.list.list'('type.closSem.v'),X63)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X64),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X62))),s('type.closSem.state'(sK0),X60))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X64),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X63))),s('type.closSem.state'(sK0),X61)))))) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X64),s('type.num.num',X59)))) ) | $spl10),
  inference(resolution,[],[f152,f123])).
fof(f812,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X49] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X41),s('type.closSem.state'(sK0),X42),s('type.closSem.state'(sK0),X43)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X41))),s('type.list.list'('type.closSem.v'),X44),s('type.list.list'('type.closSem.v'),X45)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X46),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X44))),s('type.closSem.state'(sK0),X42))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X46),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X45))),s('type.closSem.state'(sK0),X43)))))) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X46),s('type.num.num',X41)))) ) | $spl10),
  inference(resolution,[],[f152,f122])).
fof(f811,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32,X40] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X32),s('type.closSem.state'(sK0),X33),s('type.closSem.state'(sK0),X34)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X32))),s('type.list.list'('type.closSem.v'),X35),s('type.list.list'('type.closSem.v'),X36)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X37),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X35))),s('type.closSem.state'(sK0),X33))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X37),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X36))),s('type.closSem.state'(sK0),X34)))))) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X39) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X37),s('type.num.num',X32))) | s(bool,X38) != s(bool,X40)) ) | $spl10),
  inference(resolution,[],[f152,f120])).
fof(f809,plain,(
  ( ! [X14,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X14),s('type.closSem.state'(sK0),X15),s('type.closSem.state'(sK0),X16)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X14))),s('type.list.list'('type.closSem.v'),X17),s('type.list.list'('type.closSem.v'),X18)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X19),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X17))),s('type.closSem.state'(sK0),X15))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X19),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X18))),s('type.closSem.state'(sK0),X16)))))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X14)))) ) | $spl10),
  inference(resolution,[],[f152,f118])).
fof(f788,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f118])).
fof(f790,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f141,f120])).
fof(f791,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f122])).
fof(f793,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f123])).
fof(f795,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f141,f124])).
fof(f796,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f137])).
fof(f168,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl16),
  inference(cnf_transformation,[],[f168_D])).
fof(f168_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f803,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | ($spl1 | $spl52)),
  inference(subsumption_resolution,[],[f796,f461])).
fof(f802,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | s(bool,X22) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X21) != s(bool,X23)) ) | ($spl1 | $spl52)),
  inference(subsumption_resolution,[],[f795,f461])).
fof(f801,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | ($spl1 | $spl52)),
  inference(subsumption_resolution,[],[f793,f461])).
fof(f800,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | ($spl1 | $spl52)),
  inference(subsumption_resolution,[],[f791,f461])).
fof(f799,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | s(bool,X7) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X6) != s(bool,X8)) ) | ($spl1 | $spl52)),
  inference(subsumption_resolution,[],[f790,f461])).
fof(f798,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | ($spl1 | $spl52)),
  inference(subsumption_resolution,[],[f788,f461])).
fof(f265,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | ~p(s(bool,X25))) ) | $spl38),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | ~p(s(bool,X25))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f783,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f137])).
fof(f148,plain,(
  ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) | $spl9),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f782,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl9),
  inference(resolution,[],[f148,f124])).
fof(f780,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f123])).
fof(f778,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f122])).
fof(f777,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl9),
  inference(resolution,[],[f148,f120])).
fof(f775,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f118])).
fof(f771,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X56))))) ) | $spl8),
  inference(resolution,[],[f686,f137])).
fof(f686,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK16(X8,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))))) ) | $spl8),
  inference(resolution,[],[f149,f116])).
fof(f770,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl8),
  inference(resolution,[],[f686,f124])).
fof(f768,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X38))))) ) | $spl8),
  inference(resolution,[],[f686,f123])).
fof(f766,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X26))))) ) | $spl8),
  inference(resolution,[],[f686,f122])).
fof(f765,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl8),
  inference(resolution,[],[f686,f120])).
fof(f763,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X8))))) ) | $spl8),
  inference(resolution,[],[f686,f118])).
fof(f526,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl56),
  inference(cnf_transformation,[],[f526_D])).
fof(f526_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f747,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X56))))) ) | $spl4),
  inference(resolution,[],[f674,f137])).
fof(f674,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8))))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10))))) ) | $spl4),
  inference(resolution,[],[f145,f116])).
fof(f746,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl4),
  inference(resolution,[],[f674,f124])).
fof(f744,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X38))))) ) | $spl4),
  inference(resolution,[],[f674,f123])).
fof(f742,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X26))))) ) | $spl4),
  inference(resolution,[],[f674,f122])).
fof(f741,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f674,f120])).
fof(f739,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8))))) ) | $spl4),
  inference(resolution,[],[f674,f118])).
fof(f734,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89,X86] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X86),s('type.closSem.state'(sK0),X87),s('type.closSem.state'(sK0),X88)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X86))),s('type.list.list'('type.closSem.v'),X89),s('type.list.list'('type.closSem.v'),X90)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X91),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X89))),s('type.closSem.state'(sK0),X87))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X91),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X90))),s('type.closSem.state'(sK0),X88)))))) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X91),s('type.num.num',X86)))) ) | $spl10),
  inference(resolution,[],[f152,f137])).
fof(f733,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X77,X84,X82] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X77),s('type.closSem.state'(sK0),X78),s('type.closSem.state'(sK0),X79)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X77))),s('type.list.list'('type.closSem.v'),X80),s('type.list.list'('type.closSem.v'),X81)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X82),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X80))),s('type.closSem.state'(sK0),X78))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X82),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X81))),s('type.closSem.state'(sK0),X79)))))) | ~p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X84) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X82),s('type.num.num',X77))) | s(bool,X83) != s(bool,X85)) ) | $spl10),
  inference(resolution,[],[f152,f124])).
fof(f731,plain,(
  ( ! [X61,X59,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X59),s('type.closSem.state'(sK0),X60),s('type.closSem.state'(sK0),X61)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X59))),s('type.list.list'('type.closSem.v'),X62),s('type.list.list'('type.closSem.v'),X63)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X64),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X62))),s('type.closSem.state'(sK0),X60))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X64),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X63))),s('type.closSem.state'(sK0),X61)))))) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X64),s('type.num.num',X59)))) ) | $spl10),
  inference(resolution,[],[f152,f123])).
fof(f729,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X49] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X41),s('type.closSem.state'(sK0),X42),s('type.closSem.state'(sK0),X43)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X41))),s('type.list.list'('type.closSem.v'),X44),s('type.list.list'('type.closSem.v'),X45)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X46),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X44))),s('type.closSem.state'(sK0),X42))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X46),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X45))),s('type.closSem.state'(sK0),X43)))))) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X46),s('type.num.num',X41)))) ) | $spl10),
  inference(resolution,[],[f152,f122])).
fof(f728,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32,X40] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X32),s('type.closSem.state'(sK0),X33),s('type.closSem.state'(sK0),X34)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X32))),s('type.list.list'('type.closSem.v'),X35),s('type.list.list'('type.closSem.v'),X36)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X37),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X35))),s('type.closSem.state'(sK0),X33))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X37),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X36))),s('type.closSem.state'(sK0),X34)))))) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X39) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X37),s('type.num.num',X32))) | s(bool,X38) != s(bool,X40)) ) | $spl10),
  inference(resolution,[],[f152,f120])).
fof(f726,plain,(
  ( ! [X14,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X14),s('type.closSem.state'(sK0),X15),s('type.closSem.state'(sK0),X16)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X14))),s('type.list.list'('type.closSem.v'),X17),s('type.list.list'('type.closSem.v'),X18)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X19),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X17))),s('type.closSem.state'(sK0),X15))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X19),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X18))),s('type.closSem.state'(sK0),X16)))))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X14)))) ) | $spl10),
  inference(resolution,[],[f152,f118])).
fof(f708,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f137])).
fof(f707,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X33) != s(bool,X34)) ) | $spl0),
  inference(resolution,[],[f140,f124])).
fof(f704,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f123])).
fof(f703,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f122])).
fof(f702,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X18) != s(bool,X19)) ) | $spl0),
  inference(resolution,[],[f140,f120])).
fof(f700,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f118])).
fof(f695,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f137])).
fof(f694,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X32) != s(bool,X33)) ) | $spl8),
  inference(resolution,[],[f149,f124])).
fof(f225,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl34),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f169,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X3)) | p(s(bool,X4))) ) | $spl18),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X3)) | p(s(bool,X4))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f693,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f124])).
fof(f691,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f123])).
fof(f690,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f122])).
fof(f689,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X17) != s(bool,X18)) ) | $spl8),
  inference(resolution,[],[f149,f120])).
fof(f688,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X14) | s(bool,X15) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f120])).
fof(f687,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f118])).
fof(f186,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | p(s(bool,X12))) ) | $spl26),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | p(s(bool,X12))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f683,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f137])).
fof(f682,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X33) != s(bool,X34)) ) | $spl4),
  inference(resolution,[],[f145,f124])).
fof(f187,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl28),
  inference(cnf_transformation,[],[f187_D])).
fof(f187_D,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f681,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X30) | s(bool,X31) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f124])).
fof(f679,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f123])).
fof(f678,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f122])).
fof(f677,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X18) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f145,f120])).
fof(f676,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f120])).
fof(f675,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f118])).
fof(f670,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl64),
  inference(cnf_transformation,[],[f670_D])).
fof(f670_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f635,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X4) != s(bool,X5)) ) | $spl3),
  inference(resolution,[],[f143,f118])).
fof(f638,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X13) != s(bool,X14)) ) | $spl3),
  inference(resolution,[],[f143,f122])).
fof(f666,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X56))))) ) | $spl1),
  inference(resolution,[],[f491,f137])).
fof(f491,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X4))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl1),
  inference(resolution,[],[f463,f114])).
fof(f463,plain,(
  p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) | $spl1),
  inference(resolution,[],[f141,f108])).
fof(f665,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl1),
  inference(resolution,[],[f491,f124])).
fof(f663,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X38))))) ) | $spl1),
  inference(resolution,[],[f491,f123])).
fof(f661,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X26))))) ) | $spl1),
  inference(resolution,[],[f491,f122])).
fof(f660,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f491,f120])).
fof(f658,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))),X8))))) ) | $spl1),
  inference(resolution,[],[f491,f118])).
fof(f642,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f137])).
fof(f641,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X22) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f143,f124])).
fof(f639,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f123])).
fof(f637,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f122])).
fof(f636,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X7) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f143,f120])).
fof(f634,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f118])).
fof(f615,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f137])).
fof(f614,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl9),
  inference(resolution,[],[f148,f124])).
fof(f612,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f123])).
fof(f610,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f122])).
fof(f609,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl9),
  inference(resolution,[],[f148,f120])).
fof(f607,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f118])).
fof(f606,plain,(
  $false | ($spl32 | $spl38)),
  inference(equality_resolution,[],[f563])).
fof(f563,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | ($spl32 | $spl38)),
  inference(subsumption_resolution,[],[f265,f224])).
fof(f605,plain,(
  $false | $spl34),
  inference(equality_resolution,[],[f225])).
fof(f592,plain,(
  $false | ($spl9 | $spl32)),
  inference(subsumption_resolution,[],[f326,f148])).
fof(f326,plain,(
  p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) | $spl32),
  inference(equality_resolution,[],[f224])).
fof(f590,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f137])).
fof(f589,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X28) != s(bool,X29)) ) | $spl2),
  inference(resolution,[],[f142,f124])).
fof(f586,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f123])).
fof(f585,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f122])).
fof(f584,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f142,f120])).
fof(f582,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl2),
  inference(resolution,[],[f142,f118])).
fof(f579,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | ($spl9 | $spl32)),
  inference(subsumption_resolution,[],[f572,f224])).
fof(f572,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f137])).
fof(f578,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | s(bool,X22) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X21) != s(bool,X23)) ) | ($spl9 | $spl32)),
  inference(subsumption_resolution,[],[f571,f224])).
fof(f571,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl9),
  inference(resolution,[],[f148,f124])).
fof(f577,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | ($spl9 | $spl32)),
  inference(subsumption_resolution,[],[f569,f224])).
fof(f569,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f123])).
fof(f576,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | ($spl9 | $spl32)),
  inference(subsumption_resolution,[],[f567,f224])).
fof(f567,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f122])).
fof(f575,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X6) != s(bool,X8)) ) | ($spl9 | $spl32)),
  inference(subsumption_resolution,[],[f566,f224])).
fof(f566,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl9),
  inference(resolution,[],[f148,f120])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | ($spl9 | $spl32)),
  inference(subsumption_resolution,[],[f564,f224])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f118])).
fof(f562,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl58),
  inference(cnf_transformation,[],[f562_D])).
fof(f562_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f560,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X56))))) ) | $spl8),
  inference(resolution,[],[f299,f137])).
fof(f299,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,sK15(X5,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X4))))) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) ) | $spl8),
  inference(resolution,[],[f149,f114])).
fof(f559,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl8),
  inference(resolution,[],[f299,f124])).
fof(f557,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X38))))) ) | $spl8),
  inference(resolution,[],[f299,f123])).
fof(f555,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X26))))) ) | $spl8),
  inference(resolution,[],[f299,f122])).
fof(f554,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl8),
  inference(resolution,[],[f299,f120])).
fof(f552,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)),X8))))) ) | $spl8),
  inference(resolution,[],[f299,f118])).
fof(f549,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f464,f137])).
fof(f464,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) | $spl1),
  inference(resolution,[],[f141,f109])).
fof(f548,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))) | s(bool,X32) != s(bool,X33)) ) | $spl1),
  inference(resolution,[],[f464,f124])).
fof(f545,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f464,f123])).
fof(f544,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f464,f122])).
fof(f543,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))) | s(bool,X17) != s(bool,X18)) ) | $spl1),
  inference(resolution,[],[f464,f120])).
fof(f541,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f464,f118])).
fof(f198,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f535,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f137])).
fof(f144,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) | $spl5),
  inference(cnf_transformation,[],[f144_D])).
fof(f144_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f534,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f144,f124])).
fof(f532,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f123])).
fof(f530,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f122])).
fof(f529,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f144,f120])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f118])).
fof(f524,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK16(X57,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X56))))) ) | $spl4),
  inference(resolution,[],[f427,f137])).
fof(f427,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8))))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10))))) ) | $spl4),
  inference(resolution,[],[f145,f116])).
fof(f523,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK16(X51,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl4),
  inference(resolution,[],[f427,f124])).
fof(f521,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK16(X39,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X38))))) ) | $spl4),
  inference(resolution,[],[f427,f123])).
fof(f519,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK16(X27,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X26))))) ) | $spl4),
  inference(resolution,[],[f427,f122])).
fof(f518,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK16(X21,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f427,f120])).
fof(f516,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK16(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8))))) ) | $spl4),
  inference(resolution,[],[f427,f118])).
fof(f390,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) | $spl46),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  ( ! [X12,X10,X8,X9] : (s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8)))) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | ~p(s(bool,X12))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f512,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X56))))) ) | $spl4),
  inference(resolution,[],[f426,f137])).
fof(f426,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK15(X6,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X5))))) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))))) ) | $spl4),
  inference(resolution,[],[f145,f114])).
fof(f511,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl4),
  inference(resolution,[],[f426,f124])).
fof(f509,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X38))))) ) | $spl4),
  inference(resolution,[],[f426,f123])).
fof(f507,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X26))))) ) | $spl4),
  inference(resolution,[],[f426,f122])).
fof(f506,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f426,f120])).
fof(f504,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8))))) ) | $spl4),
  inference(resolution,[],[f426,f118])).
fof(f486,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f137])).
fof(f485,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X22) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f143,f124])).
fof(f483,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f123])).
fof(f481,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f122])).
fof(f480,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7))))) | s(bool,X7) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f143,f120])).
fof(f478,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) ) | $spl3),
  inference(resolution,[],[f143,f118])).
fof(f474,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f137])).
fof(f473,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f141,f124])).
fof(f471,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f123])).
fof(f469,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f122])).
fof(f468,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f141,f120])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f118])).
fof(f446,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f118])).
fof(f448,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X18) != s(bool,X19)) ) | $spl0),
  inference(resolution,[],[f140,f120])).
fof(f449,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f122])).
fof(f450,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f123])).
fof(f453,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X33) != s(bool,X34)) ) | $spl0),
  inference(resolution,[],[f140,f124])).
fof(f454,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f137])).
fof(f460,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | s(bool,X36) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f454,f166])).
fof(f459,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X33) != s(bool,X34)) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f453,f166])).
fof(f452,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X30) | s(bool,X31) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f124])).
fof(f458,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f450,f166])).
fof(f457,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f449,f166])).
fof(f456,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X18) != s(bool,X19)) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f448,f166])).
fof(f447,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl0),
  inference(resolution,[],[f140,f120])).
fof(f455,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f446,f166])).
fof(f428,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f118])).
fof(f430,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X18) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f145,f120])).
fof(f431,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f122])).
fof(f432,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f123])).
fof(f435,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X33) != s(bool,X34)) ) | $spl4),
  inference(resolution,[],[f145,f124])).
fof(f436,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f137])).
fof(f442,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | s(bool,X36) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f436,f198])).
fof(f441,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X33) != s(bool,X34)) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f435,f198])).
fof(f434,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X30) | s(bool,X31) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f124])).
fof(f440,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f432,f198])).
fof(f439,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f431,f198])).
fof(f438,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X18) != s(bool,X19)) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f430,f198])).
fof(f429,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f120])).
fof(f437,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f428,f198])).
fof(f369,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X14) | s(bool,X15) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f120])).
fof(f374,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f124])).
fof(f421,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f312,f137])).
fof(f312,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) | $spl1),
  inference(resolution,[],[f141,f109])).
fof(f420,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))) | s(bool,X32) != s(bool,X33)) ) | $spl1),
  inference(resolution,[],[f312,f124])).
fof(f417,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f312,f123])).
fof(f416,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f312,f122])).
fof(f415,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0)))) | s(bool,X17) != s(bool,X18)) ) | $spl1),
  inference(resolution,[],[f312,f120])).
fof(f413,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK10(sK2,sK1,sK0)))),s('type.list.list'('type.closSem.v'),sK11(sK2,sK1,sK0)),s('type.list.list'('type.closSem.v'),sK12(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f312,f118])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f118])).
fof(f394,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f144,f120])).
fof(f395,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f122])).
fof(f397,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f123])).
fof(f399,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f144,f124])).
fof(f400,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f137])).
fof(f407,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | ($spl5 | $spl26)),
  inference(subsumption_resolution,[],[f400,f186])).
fof(f406,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X21) != s(bool,X23)) ) | ($spl5 | $spl26)),
  inference(subsumption_resolution,[],[f399,f186])).
fof(f405,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | ($spl5 | $spl26)),
  inference(subsumption_resolution,[],[f397,f186])).
fof(f404,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | ($spl5 | $spl26)),
  inference(subsumption_resolution,[],[f395,f186])).
fof(f403,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X6) != s(bool,X8)) ) | ($spl5 | $spl26)),
  inference(subsumption_resolution,[],[f394,f186])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | ($spl5 | $spl26)),
  inference(subsumption_resolution,[],[f392,f186])).
fof(f388,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,happ(s(fun(X56,bool),X57),s(X56,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(X56,bool),X57),s(X56,sK15(X57,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X56))))) ) | $spl4),
  inference(resolution,[],[f201,f137])).
fof(f201,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK15(X6,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X5))))) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))))) ) | $spl4),
  inference(resolution,[],[f145,f114])).
fof(f387,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,happ(s(fun(X50,bool),X51),s(X50,sK15(X51,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X50)))) | s(bool,X53) != s(bool,X55)) ) | $spl4),
  inference(resolution,[],[f201,f124])).
fof(f385,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,happ(s(fun(X38,bool),X39),s(X38,sK15(X39,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X38))))) ) | $spl4),
  inference(resolution,[],[f201,f123])).
fof(f383,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,happ(s(fun(X26,bool),X27),s(X26,sK15(X27,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X26))))) ) | $spl4),
  inference(resolution,[],[f201,f122])).
fof(f382,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun(X20,bool),X21),s(X20,sK15(X21,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X20)))) | s(bool,X23) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f201,f120])).
fof(f380,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(X8,bool),X9),s(X8,sK15(X9,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)),X8))))) ) | $spl4),
  inference(resolution,[],[f201,f118])).
fof(f376,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f137])).
fof(f375,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X32) != s(bool,X33)) ) | $spl6),
  inference(resolution,[],[f147,f124])).
fof(f372,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f123])).
fof(f371,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f122])).
fof(f370,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X17) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f147,f120])).
fof(f368,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f118])).
fof(f362,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f137])).
fof(f361,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f146,f124])).
fof(f359,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f123])).
fof(f357,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f122])).
fof(f356,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f146,f120])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl7),
  inference(resolution,[],[f146,f118])).
fof(f352,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f311,f137])).
fof(f311,plain,(
  p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) | $spl1),
  inference(resolution,[],[f141,f108])).
fof(f351,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))) | s(bool,X32) != s(bool,X33)) ) | $spl1),
  inference(resolution,[],[f311,f124])).
fof(f348,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f311,f123])).
fof(f347,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f311,f122])).
fof(f346,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0)))) | s(bool,X17) != s(bool,X18)) ) | $spl1),
  inference(resolution,[],[f311,f120])).
fof(f344,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK13(sK2,sK1,sK0)),s('type.closSem.state'(sK0),sK14(sK2,sK1,sK0))))) ) | $spl1),
  inference(resolution,[],[f311,f118])).
fof(f339,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f137])).
fof(f338,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X32) != s(bool,X33)) ) | $spl6),
  inference(resolution,[],[f147,f124])).
fof(f335,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f123])).
fof(f334,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f122])).
fof(f333,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5))) | s(bool,X17) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f147,f120])).
fof(f331,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) ) | $spl6),
  inference(resolution,[],[f147,f118])).
fof(f325,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f168])).
fof(f322,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f137])).
fof(f321,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f141,f124])).
fof(f319,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f123])).
fof(f317,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f122])).
fof(f316,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f141,f120])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f118])).
fof(f309,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f137])).
fof(f308,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X32) != s(bool,X33)) ) | $spl8),
  inference(resolution,[],[f149,f124])).
fof(f307,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f124])).
fof(f305,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f123])).
fof(f304,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f122])).
fof(f303,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X17) != s(bool,X18)) ) | $spl8),
  inference(resolution,[],[f149,f120])).
fof(f302,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X14) | s(bool,X15) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f120])).
fof(f301,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f118])).
fof(f297,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f289,f267])).
fof(f267,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(equality_resolution,[],[f170])).
fof(f289,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK9(X70,X71,X72,X73,X74,X75)),s('type.num.num',X70))))) ) | $spl20),
  inference(resolution,[],[f267,f102])).
fof(f296,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f284,f267])).
fof(f284,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (p(s(bool,'const.clos_relation.exec_rel_3'(s('type.num.num',X51),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X52)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X53),s('type.closSem.state'(X52),X54))),s('type.pair.prod'('type.clos_relation.val_or_exp','type.closSem.state'(X52)),'const.pair.,_2'(s('type.clos_relation.val_or_exp',X55),s('type.closSem.state'(X52),X56))))))) ) | $spl20),
  inference(resolution,[],[f267,f102])).
fof(f295,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f283,f267])).
fof(f283,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X48),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X49),s('type.list.list'('type.closLang.exp'),X50))))) ) | $spl20),
  inference(resolution,[],[f267,f109])).
fof(f294,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f282,f267])).
fof(f282,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X45),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X46),s('type.list.list'('type.closLang.exp'),X47))))) ) | $spl20),
  inference(resolution,[],[f267,f108])).
fof(f293,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f280,f267])).
fof(f280,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X39),'const.bool.the_value_0'),s('type.num.num',sK10(X40,X41,X39)))),s('type.list.list'('type.closSem.v'),sK11(X40,X41,X39)),s('type.list.list'('type.closSem.v'),sK12(X40,X41,X39)))))) ) | $spl20),
  inference(resolution,[],[f267,f109])).
fof(f292,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f279,f267])).
fof(f279,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK10(X36,X37,X38)),s('type.closSem.state'(X38),sK13(X36,X37,X38)),s('type.closSem.state'(X38),sK14(X36,X37,X38)))))) ) | $spl20),
  inference(resolution,[],[f267,f108])).
fof(f278,plain,(
  $false | $spl20),
  inference(resolution,[],[f267,f132])).
fof(f132,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relationProps.exp_rel_thm',aHL_TRUTH)).
fof(f285,plain,(
  $false | ($spl4 | $spl20)),
  inference(resolution,[],[f267,f145])).
fof(f248,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f118])).
fof(f249,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X3) != s(bool,X4)) ) | $spl9),
  inference(resolution,[],[f148,f118])).
fof(f250,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X6) != s(bool,X8)) ) | $spl9),
  inference(resolution,[],[f148,f120])).
fof(f251,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f122])).
fof(f252,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X12) != s(bool,X13)) ) | $spl9),
  inference(resolution,[],[f148,f122])).
fof(f253,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f123])).
fof(f255,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X21) != s(bool,X23)) ) | $spl9),
  inference(resolution,[],[f148,f124])).
fof(f256,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl9),
  inference(resolution,[],[f148,f137])).
fof(f264,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | s(bool,X25) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f256,f167])).
fof(f260,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8)) ) | $spl36),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f263,plain,(
  ( ! [X23,X21,X22] : (s(bool,X22) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X21) != s(bool,X23)) ) | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f262,f167])).
fof(f262,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X21) != s(bool,X23)) ) | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f255,f167])).
fof(f261,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f253,f167])).
fof(f259,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X6) != s(bool,X8)) ) | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f258,f167])).
fof(f258,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X6) != s(bool,X8)) ) | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f250,f167])).
fof(f237,plain,(
  $false | $spl24),
  inference(resolution,[],[f172,f131])).
fof(f131,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_relationProps.exp_rel_thm',aHL_FALSITY)).
fof(f223,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f137])).
fof(f222,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X32) != s(bool,X33)) ) | $spl8),
  inference(resolution,[],[f149,f124])).
fof(f221,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f124])).
fof(f219,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X24) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f123])).
fof(f218,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f122])).
fof(f217,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7))) | s(bool,X17) != s(bool,X18)) ) | $spl8),
  inference(resolution,[],[f149,f120])).
fof(f216,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X14) | s(bool,X15) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f120])).
fof(f215,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) ) | $spl8),
  inference(resolution,[],[f149,f118])).
fof(f211,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f137])).
fof(f210,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X33) != s(bool,X34)) ) | $spl4),
  inference(resolution,[],[f145,f124])).
fof(f209,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X30) | s(bool,X31) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f124])).
fof(f207,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f123])).
fof(f206,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f122])).
fof(f205,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X18) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f145,f120])).
fof(f204,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f120])).
fof(f203,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f118])).
fof(f196,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f137])).
fof(f195,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f144,f124])).
fof(f193,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f123])).
fof(f192,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f144,f122])).
fof(f191,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f122])).
fof(f190,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f144,f120])).
fof(f189,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f144,f118])).
fof(f188,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl5),
  inference(resolution,[],[f144,f118])).
fof(f185,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f137])).
fof(f184,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X33) != s(bool,X34)) ) | $spl4),
  inference(resolution,[],[f145,f124])).
fof(f183,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X30) | s(bool,X31) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f124])).
fof(f181,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f123])).
fof(f180,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f122])).
fof(f179,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3))) | s(bool,X18) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f145,f120])).
fof(f178,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f120])).
fof(f177,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) ) | $spl4),
  inference(resolution,[],[f145,f118])).
fof(f164,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f137])).
fof(f163,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f141,f124])).
fof(f161,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f123])).
fof(f160,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f141,f122])).
fof(f159,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f122])).
fof(f158,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f141,f120])).
fof(f157,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f141,f118])).
fof(f156,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) ) | $spl1),
  inference(resolution,[],[f141,f118])).
fof(f140,plain,(
  p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f152,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(sK0),X12),s('type.closSem.state'(sK0),X13)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(sK0),X12))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X11))),s('type.closSem.state'(sK0),X13))))))) ) | $spl10),
  inference(cnf_transformation,[],[f152_D])).
fof(f152_D,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(sK0),X12),s('type.closSem.state'(sK0),X13)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(sK0),X12))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X11))),s('type.closSem.state'(sK0),X13))))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f151,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(sK0),X12))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X11))),s('type.closSem.state'(sK0),X13)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(sK0),X12),s('type.closSem.state'(sK0),X13)))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))) )),
  inference(forward_demodulation,[],[f150,f105])).
fof(f105,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X1),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(X0),X4))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X3),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X4)))))))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X1),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(X0),X4))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X3),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X4)))))))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X1),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(X0),X4))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X3),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X1))),s('type.closSem.state'(X0),X4)))))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X15] : ! [X9,X7,X10,X12] : s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X15)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X9),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.list.list'('type.closSem.v'),X10))),s('type.closSem.state'(X15),X12))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X15)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X15))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X15)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X10),s('type.closSem.state'(X15),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X9))),s('type.closSem.state'(X15),X12)))))))))),
  file('thm.clos_relationProps.exp_rel_thm','thm.clos_relation.evaluate_ev_def|split|2')).
fof(f150,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X10),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(sK0),X12))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',X14),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),X11))),s('type.closSem.state'(sK0),X13)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(sK0),X12),s('type.closSem.state'(sK0),X13)))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))) )),
  inference(forward_demodulation,[],[f96,f105])).
fof(f96,plain,(
  ( ! [X14,X12,X10,X13,X11,X9] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X10),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(sK0),X12))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X11),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(sK0),X13)))))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(sK0),X12),s('type.closSem.state'(sK0),X13)))) | p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) | ! [X9,X10,X11,X12,X13,X14] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(sK0),X12),s('type.closSem.state'(sK0),X13)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X10),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(sK0),X12))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X11),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(sK0),X13)))))))))))))) & (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2)))) | (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) & ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),sK4),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK8))),s('type.closSem.state'(sK0),sK6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),sK5),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK8))),s('type.closSem.state'(sK0),sK7))))))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f79])).
fof(f79,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | ! [X9,X10,X11,X12,X13,X14] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(X0),X12),s('type.closSem.state'(X0),X13)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X10),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(X0),X12))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X11),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(X0),X13)))))))))))))) & (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | ? [X3,X4,X5,X6,X7,X8] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X3)))) & ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X4),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X5),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X7)))))))))))))))),
  inference(rectify,[],[f78])).
fof(f78,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | ! [X3,X4,X5,X6,X7,X8] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X3)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X4),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X5),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X7)))))))))))))) & (~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) | ? [X3,X4,X5,X6,X7,X8] : (p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X3)))) & ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X4),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X5),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X7)))))))))))))))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) <~> ! [X3,X4,X5,X6,X7,X8] : (~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X3)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X4),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X5),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X7))))))))))))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) <~> ! [X3,X4,X5,X6,X7,X8] : ((~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X3))))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X4),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X5),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X7))))))))))))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) <=> ! [X3,X4,X5,X6,X7,X8] : ((p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X3))))) => p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X4),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X5),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X7))))))))))))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closLang.exp'),X2)))) <=> ! [X3,X4,X5,X6,X7,X8] : ((p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X3),s('type.closSem.state'(X0),X6),s('type.closSem.state'(X0),X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X3))),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X5)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X3))))) => p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X4),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X5),s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X8))),s('type.closSem.state'(X0),X7))))))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X7),s('type.list.list'('type.closLang.exp'),X8)))) <=> ! [X9,X10,X11,X12,X13,X14] : ((p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(X6),X12),s('type.closSem.state'(X6),X13)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9))))) => p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X6))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X10),s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(X6),X12))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X6))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X8),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X11),s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(X6),X13))))))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),X7),s('type.list.list'('type.closLang.exp'),X8)))) <=> ! [X9,X10,X11,X12,X13,X14] : ((p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',X9),s('type.closSem.state'(X6),X12),s('type.closSem.state'(X6),X13)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X9))),s('type.list.list'('type.closSem.v'),X10),s('type.list.list'('type.closSem.v'),X11)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X9))))) => p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X6))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X10),s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(X6),X12))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X6))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X8),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),X11),s('type.closSem.state'(X6),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X14))),s('type.closSem.state'(X6),X13))))))))))))))),
  file('thm.clos_relationProps.exp_rel_thm','thm.clos_relationProps.exp_rel_thm')).
fof(f149,plain,(
  p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) | $spl8),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f97,plain,(
  p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num',sK3),s('type.closSem.state'(sK0),sK6),s('type.closSem.state'(sK0),sK7)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))),
  inference(cnf_transformation,[],[f80])).
fof(f147,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) | $spl6),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f98,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num',sK3))),s('type.list.list'('type.closSem.v'),sK4),s('type.list.list'('type.closSem.v'),sK5)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))),
  inference(cnf_transformation,[],[f80])).
fof(f145,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) | $spl4),
  inference(cnf_transformation,[],[f145_D])).
fof(f145_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f99,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK8),s('type.num.num',sK3)))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))),
  inference(cnf_transformation,[],[f80])).
fof(f143,plain,(
  ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | $spl3),
  inference(cnf_transformation,[],[f143_D])).
fof(f143_D,plain,(
  ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f139,plain,(
  ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK6))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))),
  inference(forward_demodulation,[],[f138,f105])).
fof(f138,plain,(
  ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),sK4),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK8))),s('type.closSem.state'(sK0),sK6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.clos_relation.evaluate_ev_3'(s('type.num.num',sK8),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.list.list'('type.closSem.v'),sK5))),s('type.closSem.state'(sK0),sK7)))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))),
  inference(forward_demodulation,[],[f100,f105])).
fof(f100,plain,(
  ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK1),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),sK4),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK8))),s('type.closSem.state'(sK0),sK6))))))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0))),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),sK2),s('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.list.list'('type.closSem.v'),sK5),s('type.closSem.state'(sK0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK8))),s('type.closSem.state'(sK0),sK7)))))))))))) | ~p(s(bool,'const.clos_relation.exp_rel_3'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.list.list'('type.closLang.exp'),sK1),s('type.list.list'('type.closLang.exp'),sK2))))),
  inference(cnf_transformation,[],[f80])).
% SZS output end Proof for thm.clos_relationProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 250
% Passive clauses: 426
% Generated clauses: 2351
% Final active clauses: 145
% Final passive clauses: 45
% Input formulas: 40
% Initial clauses: 41
% 
% Selected by SInE selection: 16
% SInE iterations: 1
% 
% Duplicate literals: 3
% Fw subsumption resolutions: 85
% Fw demodulations: 66
% 
% Simple tautologies: 26
% Forward subsumptions: 971
% 
% Binary resolution: 1679
% Forward superposition: 80
% Equality resolution: 90
% 
% Split clauses: 793
% Split components: 1713
% SAT solver clauses: 815
% SAT solver unit clauses: 11
% SAT solver binary clauses: 14
% 
% Sat splits: 793
% Sat splitting refutations: 21
% 
% Memory used [KB]: 4605
% Time elapsed: 0.223 s
% ------------------------------
% ------------------------------
% Success in time 12.221 s
