Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.labProps
% (71088)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 164
% Passive clauses: 2326
% Generated clauses: 24568
% Final active clauses: 163
% Final passive clauses: 2005
% Input formulas: 187
% Initial clauses: 198
% 
% Duplicate literals: 757
% Trivial inequalities: 187
% Fw subsumption resolutions: 2006
% Fw demodulations: 2560
% 
% Simple tautologies: 206
% Equational tautologies: 27
% Forward subsumptions: 15285
% 
% Binary resolution: 3
% Forward superposition: 10734
% Backward superposition: 7698
% Self superposition: 262
% Equality factoring: 79
% Equality resolution: 15
% 
% Split clauses: 242
% Split components: 484
% SAT solver clauses: 652
% SAT solver unit clauses: 44
% SAT solver binary clauses: 139
% 
% Sat splits: 242
% 
% TWLsolver clauses: 652
% TWLsolver calls for satisfiability: 118
% 
% Memory used [KB]: 20212
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.labProps
% (71108)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 752
% Generated clauses: 13644
% Final active clauses: 129
% Final passive clauses: 591
% Input formulas: 187
% Initial clauses: 198
% 
% Duplicate literals: 469
% Trivial inequalities: 125
% Fw subsumption resolutions: 408
% Fw demodulations: 1131
% Bw demodulations: 28
% Condensations: 1195
% Global subsumptions: 4
% 
% Simple tautologies: 120
% Equational tautologies: 22
% Forward subsumptions: 7383
% 
% Binary resolution: 4
% Forward superposition: 5297
% Backward superposition: 3788
% Self superposition: 165
% Equality factoring: 60
% Equality resolution: 11
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 1316
% SAT solver unit clauses: 376
% SAT solver binary clauses: 29
% SAT solver learnt clauses: 8
% 
% Sat splits: 14
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1300
% TWLsolver calls for satisfiability: 272
% 
% Memory used [KB]: 13816
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.labProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.labProps
% SZS output start Proof for thm.labProps
fof(f48690,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f48615])).
fof(f48615,plain,(
  s('type.num.num',sK2) != s('type.num.num',sK2)),
  inference(superposition,[],[f566,f48548])).
fof(f48548,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',X2) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),'const.labSem.upd_pc_2'(s('type.num.num',X2),s('type.labSem.state'(X0,X1),X3)))))) )),
  inference(forward_demodulation,[],[f48473,f766])).
fof(f766,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f592,f590])).
fof(f590,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : ! [X8,X11] : s(X7,X8) = s(X7,'const.combin.K_2'(s(X7,X8),s(X6,X11)))),
  file('thm.labProps.update_simps|split|2','thm.combin.K_THM')).
fof(f592,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f188])).
fof(f188,axiom,(
  ! [X7,X6] : ! [X16,X17] : s(X7,'const.combin.K_2'(s(X7,X16),s(X6,X17))) = s(X7,happ(s(fun(X6,X7),'const.combin.K_1'(s(X7,X16))),s(X6,X17)))),
  file('thm.labProps.update_simps|split|2','ARITY_1_2_const.combin.K')).
fof(f48473,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),'const.labSem.upd_pc_2'(s('type.num.num',X2),s('type.labSem.state'(X0,X1),X3))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),X3)))))) )),
  inference(superposition,[],[f691,f656])).
fof(f656,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.labSem.state'(X0,X1),'const.labSem.upd_pc_2'(s('type.num.num',X2),s('type.labSem.state'(X0,X1),X3))) = s('type.labSem.state'(X0,X1),'const.labSem.state_pc_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.labSem.state'(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0,X1,X2,X3] : s('type.labSem.state'(X0,X1),'const.labSem.upd_pc_2'(s('type.num.num',X2),s('type.labSem.state'(X0,X1),X3))) = s('type.labSem.state'(X0,X1),'const.labSem.state_pc_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.labSem.state'(X0,X1),X3)))),
  inference(flattening,[],[f335])).
fof(f335,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.labSem.state'(X0,X1),'const.labSem.upd_pc_2'(s('type.num.num',X2),s('type.labSem.state'(X0,X1),X3))) = s('type.labSem.state'(X0,X1),'const.labSem.state_pc_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.labSem.state'(X0,X1),X3)))),
  inference(rectify,[],[f185])).
fof(f185,axiom,(
  ! [X7,X6] : ! [X15,X9] : s('type.labSem.state'(X7,X6),'const.labSem.upd_pc_2'(s('type.num.num',X15),s('type.labSem.state'(X7,X6),X9))) = s('type.labSem.state'(X7,X6),'const.labSem.state_pc_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X15))),s('type.labSem.state'(X7,X6),X9)))),
  file('thm.labProps.update_simps|split|2','thm.labSem.upd_pc_def')).
fof(f691,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X3),s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),X2))))) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),'const.labSem.state_pc_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.labSem.state'(X0,X1),X2)))))) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X3),s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),X2))))) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),'const.labSem.state_pc_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.labSem.state'(X0,X1),X2)))))),
  inference(flattening,[],[f405])).
fof(f405,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X3),s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),X2))))) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),'const.labSem.state_pc_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.labSem.state'(X0,X1),X2)))))),
  inference(rectify,[],[f175])).
fof(f175,axiom,(
  ! [X7,X12] : ! [X9,X13] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X13),s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X7,X12),X9))))) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X7,X12),'const.labSem.state_pc_fupd_2'(s(fun('type.num.num','type.num.num'),X13),s('type.labSem.state'(X7,X12),X9)))))),
  file('thm.labProps.update_simps|split|2','thm.labSem.labSem_2|48|8_prove_ltb')).
fof(f566,plain,(
  s('type.num.num',sK2) != s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(sK0,sK1),'const.labSem.upd_pc_2'(s('type.num.num',sK2),s('type.labSem.state'(sK0,sK1),sK3)))))),
  inference(cnf_transformation,[],[f553])).
fof(f553,plain,(
  s('type.num.num',sK2) != s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(sK0,sK1),'const.labSem.upd_pc_2'(s('type.num.num',sK2),s('type.labSem.state'(sK0,sK1),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f551])).
fof(f551,plain,(
  ? [X0,X1,X2,X3] : s('type.num.num',X2) != s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),'const.labSem.upd_pc_2'(s('type.num.num',X2),s('type.labSem.state'(X0,X1),X3)))))),
  inference(ennf_transformation,[],[f190])).
fof(f190,plain,(
  ~! [X0,X1,X2,X3] : s('type.num.num',X2) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),'const.labSem.upd_pc_2'(s('type.num.num',X2),s('type.labSem.state'(X0,X1),X3)))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ~! [X0,X1] : ! [X2,X3] : s('type.num.num',X2) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X0,X1),'const.labSem.upd_pc_2'(s('type.num.num',X2),s('type.labSem.state'(X0,X1),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : s('type.num.num',X8) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X6,X7),'const.labSem.upd_pc_2'(s('type.num.num',X8),s('type.labSem.state'(X6,X7),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : s('type.num.num',X8) = s('type.num.num','const.labSem.state_pc_1'(s('type.labSem.state'(X6,X7),'const.labSem.upd_pc_2'(s('type.num.num',X8),s('type.labSem.state'(X6,X7),X9)))))),
  file('thm.labProps.update_simps|split|2','thm.labProps.update_simps|split|2')).
% SZS output end Proof for thm.labProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 231
% Passive clauses: 1938
% Generated clauses: 47982
% Final active clauses: 213
% Final passive clauses: 1534
% Input formulas: 187
% Initial clauses: 198
% 
% Duplicate literals: 1592
% Trivial inequalities: 327
% Fw subsumption resolutions: 3395
% Fw demodulations: 5160
% Bw demodulations: 165
% Condensations: 5774
% 
% Simple tautologies: 314
% Equational tautologies: 116
% Forward subsumptions: 29186
% Bw demodulations to eq. taut.: 24
% 
% Unit resulting resolution: 687
% Forward superposition: 15884
% Backward superposition: 14191
% Self superposition: 339
% Equality factoring: 109
% Equality resolution: 21
% 
% Split clauses: 114
% Split components: 228
% SAT solver clauses: 2169
% SAT solver unit clauses: 219
% SAT solver binary clauses: 435
% 
% Sat splits: 114
% Sat splitting refutations: 1
% 
% Memory used [KB]: 28528
% Time elapsed: 3.942 s
% ------------------------------
% ------------------------------
% Success in time 5.34 s
