Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closSem
% (124861)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 101
% Passive clauses: 1668
% Generated clauses: 10664
% Final active clauses: 78
% Final passive clauses: 1417
% Input formulas: 37
% Initial clauses: 40
% 
% Duplicate literals: 429
% Trivial inequalities: 46
% Fw subsumption resolutions: 383
% Fw demodulations: 1017
% Bw demodulations: 20
% 
% Simple tautologies: 57
% Equational tautologies: 20
% Forward subsumptions: 6842
% 
% Binary resolution: 88
% Forward superposition: 4898
% Backward superposition: 3660
% Self superposition: 234
% Equality factoring: 56
% Equality resolution: 4
% 
% Split clauses: 193
% Split components: 386
% SAT solver clauses: 304
% SAT solver unit clauses: 3
% SAT solver binary clauses: 66
% 
% Sat splits: 193
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 304
% TWLsolver calls for satisfiability: 80
% 
% Memory used [KB]: 18677
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closSem
% (124878)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 77
% Passive clauses: 1164
% Generated clauses: 11419
% Final active clauses: 76
% Final passive clauses: 1058
% Input formulas: 37
% Initial clauses: 40
% 
% Duplicate literals: 498
% Trivial inequalities: 78
% Fw subsumption resolutions: 337
% Fw demodulations: 998
% Bw demodulations: 27
% Condensations: 895
% 
% Simple tautologies: 99
% Equational tautologies: 22
% Forward subsumptions: 5676
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 42
% Forward superposition: 3491
% Backward superposition: 3557
% Self superposition: 222
% Equality factoring: 60
% Equality resolution: 7
% 
% Split clauses: 64
% Split components: 128
% SAT solver clauses: 1662
% SAT solver unit clauses: 68
% SAT solver binary clauses: 59
% 
% Sat splits: 64
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1662
% TWLsolver calls for satisfiability: 136
% 
% Memory used [KB]: 23922
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.closSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closSem
% SZS output start Proof for thm.closSem
fof(f10839,plain,(
  $false),
  inference(subsumption_resolution,[],[f10838,f116])).
fof(f116,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.closSem.state_fupdfupds_comp|split|1',aHL_TRUTH)).
fof(f10838,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f10836,f644])).
fof(f644,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(equality_resolution,[],[f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(condensation,[],[f436])).
fof(f436,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f433,f159])).
fof(f159,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f117,f118])).
fof(f118,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.closSem.state_fupdfupds_comp|split|1',aHL_BOOL_CASES)).
fof(f117,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.closSem.state_fupdfupds_comp|split|1',aHL_FALSITY)).
fof(f433,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f134,f157])).
fof(f157,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f118,f118])).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X35,X36,X37] : (p(s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),'=_0'),s(X35,X36))),s(X35,X37)))) <=> s(X35,X36) = s(X35,X37))),
  file('thm.closSem.state_fupdfupds_comp|split|1',arity_0_HL_EQ)).
fof(f10836,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.closSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.closSem.state'(sK1)),fun(fun(sK0,'type.closSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK3))))),s(fun(sK0,'type.closSem.state'(sK1)),sK2))))),s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK3))))),s(fun(sK0,'type.closSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f10784,f2168])).
fof(f2168,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.closSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.closSem.state'(sK1)),fun(fun(sK0,'type.closSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK4))),s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK3))),s(fun(sK0,'type.closSem.state'(sK1)),sK2))))))),s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK3))))),s(fun(sK0,'type.closSem.state'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f115,f133])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f114])).
fof(f115,plain,(
  s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK3))))),s(fun(sK0,'type.closSem.state'(sK1)),sK2))) != s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK4))),s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK3))),s(fun(sK0,'type.closSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK4),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK3))))),s(fun(sK0,'type.closSem.state'(sK1)),sK2))) != s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK4))),s(fun(sK0,'type.closSem.state'(sK1)),'const.combin.o_2'(s(fun('type.closSem.state'(sK1),'type.closSem.state'(sK1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),sK3))),s(fun(sK0,'type.closSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f105])).
fof(f105,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X4),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X3))))),s(fun(X0,'type.closSem.state'(X1)),X2))) != s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X4))),s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X3))),s(fun(X0,'type.closSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X4),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X3))))),s(fun(X0,'type.closSem.state'(X1)),X2))) = s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X4))),s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X3))),s(fun(X0,'type.closSem.state'(X1)),X2)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X4),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X3))))),s(fun(X0,'type.closSem.state'(X1)),X2))) = s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X4))),s(fun(X0,'type.closSem.state'(X1)),'const.combin.o_2'(s(fun('type.closSem.state'(X1),'type.closSem.state'(X1)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X3))),s(fun(X0,'type.closSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.closSem.state'(X7)),'const.combin.o_2'(s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X10),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X9))))),s(fun(X6,'type.closSem.state'(X7)),X8))) = s(fun(X6,'type.closSem.state'(X7)),'const.combin.o_2'(s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X10))),s(fun(X6,'type.closSem.state'(X7)),'const.combin.o_2'(s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X9))),s(fun(X6,'type.closSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.closSem.state'(X7)),'const.combin.o_2'(s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X10),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X9))))),s(fun(X6,'type.closSem.state'(X7)),X8))) = s(fun(X6,'type.closSem.state'(X7)),'const.combin.o_2'(s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X10))),s(fun(X6,'type.closSem.state'(X7)),'const.combin.o_2'(s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X9))),s(fun(X6,'type.closSem.state'(X7)),X8)))))),
  file('thm.closSem.state_fupdfupds_comp|split|1','thm.closSem.state_fupdfupds_comp|split|1')).
fof(f10784,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.closSem.state'(X0)),'const.combin.o_2'(s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X1),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X2))))),s(fun(X3,'type.closSem.state'(X0)),X4))) = s(fun(X3,'type.closSem.state'(X0)),'const.combin.o_2'(s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X1))),s(fun(X3,'type.closSem.state'(X0)),'const.combin.o_2'(s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X2))),s(fun(X3,'type.closSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f140,f120])).
fof(f120,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X2),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X1))))) = s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.combin.o_2'(s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X2))),s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X1)))))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X2),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X1))))) = s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.combin.o_2'(s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X2))),s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X1)))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X2),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X1))))) = s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.combin.o_2'(s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X2))),s(fun('type.closSem.state'(X0),'type.closSem.state'(X0)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X10),s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X9))))) = s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.combin.o_2'(s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X10))),s(fun('type.closSem.state'(X7),'type.closSem.state'(X7)),'const.closSem.state_globals_fupd_1'(s(fun('type.list.list'('type.option.option'('type.closSem.v')),'type.list.list'('type.option.option'('type.closSem.v'))),X9)))))),
  file('thm.closSem.state_fupdfupds_comp|split|1','thm.closSem.state_fupdfupds_comp|split|0')).
fof(f140,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X14,X6,X18,X16] : ! [X10,X9,X8] : s(fun(X18,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X18,X6),'const.combin.o_2'(s(fun(X16,X6),X9),s(fun(X18,X16),X8))))) = s(fun(X18,X14),'const.combin.o_2'(s(fun(X16,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X16,X6),X9))),s(fun(X18,X16),X8)))),
  file('thm.closSem.state_fupdfupds_comp|split|1','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.closSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 98
% Passive clauses: 1628
% Generated clauses: 11598
% Final active clauses: 57
% Final passive clauses: 678
% Input formulas: 37
% Initial clauses: 40
% 
% Duplicate literals: 440
% Trivial inequalities: 38
% Fw subsumption resolutions: 325
% Fw demodulations: 1868
% Bw demodulations: 697
% Condensations: 458
% Global subsumptions: 3
% 
% Simple tautologies: 98
% Equational tautologies: 87
% Forward subsumptions: 6506
% Fw demodulations to eq. taut.: 7
% Bw demodulations to eq. taut.: 28
% 
% Unit resulting resolution: 434
% Forward superposition: 3066
% Backward superposition: 2941
% Self superposition: 183
% Equality factoring: 54
% Equality resolution: 1
% 
% Split clauses: 33
% Split components: 66
% SAT solver clauses: 1763
% SAT solver unit clauses: 325
% SAT solver binary clauses: 644
% 
% Sat splits: 33
% Sat splitting refutations: 64
% 
% Memory used [KB]: 15351
% Time elapsed: 0.692 s
% ------------------------------
% ------------------------------
% Success in time 2.041 s
