Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.bigClock
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bigClock
% SZS output start Proof for thm.bigClock
fof(f815,plain,(
  $false),
  inference(subsumption_resolution,[],[f814,f48])).
fof(f48,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bigClock.bigClock_|49||49|_prove_ltb',aHL_TRUTH)).
fof(f814,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f813,f70])).
fof(f70,plain,(
  s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0')))),
  inference(resolution,[],[f55,f50])).
fof(f50,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'))))),
  inference(cnf_transformation,[],[f13])).
fof(f13,axiom,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'))))),
  file('thm.bigClock.bigClock_|49||49|_prove_ltb','thm.prim_rec.WF_LESS')).
fof(f55,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : (s(bool,'T_0') = s(bool,X6) <=> p(s(bool,X6)))),
  file('thm.bigClock.bigClock_|49||49|_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f813,plain,(
  ~p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'))))),
  inference(subsumption_resolution,[],[f812,f48])).
fof(f812,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'))))),
  inference(forward_demodulation,[],[f804,f105])).
fof(f105,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.prim_rec.measure_1'(s(fun(X0,'type.num.num'),X1)))))) )),
  inference(resolution,[],[f62,f55])).
fof(f62,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.prim_rec.measure_1'(s(fun(X0,'type.num.num'),X1))))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.prim_rec.measure_1'(s(fun(X0,'type.num.num'),X1))))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.prim_rec.measure_1'(s(fun(X0,'type.num.num'),X1))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : ! [X11] : p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),'const.prim_rec.measure_1'(s(fun(X7,'type.num.num'),X11))))))),
  file('thm.bigClock.bigClock_|49||49|_prove_ltb','thm.prim_rec.WF_measure')).
fof(f804,plain,(
  ~p(s(bool,'const.relation.WF_1'(s(fun('type.ast.exp',fun('type.ast.exp',bool)),'const.prim_rec.measure_1'(s(fun('type.ast.exp','type.num.num'),'const.ast.exp_size_0')))))) | ~p(s(bool,'const.relation.WF_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'))))),
  inference(resolution,[],[f65,f47])).
fof(f47,plain,(
  ~p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'('type.num.num','type.ast.exp'),fun('type.pair.prod'('type.num.num','type.ast.exp'),bool)),'const.pair.LEX_2'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.ast.exp',fun('type.ast.exp',bool)),'const.prim_rec.measure_1'(s(fun('type.ast.exp','type.num.num'),'const.ast.exp_size_0'))))))))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'('type.num.num','type.ast.exp'),fun('type.pair.prod'('type.num.num','type.ast.exp'),bool)),'const.pair.LEX_2'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.ast.exp',fun('type.ast.exp',bool)),'const.prim_rec.measure_1'(s(fun('type.ast.exp','type.num.num'),'const.ast.exp_size_0'))))))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'('type.num.num','type.ast.exp'),fun('type.pair.prod'('type.num.num','type.ast.exp'),bool)),'const.pair.LEX_2'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.ast.exp',fun('type.ast.exp',bool)),'const.prim_rec.measure_1'(s(fun('type.ast.exp','type.num.num'),'const.ast.exp_size_0'))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'('type.num.num','type.ast.exp'),fun('type.pair.prod'('type.num.num','type.ast.exp'),bool)),'const.pair.LEX_2'(s(fun('type.num.num',fun('type.num.num',bool)),'const.prim_rec.<_0'),s(fun('type.ast.exp',fun('type.ast.exp',bool)),'const.prim_rec.measure_1'(s(fun('type.ast.exp','type.num.num'),'const.ast.exp_size_0'))))))))),
  file('thm.bigClock.bigClock_|49||49|_prove_ltb','thm.bigClock.bigClock_|49||49|_prove_ltb')).
fof(f65,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'const.pair.LEX_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,fun(X1,bool)),X3)))))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'const.pair.LEX_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,fun(X1,bool)),X3)))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3))))) | p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'const.pair.LEX_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,fun(X1,bool)),X3)))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3))))) => p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'const.pair.LEX_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,fun(X1,bool)),X3)))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3] : ((p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3))))) => p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'const.pair.LEX_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,fun(X1,bool)),X3)))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7,X8] : ! [X9,X10] : ((p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X9)))) & p(s(bool,'const.relation.WF_1'(s(fun(X8,fun(X8,bool)),X10))))) => p(s(bool,'const.relation.WF_1'(s(fun('type.pair.prod'(X7,X8),fun('type.pair.prod'(X7,X8),bool)),'const.pair.LEX_2'(s(fun(X7,fun(X7,bool)),X9),s(fun(X8,fun(X8,bool)),X10)))))))),
  file('thm.bigClock.bigClock_|49||49|_prove_ltb','thm.pair.WF_LEX')).
% SZS output end Proof for thm.bigClock
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 36
% Passive clauses: 164
% Generated clauses: 741
% Final active clauses: 27
% Final passive clauses: 102
% Input formulas: 17
% Initial clauses: 21
% 
% Duplicate literals: 75
% Trivial inequalities: 20
% Fw subsumption resolutions: 59
% Fw demodulations: 36
% Bw demodulations: 8
% 
% Simple tautologies: 52
% Equational tautologies: 16
% Forward subsumptions: 320
% 
% Binary resolution: 25
% Forward superposition: 238
% Backward superposition: 190
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 8
% 
% Split clauses: 8
% Split components: 21
% SAT solver clauses: 13
% SAT solver unit clauses: 2
% SAT solver binary clauses: 4
% 
% Sat splits: 8
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 13
% TWLsolver calls for satisfiability: 5
% 
% Memory used [KB]: 1023
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
% Success in time 0.03 s
