Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relation
% (44443)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 118
% Passive clauses: 2949
% Generated clauses: 14432
% Final active clauses: 109
% Final passive clauses: 2706
% Input formulas: 37
% Initial clauses: 69
% 
% Duplicate literals: 494
% Trivial inequalities: 215
% Fw subsumption resolutions: 1181
% Fw demodulations: 1650
% Bw demodulations: 7
% 
% Simple tautologies: 260
% Equational tautologies: 41
% Forward subsumptions: 7401
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 216
% Factoring: 1
% Forward superposition: 5474
% Backward superposition: 4778
% Self superposition: 190
% Equality factoring: 84
% Equality resolution: 23
% 
% Split clauses: 137
% Split components: 280
% SAT solver clauses: 324
% SAT solver unit clauses: 17
% SAT solver binary clauses: 85
% 
% Sat splits: 137
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 323
% TWLsolver calls for satisfiability: 92
% 
% Memory used [KB]: 26353
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_relation
% (44583)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 92
% Passive clauses: 1152
% Generated clauses: 8722
% Final active clauses: 88
% Final passive clauses: 1060
% Input formulas: 37
% Initial clauses: 69
% 
% Duplicate literals: 119
% Trivial inequalities: 95
% Fw subsumption resolutions: 387
% Fw demodulations: 834
% Bw demodulations: 4
% Condensations: 257
% Global subsumptions: 7
% 
% Simple tautologies: 147
% Equational tautologies: 9
% Forward subsumptions: 4213
% 
% Binary resolution: 43
% Factoring: 1
% Forward superposition: 3447
% Backward superposition: 2131
% Self superposition: 127
% Equality factoring: 20
% Equality resolution: 27
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 1534
% SAT solver unit clauses: 101
% SAT solver binary clauses: 39
% SAT solver learnt clauses: 14
% 
% Sat splits: 23
% 
% TWLsolver clauses: 1506
% TWLsolver calls for satisfiability: 558
% 
% Memory used [KB]: 22515
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_relation
% (44620)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 504
% Passive clauses: 13072
% Generated clauses: 110516
% Final active clauses: 126
% Final passive clauses: 1033
% Input formulas: 37
% Initial clauses: 69
% Discarded non-redundant clauses: 4621
% 
% Duplicate literals: 3650
% Trivial inequalities: 773
% Fw subsumption resolutions: 7193
% Fw demodulations: 18751
% Bw demodulations: 5198
% Condensations: 6118
% Global subsumptions: 91
% 
% Simple tautologies: 798
% Equational tautologies: 319
% Forward subsumptions: 59411
% Fw demodulations to eq. taut.: 20
% Bw demodulations to eq. taut.: 556
% 
% Unit resulting resolution: 1950
% Factoring: 19
% Forward superposition: 38105
% Backward superposition: 23118
% Self superposition: 558
% Equality factoring: 210
% Equality resolution: 49
% 
% Split clauses: 77
% Split components: 155
% SAT solver clauses: 13524
% SAT solver unit clauses: 349
% SAT solver binary clauses: 3518
% 
% Sat splits: 77
% Sat splitting refutations: 205
% 
% Memory used [KB]: 93644
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.clos_relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relation
% SZS output start Proof for thm.clos_relation
fof(f1659,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1369,f1373,f1413,f1551,f1553,f1371,f1594])).
fof(f1594,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f1557,f156])).
fof(f156,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp',sK1)))),
  inference(resolution,[],[f102,f143])).
fof(f143,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X27] : (s(bool,'T_0') = s(bool,X27) <=> p(s(bool,X27)))),
  file('thm.clos_relation.val_or_exp_induction','thm.bool.EQ_CLAUSES|split|0')).
fof(f102,plain,(
  ~p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp',sK1))))),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X9,X10] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X9),s('type.num.num',X10)))))) & ! [X4,X5,X6,X7,X8] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X4),s('type.closLang.exp',X5),s('type.list.list'('type.closSem.v'),X6),s('type.list.list'('type.closSem.v'),X7),s('type.num.num',X8)))))) & ! [X2,X3] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3)))))) & ~p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp',sK1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f79])).
fof(f79,plain,(
  ? [X0] : (! [X9,X10] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X9),s('type.num.num',X10)))))) & ! [X4,X5,X6,X7,X8] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X4),s('type.closLang.exp',X5),s('type.list.list'('type.closSem.v'),X6),s('type.list.list'('type.closSem.v'),X7),s('type.num.num',X8)))))) & ! [X2,X3] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3)))))) & ? [X1] : ~p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp',X1)))))),
  inference(rectify,[],[f73])).
fof(f73,plain,(
  ? [X0] : (! [X8,X9] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X8),s('type.num.num',X9)))))) & ! [X3,X4,X5,X6,X7] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X3),s('type.closLang.exp',X4),s('type.list.list'('type.closSem.v'),X5),s('type.list.list'('type.closSem.v'),X6),s('type.num.num',X7)))))) & ! [X1,X2] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X2)))))) & ? [X10] : ~p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp',X10)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ? [X0] : ((! [X8,X9] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X8),s('type.num.num',X9)))))) & ! [X3,X4,X5,X6,X7] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X3),s('type.closLang.exp',X4),s('type.list.list'('type.closSem.v'),X5),s('type.list.list'('type.closSem.v'),X6),s('type.num.num',X7)))))) & ! [X1,X2] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X2))))))) & ? [X10] : ~p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp',X10)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ((! [X8,X9] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X8),s('type.num.num',X9)))))) & ! [X3,X4,X5,X6,X7] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X3),s('type.closLang.exp',X4),s('type.list.list'('type.closSem.v'),X5),s('type.list.list'('type.closSem.v'),X6),s('type.num.num',X7)))))) & ! [X1,X2] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X2))))))) => ! [X10] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X0),s('type.clos_relation.val_or_exp',X10)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ((! [X7,X8] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X6),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X7),s('type.num.num',X8)))))) & ! [X9,X10,X11,X12,X8] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X6),s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X9),s('type.closLang.exp',X10),s('type.list.list'('type.closSem.v'),X11),s('type.list.list'('type.closSem.v'),X12),s('type.num.num',X8)))))) & ! [X11,X12] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X6),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),X12))))))) => ! [X7] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X6),s('type.clos_relation.val_or_exp',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ((! [X7,X8] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X6),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X7),s('type.num.num',X8)))))) & ! [X9,X10,X11,X12,X8] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X6),s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X9),s('type.closLang.exp',X10),s('type.list.list'('type.closSem.v'),X11),s('type.list.list'('type.closSem.v'),X12),s('type.num.num',X8)))))) & ! [X11,X12] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X6),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),X12))))))) => ! [X7] : p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),X6),s('type.clos_relation.val_or_exp',X7)))))),
  file('thm.clos_relation.val_or_exp_induction','thm.clos_relation.val_or_exp_induction')).
fof(f1557,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp',sK1))) | $spl0),
  inference(superposition,[],[f182,f1371])).
fof(f182,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X10),s('type.num.num',X11)))))) )),
  inference(resolution,[],[f99,f144])).
fof(f144,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f96])).
fof(f99,plain,(
  ( ! [X10,X9] : (p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X9),s('type.num.num',X10))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f1371,plain,(
  s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(sK1)),s('type.num.num',sK13(sK1)))) | $spl0),
  inference(cnf_transformation,[],[f1371_D])).
fof(f1371_D,plain,(
  s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(sK1)),s('type.num.num',sK13(sK1)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1553,plain,(
  s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),sK14(sK1)),s('type.closLang.exp',sK15(sK1)),s('type.list.list'('type.closSem.v'),sK16(sK1)),s('type.list.list'('type.closSem.v'),sK17(sK1)),s('type.num.num',sK18(sK1)))) | $spl4),
  inference(cnf_transformation,[],[f1553_D])).
fof(f1553_D,plain,(
  s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),sK14(sK1)),s('type.closLang.exp',sK15(sK1)),s('type.list.list'('type.closSem.v'),sK16(sK1)),s('type.list.list'('type.closSem.v'),sK17(sK1)),s('type.num.num',sK18(sK1)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1551,plain,(
  s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),sK14(sK1)),s('type.closLang.exp',sK15(sK1)),s('type.list.list'('type.closSem.v'),sK16(sK1)),s('type.list.list'('type.closSem.v'),sK17(sK1)),s('type.num.num',sK18(sK1)))) | s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(sK1)),s('type.num.num',sK13(sK1))))),
  inference(subsumption_resolution,[],[f1526,f147])).
fof(f147,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_relation.val_or_exp_induction',aHL_FALSITY)).
fof(f1526,plain,(
  p(s(bool,'F_0')) | s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),sK14(sK1)),s('type.closLang.exp',sK15(sK1)),s('type.list.list'('type.closSem.v'),sK16(sK1)),s('type.list.list'('type.closSem.v'),sK17(sK1)),s('type.num.num',sK18(sK1)))) | s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(sK1)),s('type.num.num',sK13(sK1))))),
  inference(superposition,[],[f400,f157])).
fof(f157,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp',sK1)))),
  inference(resolution,[],[f102,f142])).
fof(f142,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X27] : (s(bool,'F_0') = s(bool,X27) <=> ~p(s(bool,X27)))),
  file('thm.clos_relation.val_or_exp_induction','thm.bool.EQ_CLAUSES|split|2')).
fof(f400,plain,(
  ( ! [X36] : (p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp',X36)))) | s('type.clos_relation.val_or_exp',X36) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),sK14(X36)),s('type.closLang.exp',sK15(X36)),s('type.list.list'('type.closSem.v'),sK16(X36)),s('type.list.list'('type.closSem.v'),sK17(X36)),s('type.num.num',sK18(X36)))) | s('type.clos_relation.val_or_exp',X36) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(X36)),s('type.num.num',sK13(X36))))) )),
  inference(superposition,[],[f101,f115])).
fof(f115,plain,(
  ( ! [X0] : (s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK19(X0)),s('type.list.list'('type.closSem.v'),sK20(X0)))) | s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),sK14(X0)),s('type.closLang.exp',sK15(X0)),s('type.list.list'('type.closSem.v'),sK16(X0)),s('type.list.list'('type.closSem.v'),sK17(X0)),s('type.num.num',sK18(X0)))) | s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(X0)),s('type.num.num',sK13(X0))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(X0)),s('type.num.num',sK13(X0)))) | s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),sK14(X0)),s('type.closLang.exp',sK15(X0)),s('type.list.list'('type.closSem.v'),sK16(X0)),s('type.list.list'('type.closSem.v'),sK17(X0)),s('type.num.num',sK18(X0)))) | s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK19(X0)),s('type.list.list'('type.closSem.v'),sK20(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17,sK18,sK19,sK20])],[f46])).
fof(f46,plain,(
  ! [X0] : (? [X8,X9] : s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X8),s('type.num.num',X9))) | ? [X3,X4,X5,X6,X7] : s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X3),s('type.closLang.exp',X4),s('type.list.list'('type.closSem.v'),X5),s('type.list.list'('type.closSem.v'),X6),s('type.num.num',X7))) | ? [X1,X2] : s('type.clos_relation.val_or_exp',X0) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.list.list'('type.closSem.v'),X2))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X38] : (? [X7,X8] : s('type.clos_relation.val_or_exp',X38) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',X7),s('type.num.num',X8))) | ? [X9,X10,X11,X12,X8] : s('type.clos_relation.val_or_exp',X38) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X9),s('type.closLang.exp',X10),s('type.list.list'('type.closSem.v'),X11),s('type.list.list'('type.closSem.v'),X12),s('type.num.num',X8))) | ? [X11,X12] : s('type.clos_relation.val_or_exp',X38) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X11),s('type.list.list'('type.closSem.v'),X12))))),
  file('thm.clos_relation.val_or_exp_induction','thm.clos_relation.val_or_exp_nchotomy')).
fof(f101,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.list.list'('type.closSem.v'),X3))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f1413,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1377,f156])).
fof(f1377,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp',sK1))) | $spl2),
  inference(superposition,[],[f342,f1373])).
fof(f342,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),X10),s('type.list.list'('type.closSem.v'),X11)))))) )),
  inference(resolution,[],[f101,f144])).
fof(f1373,plain,(
  s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK19(sK1)),s('type.list.list'('type.closSem.v'),sK20(sK1)))) | $spl2),
  inference(cnf_transformation,[],[f1373_D])).
fof(f1373_D,plain,(
  s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK19(sK1)),s('type.list.list'('type.closSem.v'),sK20(sK1)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1369,plain,(
  s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK19(sK1)),s('type.list.list'('type.closSem.v'),sK20(sK1)))) | s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(sK1)),s('type.num.num',sK13(sK1))))),
  inference(subsumption_resolution,[],[f1349,f147])).
fof(f1349,plain,(
  p(s(bool,'F_0')) | s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK19(sK1)),s('type.list.list'('type.closSem.v'),sK20(sK1)))) | s('type.clos_relation.val_or_exp',sK1) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(sK1)),s('type.num.num',sK13(sK1))))),
  inference(superposition,[],[f559,f157])).
fof(f559,plain,(
  ( ! [X36] : (p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp',X36)))) | s('type.clos_relation.val_or_exp',X36) = s('type.clos_relation.val_or_exp','const.clos_relation.Exp_2'(s('type.list.list'('type.closLang.exp'),sK19(X36)),s('type.list.list'('type.closSem.v'),sK20(X36)))) | s('type.clos_relation.val_or_exp',X36) = s('type.clos_relation.val_or_exp','const.clos_relation.Val_2'(s('type.closSem.v',sK12(X36)),s('type.num.num',sK13(X36))))) )),
  inference(superposition,[],[f100,f115])).
fof(f100,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,happ(s(fun('type.clos_relation.val_or_exp',bool),sK0),s('type.clos_relation.val_or_exp','const.clos_relation.Exp|49|_5'(s('type.option.option'('type.num.num'),X4),s('type.closLang.exp',X5),s('type.list.list'('type.closSem.v'),X6),s('type.list.list'('type.closSem.v'),X7),s('type.num.num',X8))))))) )),
  inference(cnf_transformation,[],[f80])).
% SZS output end Proof for thm.clos_relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 53
% Passive clauses: 174
% Generated clauses: 1503
% Final active clauses: 45
% Final passive clauses: 135
% Input formulas: 37
% Initial clauses: 55
% 
% Selected by SInE selection: 24
% SInE iterations: 1
% 
% Duplicate literals: 14
% Trivial inequalities: 18
% Fw subsumption resolutions: 102
% Fw demodulations: 177
% Bw demodulations: 7
% 
% Simple tautologies: 35
% Equational tautologies: 1
% Forward subsumptions: 946
% 
% Binary resolution: 46
% Forward superposition: 965
% Backward superposition: 169
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 5
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Sat splits: 2
% Sat splitting refutations: 2
% 
% Memory used [KB]: 4733
% Time elapsed: 0.099 s
% ------------------------------
% ------------------------------
% Success in time 11.85 s
