Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_removeProof
% (46167)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 237
% Passive clauses: 3065
% Generated clauses: 31340
% Final active clauses: 225
% Final passive clauses: 2545
% Input formulas: 21
% Initial clauses: 44
% 
% Duplicate literals: 933
% Trivial inequalities: 1283
% Fw subsumption resolutions: 4036
% Fw demodulations: 7357
% Bw demodulations: 6
% 
% Simple tautologies: 547
% Equational tautologies: 29
% Forward subsumptions: 12497
% 
% Binary resolution: 457
% Forward superposition: 9880
% Backward superposition: 6999
% Self superposition: 86
% Equality factoring: 55
% Equality resolution: 66
% 
% Split clauses: 576
% Split components: 1182
% SAT solver clauses: 1412
% SAT solver unit clauses: 21
% SAT solver binary clauses: 219
% 
% Sat splits: 576
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1407
% TWLsolver calls for satisfiability: 205
% 
% Memory used [KB]: 11513
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_removeProof
% (46185)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 163
% Passive clauses: 870
% Generated clauses: 19292
% Final active clauses: 158
% Final passive clauses: 657
% Input formulas: 21
% Initial clauses: 44
% 
% Duplicate literals: 714
% Trivial inequalities: 444
% Fw subsumption resolutions: 1727
% Fw demodulations: 3316
% Bw demodulations: 54
% Condensations: 1146
% Global subsumptions: 1
% 
% Simple tautologies: 381
% Equational tautologies: 35
% Forward subsumptions: 9029
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 280
% Forward superposition: 5139
% Backward superposition: 5365
% Self superposition: 79
% Equality factoring: 56
% Equality resolution: 39
% 
% Split clauses: 127
% Split components: 271
% SAT solver clauses: 1568
% SAT solver unit clauses: 43
% SAT solver binary clauses: 45
% SAT solver learnt clauses: 2
% 
% Sat splits: 127
% 
% TWLsolver clauses: 1564
% TWLsolver calls for satisfiability: 304
% 
% Memory used [KB]: 11257
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_removeProof
% (46206)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 589
% Passive clauses: 5841
% Generated clauses: 463593
% Final active clauses: 40
% Final passive clauses: 217
% Input formulas: 21
% Initial clauses: 44
% 
% Duplicate literals: 4300
% Trivial inequalities: 2561
% Fw subsumption resolutions: 12257
% Fw demodulations: 234185
% Bw demodulations: 4336
% Condensations: 9393
% Global subsumptions: 1081
% 
% Simple tautologies: 1513
% Equational tautologies: 9329
% Forward subsumptions: 180634
% Fw demodulations to eq. taut.: 105
% Bw demodulations to eq. taut.: 94
% 
% Unit resulting resolution: 122714
% Forward superposition: 33547
% Backward superposition: 23368
% Self superposition: 289
% Equality factoring: 164
% Equality resolution: 117
% 
% Split clauses: 246
% Split components: 529
% SAT solver clauses: 9149
% SAT solver unit clauses: 696
% SAT solver binary clauses: 915
% 
% Sat splits: 246
% Sat splitting refutations: 1732
% 
% Memory used [KB]: 122172
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.clos_removeProof
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 29
% Passive clauses: 3
% Generated clauses: 8
% Final active clauses: 29
% Input formulas: 21
% Initial clauses: 41
% 
% Selected by SInE selection: 19
% SInE iterations: 1
% 
% Duplicate literals: 1
% 
% Simple tautologies: 13
% Forward subsumptions: 6
% 
% Forward superposition: 7
% 
% SAT solver clauses: 1
% 
% Memory used [KB]: 767
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.clos_removeProof
% (46507)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 136
% Passive clauses: 981
% Generated clauses: 24463
% Final active clauses: 125
% Final passive clauses: 823
% Input formulas: 21
% Initial clauses: 44
% 
% Duplicate literals: 547
% Trivial inequalities: 355
% Fw subsumption resolutions: 2774
% Fw demodulations: 8439
% Bw demodulations: 25
% Global subsumptions: 15
% 
% Simple tautologies: 559
% Equational tautologies: 103
% Forward subsumptions: 10332
% Fw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 1778
% Factoring: 2
% Forward superposition: 5821
% Backward superposition: 4549
% Self superposition: 63
% Equality factoring: 29
% Equality resolution: 11
% 
% Split clauses: 5
% Split components: 11
% SAT solver clauses: 1019
% SAT solver unit clauses: 59
% SAT solver binary clauses: 121
% 
% Sat splits: 5
% Sat splitting refutations: 16
% 
% Memory used [KB]: 8955
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 25 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.clos_removeProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_removeProof
% SZS output start Proof for thm.clos_removeProof
fof(f35271,plain,(
  $false),
  inference(equality_resolution,[],[f35185])).
fof(f35185,plain,(
  ( ! [X3] : (s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),X3) != s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),sK0)))) )),
  inference(resolution,[],[f2083,f108])).
fof(f108,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X31,X32,X33] : (p(s(bool,happ(s(fun(X31,bool),'=_1'(s(X31,X32))),s(X31,X33)))) <=> s(X31,X32) = s(X31,X33))),
  file('thm.clos_removeProof.clos_removeProof_44_prove_ltb',arity_1_HL_EQ)).
fof(f2083,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),bool),'=_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),sK0))))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),X0))))) )),
  inference(resolution,[],[f617,f103])).
fof(f103,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK6(X2,X1,X0)),s(X1,sK7(X2,X1,X0))))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ((! [X8] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK6(X2,X1,X0)),s(X1,sK7(X2,X1,X0)))))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK8(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((! [X8] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | ? [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ? [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(rectify,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : ((! [X3] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ? [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))))))) & (! [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ? [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2] : (? [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ? [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ! [X2] : (? [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ? [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X16,X20] : ! [X21] : (? [X22] : p(s(bool,happ(s(fun('type.pair.prod'(X16,X20),bool),X21),s('type.pair.prod'(X16,X20),X22)))) <=> ? [X23,X24] : p(s(bool,happ(s(fun('type.pair.prod'(X16,X20),bool),X21),s('type.pair.prod'(X16,X20),'const.pair.,_2'(s(X16,X23),s(X20,X24)))))))),
  file('thm.clos_removeProof.clos_removeProof_44_prove_ltb','thm.pair.EXISTS_PROD')).
fof(f617,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),bool),'=_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),sK0))))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X0),s('type.db_vars.db_var_set',X1))))))) )),
  inference(equality_resolution,[],[f145])).
fof(f145,plain,(
  ( ! [X4,X2,X3] : (s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),X4) != s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),sK0))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),bool),'=_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),X4))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X2),s('type.db_vars.db_var_set',X3))))))) )),
  inference(superposition,[],[f79,f107])).
fof(f107,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2) | ~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f79,plain,(
  ( ! [X2,X1] : (s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),sK0))) != s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.db_vars.db_var_set',X2)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X1,X2] : s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),sK0))) != s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.db_vars.db_var_set',X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f48])).
fof(f48,plain,(
  ? [X0] : ! [X1,X2] : s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),X0))) != s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.db_vars.db_var_set',X2)))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ? [X1,X2] : s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),X0))) = s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X1),s('type.db_vars.db_var_set',X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ? [X7,X8] : s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),X6))) = s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.db_vars.db_var_set',X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ? [X7,X8] : s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.clos_remove.remove_1'(s('type.list.list'('type.closLang.exp'),X6))) = s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.db_vars.db_var_set'),'const.pair.,_2'(s('type.list.list'('type.closLang.exp'),X7),s('type.db_vars.db_var_set',X8)))),
  file('thm.clos_removeProof.clos_removeProof_44_prove_ltb','thm.clos_removeProof.clos_removeProof_44_prove_ltb')).
% SZS output end Proof for thm.clos_removeProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 266
% Passive clauses: 1678
% Generated clauses: 34756
% Final active clauses: 223
% Final passive clauses: 1107
% Input formulas: 21
% Initial clauses: 44
% 
% Duplicate literals: 551
% Trivial inequalities: 1850
% Fw subsumption resolutions: 2641
% Fw demodulations: 11471
% Bw demodulations: 160
% Condensations: 343
% 
% Simple tautologies: 746
% Equational tautologies: 12
% Forward subsumptions: 15491
% Backward subsumptions: 217
% 
% Binary resolution: 769
% Forward superposition: 8284
% Backward superposition: 8563
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 96
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 1681
% SAT solver unit clauses: 30
% SAT solver binary clauses: 387
% 
% Sat splits: 1
% 
% Memory used [KB]: 12537
% Time elapsed: 1.018 s
% ------------------------------
% ------------------------------
% Success in time 13.51 s
