Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relationProps
% (91028)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 339
% Passive clauses: 2620
% Generated clauses: 10281
% Final active clauses: 297
% Final passive clauses: 2031
% Input formulas: 38
% Initial clauses: 120
% 
% Duplicate literals: 383
% Trivial inequalities: 79
% Fw subsumption resolutions: 801
% Fw demodulations: 487
% Bw demodulations: 38
% 
% Simple tautologies: 267
% Equational tautologies: 43
% Forward subsumptions: 5420
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 4001
% Factoring: 14
% Forward superposition: 2537
% Backward superposition: 1560
% Self superposition: 57
% Equality factoring: 50
% Equality resolution: 21
% 
% Split clauses: 220
% Split components: 447
% SAT solver clauses: 379
% SAT solver unit clauses: 25
% SAT solver binary clauses: 141
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 2
% 
% Sat splits: 220
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 373
% TWLsolver calls for satisfiability: 169
% 
% Memory used [KB]: 5756
% Time elapsed: 0.620 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_relationProps
% (91053)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 228
% Passive clauses: 1601
% Generated clauses: 7022
% Final active clauses: 202
% Final passive clauses: 1341
% Input formulas: 38
% Initial clauses: 120
% 
% Duplicate literals: 119
% Trivial inequalities: 50
% Fw subsumption resolutions: 264
% Fw demodulations: 514
% Bw demodulations: 36
% Condensations: 184
% Global subsumptions: 13
% 
% Simple tautologies: 148
% Equational tautologies: 24
% Forward subsumptions: 2050
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 754
% Factoring: 5
% Forward superposition: 1649
% Backward superposition: 1508
% Self superposition: 40
% Equality factoring: 30
% Equality resolution: 23
% 
% Split clauses: 26
% Split components: 60
% SAT solver clauses: 2110
% SAT solver unit clauses: 133
% SAT solver binary clauses: 71
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 13
% 
% Sat splits: 26
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 2056
% TWLsolver calls for satisfiability: 1070
% 
% Memory used [KB]: 12920
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_relationProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relationProps
% SZS output start Proof for thm.clos_relationProps
fof(f1985,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f193,f319,f318,f192,f321,f191,f323,f317,f867,f871,f874,f877,f880,f883,f886,f889,f911,f868,f913,f872,f915,f875,f917,f878,f919,f881,f921,f884,f923,f887,f925,f869,f890,f945,f946,f950,f952,f954,f1535,f1541,f1548,f1554,f1559,f1562,f1565,f1568,f1571,f1574,f1577,f1580,f1583,f1586,f1592,f1599,f1607,f1613,f1618,f1625,f1631,f1638,f1643,f1646,f1649,f1652,f1655,f1658,f1661,f1664,f1667,f1670,f1675,f1683,f1690,f1697,f1856,f1861,f1864,f1865,f1868,f1917,f1926,f1927,f1974,f1983,f1984])).
fof(f1984,plain,(
  $false | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f1951])).
fof(f1951,plain,(
  s(sK7,sK10) != s(sK7,sK10) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f264,f1938,f268])).
fof(f268,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X7] : ! [X17,X18,X12,X13] : (s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13))) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X17),s(X7,X18))) <=> (s(X6,X12) = s(X6,X17) & s(X7,X13) = s(X7,X18)))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.pair.CLOSED_PAIR_EQ')).
fof(f1938,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X2)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1937,f317])).
fof(f1937,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(X0,X2)))) ) | $spl2),
  inference(forward_demodulation,[],[f1936,f264])).
fof(f1936,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK9,sK17(sK11,sK9,sK8)))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(X0,X2)))) ) | $spl2),
  inference(forward_demodulation,[],[f1935,f310])).
fof(f310,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X7,X8] : ! [X18,X17,X11] : s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X7,X8),'const.pair.,_2'(s(X7,X17),s(X8,X18))),s(fun(X7,fun(X8,X6)),X11))) = s(X6,happ(s(fun(X8,X6),happ(s(fun(X7,fun(X8,X6)),X11),s(X7,X17))),s(X8,X18)))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.pair.pair_case_thm')).
fof(f1935,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK16(sK11,sK9,sK8)))),s(sK9,sK17(sK11,sK9,sK8)))),s(X0,X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1854,f266])).
fof(f266,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X0,X2) = s(X0,X4)) )),
  inference(cnf_transformation,[],[f170])).
fof(f1854,plain,(
  s(sK7,sK10) != s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK16(sK11,sK9,sK8)))),s(sK9,sK17(sK11,sK9,sK8)))) | $spl2),
  inference(unit_resulting_resolution,[],[f264,f319])).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : ! [X17] : ? [X20,X21] : s('type.pair.prod'(X6,X7),X17) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X20),s(X7,X21)))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.pair.ABS_PAIR_THM')).
fof(f1983,plain,(
  $false | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f1948])).
fof(f1948,plain,(
  s(sK7,sK10) != s(sK7,sK10) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f264,f1938,f268])).
fof(f1974,plain,(
  $false | ($spl0 | $spl2)),
  inference(equality_resolution,[],[f1938])).
fof(f1927,plain,(
  $false | ($spl0 | $spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f1888])).
fof(f1888,plain,(
  s(sK7,sK10) != s(sK7,sK10) | ($spl0 | $spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f323,f1875,f268])).
fof(f1875,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X2)))) ) | ($spl0 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f1874,f317])).
fof(f1874,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(X0,X2)))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f1873,f323])).
fof(f1873,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(X0,X2)))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f1872,f310])).
fof(f1872,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(X0,X2)))) ) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f1853,f266])).
fof(f1853,plain,(
  s(sK7,sK10) != s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f323,f319])).
fof(f1926,plain,(
  $false | ($spl0 | $spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f1885])).
fof(f1885,plain,(
  s(sK7,sK10) != s(sK7,sK10) | ($spl0 | $spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f323,f1875,f268])).
fof(f1917,plain,(
  $false | ($spl0 | $spl2 | $spl6)),
  inference(equality_resolution,[],[f1875])).
fof(f1868,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1860,f323])).
fof(f1860,plain,(
  s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))) | ($spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f1859])).
fof(f1859,plain,(
  s(sK7,sK10) != s(sK7,sK10) | s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))) | ($spl2 | $spl4)),
  inference(superposition,[],[f319,f321])).
fof(f1865,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1853,f321])).
fof(f1864,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1863,f321])).
fof(f1863,plain,(
  s(sK7,sK10) != s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f1862,f1208])).
fof(f1208,plain,(
  s(sK8,sK13) = s(sK8,sK16(sK11,sK9,sK8)) | $spl6),
  inference(unit_resulting_resolution,[],[f264,f1170])).
fof(f1170,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,X0),s(sK9,X1))) | s(sK8,sK13) = s(sK8,X0)) ) | $spl6),
  inference(superposition,[],[f266,f323])).
fof(f1862,plain,(
  s(sK7,sK10) != s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK16(sK11,sK9,sK8)))),s(sK9,sK14))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f1854,f1351])).
fof(f1351,plain,(
  s(sK9,sK14) = s(sK9,sK17(sK11,sK9,sK8)) | $spl6),
  inference(unit_resulting_resolution,[],[f264,f1252])).
fof(f1252,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,X0),s(sK9,X1))) | s(sK9,sK14) = s(sK9,X1)) ) | $spl6),
  inference(superposition,[],[f267,f323])).
fof(f267,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f170])).
fof(f1861,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1855,f323])).
fof(f1855,plain,(
  s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))) | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f321,f319])).
fof(f1856,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f323,f321,f319])).
fof(f1697,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1696,f1144])).
fof(f1144,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,sK10),s(X0,X1))) != s('type.pair.prod'(sK7,X0),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(X0,X2)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f318,f266])).
fof(f1696,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1695,f323])).
fof(f1695,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1694,f310])).
fof(f1694,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1406,f323])).
fof(f1406,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f323,f268])).
fof(f1690,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1689,f1144])).
fof(f1689,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1688,f323])).
fof(f1688,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1411,f310])).
fof(f1411,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f323,f268])).
fof(f1683,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1682,f1144])).
fof(f1682,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(X0,X1),X2)))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1681,f323])).
fof(f1681,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f1680,f310])).
fof(f1680,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f1416,f264])).
fof(f1416,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0))))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f321,f264,f268])).
fof(f1675,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1674,f1144])).
fof(f1674,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0))))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1673,f323])).
fof(f1673,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f1421,f310])).
fof(f1421,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f321,f264,f268])).
fof(f1670,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1669,f1224])).
fof(f1224,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,sK7),'const.pair.,_2'(s(X0,X1),s(sK7,sK10))) != s('type.pair.prod'(X0,sK7),'const.pair.,_2'(s(X0,X2),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f318,f267])).
fof(f1669,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1668,f323])).
fof(f1668,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1424,f310])).
fof(f1424,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f323,f321,f268])).
fof(f1667,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1666,f1224])).
fof(f1666,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12)))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1665,f323])).
fof(f1665,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12)))))) ) | $spl4),
  inference(forward_demodulation,[],[f1425,f310])).
fof(f1425,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f264,f321,f268])).
fof(f1664,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1663,f1144])).
fof(f1663,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1662,f323])).
fof(f1662,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | $spl4),
  inference(forward_demodulation,[],[f1426,f310])).
fof(f1426,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | $spl4),
  inference(unit_resulting_resolution,[],[f321,f321,f268])).
fof(f1661,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1660,f1224])).
fof(f1660,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1659,f323])).
fof(f1659,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1427,f310])).
fof(f1427,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f1208,f321,f268])).
fof(f1658,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1657,f1224])).
fof(f1657,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1656,f323])).
fof(f1656,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1428,f310])).
fof(f1428,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f1351,f321,f268])).
fof(f1655,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1654,f1224])).
fof(f1654,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1653,f323])).
fof(f1653,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1429,f310])).
fof(f1429,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f323,f321,f268])).
fof(f1652,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1651,f1224])).
fof(f1651,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,sK10)))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1650,f323])).
fof(f1650,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12)))))) ) | $spl4),
  inference(forward_demodulation,[],[f1430,f310])).
fof(f1430,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f264,f321,f268])).
fof(f1649,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1648,f1224])).
fof(f1648,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1647,f323])).
fof(f1647,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | $spl4),
  inference(forward_demodulation,[],[f1431,f310])).
fof(f1431,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | $spl4),
  inference(unit_resulting_resolution,[],[f321,f321,f268])).
fof(f1646,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1645,f1224])).
fof(f1645,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1644,f323])).
fof(f1644,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1432,f310])).
fof(f1432,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f1208,f321,f268])).
fof(f1643,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1642,f1224])).
fof(f1642,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1641,f323])).
fof(f1641,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1433,f310])).
fof(f1433,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f1351,f321,f268])).
fof(f1638,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1637,f1144])).
fof(f1637,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK13))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1636,f323])).
fof(f1636,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK13))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1635,f310])).
fof(f1635,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK13))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1436,f1208])).
fof(f1436,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK13))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK8,sK16(sK11,sK9,sK8)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f1208,f268])).
fof(f1631,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1630,f1144])).
fof(f1630,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK16(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1629,f323])).
fof(f1629,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK16(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1441,f310])).
fof(f1441,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK16(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f1208,f268])).
fof(f1625,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1624,f1144])).
fof(f1624,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK14))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1623,f323])).
fof(f1623,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK14))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1622,f310])).
fof(f1622,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK14))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1446,f1351])).
fof(f1446,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK14))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK9,sK17(sK11,sK9,sK8)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f1351,f268])).
fof(f1618,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1617,f1144])).
fof(f1617,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK17(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1616,f323])).
fof(f1616,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK17(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1451,f310])).
fof(f1451,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK17(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f1351,f268])).
fof(f1613,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1612,f1144])).
fof(f1612,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1611,f323])).
fof(f1611,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1456,f310])).
fof(f1456,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f323,f268])).
fof(f1607,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1606,f1144])).
fof(f1606,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1605,f323])).
fof(f1605,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1604,f310])).
fof(f1604,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(sK8,sK9),sK11))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1461,f323])).
fof(f1461,plain,(
  s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(sK8,sK9),sK11))) = s('type.pair.prod'(sK7,'type.pair.prod'(sK8,sK9)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f323,f268])).
fof(f1599,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1598,f1144])).
fof(f1598,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0))))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1597,f323])).
fof(f1597,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f1466,f310])).
fof(f1466,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f321,f264,f268])).
fof(f1592,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1591,f1144])).
fof(f1591,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(X0,X1),X2)))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1590,f323])).
fof(f1590,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f1589,f310])).
fof(f1589,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(X0,X1),X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f1471,f264])).
fof(f1471,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,sK10),s('type.pair.prod'(X0,X1),X2))) = s('type.pair.prod'(sK7,'type.pair.prod'(X0,X1)),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0))))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f321,f264,f268])).
fof(f1586,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1585,f1224])).
fof(f1585,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1584,f323])).
fof(f1584,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1474,f310])).
fof(f1474,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f323,f321,f268])).
fof(f1583,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1582,f1224])).
fof(f1582,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,sK10)))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1581,f323])).
fof(f1581,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12)))))) ) | $spl4),
  inference(forward_demodulation,[],[f1475,f310])).
fof(f1475,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f264,f321,f268])).
fof(f1580,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1579,f1224])).
fof(f1579,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1578,f323])).
fof(f1578,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | $spl4),
  inference(forward_demodulation,[],[f1476,f310])).
fof(f1476,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | $spl4),
  inference(unit_resulting_resolution,[],[f321,f321,f268])).
fof(f1577,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1576,f1224])).
fof(f1576,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1575,f323])).
fof(f1575,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1477,f310])).
fof(f1477,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f1208,f321,f268])).
fof(f1574,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1573,f1224])).
fof(f1573,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1572,f323])).
fof(f1572,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1478,f310])).
fof(f1478,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f1351,f321,f268])).
fof(f1571,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1570,f1224])).
fof(f1570,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1569,f323])).
fof(f1569,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1479,f310])).
fof(f1479,plain,(
  s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),sK11),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(sK8,sK9),sK7),'const.pair.,_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f323,f321,f268])).
fof(f1568,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1567,f1224])).
fof(f1567,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12)))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1566,f323])).
fof(f1566,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12)))))) ) | $spl4),
  inference(forward_demodulation,[],[f1480,f310])).
fof(f1480,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),X2),s(sK7,sK10))) = s('type.pair.prod'('type.pair.prod'(X0,X1),sK7),'const.pair.,_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK16(X2,X1,X0)),s(X1,sK17(X2,X1,X0)))),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f264,f321,f268])).
fof(f1565,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1564,f1144])).
fof(f1564,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1563,f323])).
fof(f1563,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | $spl4),
  inference(forward_demodulation,[],[f1481,f310])).
fof(f1481,plain,(
  s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,sK10),s(sK7,sK10))) = s('type.pair.prod'(sK7,sK7),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | $spl4),
  inference(unit_resulting_resolution,[],[f321,f321,f268])).
fof(f1562,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1561,f1224])).
fof(f1561,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1560,f323])).
fof(f1560,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1482,f310])).
fof(f1482,plain,(
  s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK13),s(sK7,sK10))) = s('type.pair.prod'(sK8,sK7),'const.pair.,_2'(s(sK8,sK16(sK11,sK9,sK8)),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f1208,f321,f268])).
fof(f1559,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1558,f1224])).
fof(f1558,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1557,f323])).
fof(f1557,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1483,f310])).
fof(f1483,plain,(
  s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK14),s(sK7,sK10))) = s('type.pair.prod'(sK9,sK7),'const.pair.,_2'(s(sK9,sK17(sK11,sK9,sK8)),s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f1351,f321,f268])).
fof(f1554,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1553,f1144])).
fof(f1553,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK16(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1552,f323])).
fof(f1552,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK16(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1486,f310])).
fof(f1486,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK16(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f1208,f268])).
fof(f1548,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1547,f1144])).
fof(f1547,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK13))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1546,f323])).
fof(f1546,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK13))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1545,f310])).
fof(f1545,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK13))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK8,sK13))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1491,f1208])).
fof(f1491,plain,(
  s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,sK10),s(sK8,sK13))) = s('type.pair.prod'(sK7,sK8),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK8,sK16(sK11,sK9,sK8)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f1208,f268])).
fof(f1541,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1540,f1144])).
fof(f1540,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK17(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1539,f323])).
fof(f1539,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK17(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1496,f310])).
fof(f1496,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK17(sK11,sK9,sK8)))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f1351,f268])).
fof(f1535,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1534,f1144])).
fof(f1534,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK14))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1533,f323])).
fof(f1533,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK14))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1532,f310])).
fof(f1532,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK14))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK9,sK14))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1501,f1351])).
fof(f1501,plain,(
  s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,sK10),s(sK9,sK14))) = s('type.pair.prod'(sK7,sK9),'const.pair.,_2'(s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))),s(sK9,sK17(sK11,sK9,sK8)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f321,f1351,f268])).
fof(f954,plain,(
  $false | ($spl4 | $spl6 | $spl10)),
  inference(global_subsumption,[],[f953,f193,f192,f191,f194,f324,f201,f209,f208,f213,f212,f240,f239,f238,f244,f241,f251,f250,f325,f264,f265,f268,f267,f266,f275,f274,f273,f272,f271,f270,f269,f279,f278,f277,f276,f283,f282,f281,f280,f287,f286,f285,f284,f291,f290,f289,f288,f298,f297,f296,f311,f294,f293,f292,f302,f301,f300,f299,f306,f305,f304,f303,f309,f310,f326,f327,f328,f329,f330,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f349,f352,f353,f355,f362,f375,f411,f420,f419,f430,f428,f442,f452,f458,f462,f467,f475,f479,f485,f506,f528,f549,f550,f575,f621,f649,f867,f871,f874,f877,f880,f883,f886,f889,f669,f891,f892,f893,f894,f895,f896,f897,f898,f899,f900,f901,f902,f903,f696,f911,f913,f915,f917,f919,f921,f923,f925,f869,f944,f949])).
fof(f949,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK15(X2,X1,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f933,f194])).
fof(f933,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X2) = s(X0,sK15(X2,X1,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f869,f239])).
fof(f944,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f926,f869])).
fof(f926,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f869,f212])).
fof(f696,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,happ(s(fun(bool,bool),X74),s(bool,'T_0')))) | s(bool,X73) = s(bool,sK15(X73,X74,bool)) | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f239,f375])).
fof(f903,plain,(
  ( ! [X70,X68,X69,X67] : (s(X67,X69) = s(X67,sK15(X69,X68,X67)) | s(bool,X70) = s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f694,f194])).
fof(f694,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | s(X67,X69) = s(X67,sK15(X69,X68,X67)) | s(bool,X70) = s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f239,f375])).
fof(f902,plain,(
  ( ! [X54,X56,X55,X53] : (~sP5(X53,X54,X55) | sP6(X53,X54,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f687,f194])).
fof(f687,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~sP5(X53,X54,X55) | sP6(X53,X54,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f311,f375])).
fof(f901,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X49,X51,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f686,f194])).
fof(f686,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X49,X51,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f306,f375])).
fof(f900,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | sP6(X45,X46,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f685,f194])).
fof(f685,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | sP6(X45,X46,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f305,f375])).
fof(f899,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,X42)) | sP6(X41,X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f684,f194])).
fof(f684,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | sP6(X41,X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f304,f375])).
fof(f898,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | ~sP5(X39,X38,X37) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f683,f194])).
fof(f683,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~sP5(X39,X38,X37) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f299,f375])).
fof(f897,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | ~sP2(X34,X33,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f682,f194])).
fof(f682,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP2(X34,X33,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f280,f375])).
fof(f896,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP3(X30,X29,X31) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f681,f194])).
fof(f681,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP3(X30,X29,X31) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f276,f375])).
fof(f895,plain,(
  ( ! [X28,X26,X27,X25] : (sP4(X26,X27,X25) | ~sP3(X26,X25,X27) | ~sP2(X27,X25,X26) | ~sP1(X25,X27,X26) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f680,f194])).
fof(f680,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'T_0')) | sP4(X26,X27,X25) | ~sP3(X26,X25,X27) | ~sP2(X27,X25,X26) | ~sP1(X25,X27,X26) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f273,f375])).
fof(f894,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f679,f194])).
fof(f679,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f251,f375])).
fof(f893,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f678,f194])).
fof(f678,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f250,f375])).
fof(f892,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f677,f194])).
fof(f677,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f244,f375])).
fof(f891,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f673,f194])).
fof(f673,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f212,f375])).
fof(f669,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,happ(s(fun(bool,bool),X65),s(bool,X64)))) | s(bool,'T_0') = s(bool,sK15('T_0',X65,bool)) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f239,f375])).
fof(f649,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f208,f375])).
fof(f621,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,happ(s(fun(bool,bool),X76),s(bool,X75)))) | s(bool,X74) = s(bool,sK15(X74,X76,bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f239,f375])).
fof(f575,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f375,f375])).
fof(f550,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0')))) | s(bool,X4) = s(bool,sK15(X4,X5,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f239,f201])).
fof(f549,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) | s(bool,'F_0') = s(bool,sK15('F_0',X3,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f239,f201])).
fof(f528,plain,(
  sP4('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f334,f331,f329,f324,f274])).
fof(f506,plain,(
  sP4('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f336,f331,f328,f194,f273])).
fof(f485,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f479,f294])).
fof(f479,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f194,f324,f194,f299])).
fof(f475,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f339,f194,f324,f298])).
fof(f467,plain,(
  ~sP4('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f462,f269])).
fof(f462,plain,(
  ~sP1('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f324,f324,f324,f284])).
fof(f458,plain,(
  ~sP4('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f452,f270])).
fof(f452,plain,(
  ~sP2('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f324,f194,f194,f280])).
fof(f442,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f324,f194,f194,f276])).
fof(f428,plain,(
  ~sP4('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f194,f340,f290])).
fof(f430,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f194,f340,f290])).
fof(f419,plain,(
  sP4('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f324,f340,f288])).
fof(f420,plain,(
  sP4('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f324,f340,f288])).
fof(f411,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f408,f324])).
fof(f408,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f241,f201])).
fof(f375,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f201,f201])).
fof(f362,plain,(
  ( ! [X0] : (sP6('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f194,f194,f305])).
fof(f355,plain,(
  ( ! [X0] : (sP6('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f324,f194,f304])).
fof(f353,plain,(
  ( ! [X0] : (~sP6('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f194,f324,f293])).
fof(f352,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f324,f324,f292])).
fof(f349,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f324,f324,f325])).
fof(f340,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f324,f208])).
fof(f339,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f324,f302])).
fof(f338,plain,(
  ( ! [X0,X1] : (sP5(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f194,f301])).
fof(f337,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f324,f300])).
fof(f336,plain,(
  ( ! [X0,X1] : (sP1('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f194,f287])).
fof(f335,plain,(
  ( ! [X0,X1] : (sP1(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f194,f286])).
fof(f334,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f194,f285])).
fof(f333,plain,(
  ( ! [X0,X1] : (sP2('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f324,f283])).
fof(f332,plain,(
  ( ! [X0,X1] : (sP2(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f324,f282])).
fof(f331,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f194,f281])).
fof(f330,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f324,f279])).
fof(f329,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f324,f278])).
fof(f328,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f194,f277])).
fof(f327,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f194,f251])).
fof(f326,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f194,f250])).
fof(f309,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3,X4] : s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X7,X8,X6] : ! [X22,X23] : s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X7,X8),X22),s(fun(X7,fun(X8,X6)),X23))) = s(X6,happ(s(fun(fun(X7,fun(X8,X6)),X6),happ(s(fun('type.pair.prod'(X7,X8),fun(fun(X7,fun(X8,X6)),X6)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X7,X8),X22))),s(fun(X7,fun(X8,X6)),X23)))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','ARITY_0_2_const.pair.pair_CASE')).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP6(X2,X0,X1))),
  inference(definition_folding,[],[f125,f134,f133])).
fof(f133,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f134,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X21,X20,X10] : ((p(s(bool,X10)) <=> (p(s(bool,X20)) => p(s(bool,X21)))) <=> ((p(s(bool,X10)) | p(s(bool,X20))) & (p(s(bool,X10)) | ~p(s(bool,X21))) & (~p(s(bool,X20)) | p(s(bool,X21)) | ~p(s(bool,X10)))))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.sat.sat_27_prove_ltb')).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f187])).
fof(f187,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f188])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f188])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f188])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP5(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP5(X0,X1,X2) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f184])).
fof(f184,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP5(X2,X0,X1) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP5(X2,X0,X1)) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f134])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f185])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f185])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP5(X0,X1,X2) | sP6(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f295])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f185])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f185])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP4(X2,X1,X0)) & (~sP4(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP4(X2,X1,X0))),
  inference(definition_folding,[],[f92,f131,f130,f129,f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : (sP1(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f129,plain,(
  ! [X1,X0,X2] : (sP2(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f130,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f131,plain,(
  ! [X2,X1,X0] : (sP4(X2,X1,X0) <=> (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f92,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X21,X20,X10] : ((p(s(bool,X10)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X10)) | p(s(bool,X20)) | p(s(bool,X21))) & (p(s(bool,X10)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X10))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X10)))))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(nnf_transformation,[],[f128])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f181])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f181])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(nnf_transformation,[],[f129])).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f179])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f175])).
fof(f175,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f130])).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f176])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f176])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (sP1(X2,X1,X0) & sP2(X1,X2,X0) & sP3(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP1(X2,X1,X0) | ~sP2(X1,X2,X0) | ~sP3(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f172])).
fof(f172,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP4(X2,X1,X0)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP4(X2,X1,X0)))),
  inference(nnf_transformation,[],[f131])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP2(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP3(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f173])).
fof(f273,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2) | ~sP3(X0,X2,X1) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f274,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X2,X1) | p(s(bool,X1)) | sP4(X0,X1,X2) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X2,X1) | p(s(bool,X0)) | sP4(X0,X1,X2) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f265,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK18(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK18(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK18(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK18(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb',aHL_EXT)).
fof(f325,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f249,f324])).
fof(f249,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f250,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f162])).
fof(f251,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f162])).
fof(f241,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f244,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f111])).
fof(f238,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK15(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X19,X12] : (? [X17] : (s(X6,X12) = s(X6,X17) & p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X17))))) <=> p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X12)))))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.bool.UNWIND_THM2')).
fof(f239,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK15(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f160])).
fof(f240,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f212,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f213,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f143])).
fof(f208,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f209,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f201,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb',aHL_BOOL_CASES)).
fof(f324,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f197,f194])).
fof(f197,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f194,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb',aHL_TRUTH)).
fof(f953,plain,(
  ( ! [X1] : (s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) = s(sK7,sK15(happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)),X1,sK7))) ) | ($spl4 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f939,f194])).
fof(f939,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) = s(sK7,sK15(happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)),X1,sK7))) ) | ($spl4 | $spl6 | $spl10)),
  inference(backward_demodulation,[],[f869,f554])).
fof(f554,plain,(
  ( ! [X1] : (s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) = s(sK7,sK15(happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)),X1,sK7)) | ~p(s(bool,happ(s(fun(sK7,bool),X1),s(sK7,sK10))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f553,f323])).
fof(f553,plain,(
  ( ! [X1] : (s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))),s(fun(sK8,fun(sK9,sK7)),sK12))) = s(sK7,sK15(happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)),X1,sK7)) | ~p(s(bool,happ(s(fun(sK7,bool),X1),s(sK7,sK10))))) ) | $spl4),
  inference(forward_demodulation,[],[f548,f310])).
fof(f548,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK7,bool),X1),s(sK7,sK10)))) | s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))) = s(sK7,sK15(happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14)),X1,sK7))) ) | $spl4),
  inference(superposition,[],[f239,f321])).
fof(f952,plain,(
  $false | ($spl6 | $spl10)),
  inference(global_subsumption,[],[f951,f193,f192,f191,f194,f324,f201,f209,f208,f213,f212,f240,f239,f238,f244,f241,f251,f250,f325,f264,f265,f268,f267,f266,f275,f274,f273,f272,f271,f270,f269,f279,f278,f277,f276,f283,f282,f281,f280,f287,f286,f285,f284,f291,f290,f289,f288,f298,f297,f296,f311,f294,f293,f292,f302,f301,f300,f299,f306,f305,f304,f303,f309,f310,f326,f327,f328,f329,f330,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f349,f352,f353,f355,f362,f375,f411,f420,f419,f430,f428,f442,f452,f458,f462,f467,f475,f479,f485,f506,f528,f549,f550,f575,f621,f649,f867,f871,f874,f877,f880,f883,f886,f889,f669,f891,f892,f893,f894,f895,f896,f897,f898,f899,f900,f901,f902,f903,f696,f911,f913,f915,f917,f919,f921,f923,f925,f869,f944,f949])).
fof(f951,plain,(
  ( ! [X0] : (s('type.pair.prod'(sK8,sK9),sK11) = s('type.pair.prod'(sK8,sK9),sK15('const.pair.,_2'(s(sK8,sK13),s(sK9,sK14)),X0,'type.pair.prod'(sK8,sK9)))) ) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f938,f194])).
fof(f938,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.pair.prod'(sK8,sK9),sK11) = s('type.pair.prod'(sK8,sK9),sK15('const.pair.,_2'(s(sK8,sK13),s(sK9,sK14)),X0,'type.pair.prod'(sK8,sK9)))) ) | ($spl6 | $spl10)),
  inference(backward_demodulation,[],[f869,f552])).
fof(f552,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK8,sK9),bool),X0),s('type.pair.prod'(sK8,sK9),sK11)))) | s('type.pair.prod'(sK8,sK9),sK11) = s('type.pair.prod'(sK8,sK9),sK15('const.pair.,_2'(s(sK8,sK13),s(sK9,sK14)),X0,'type.pair.prod'(sK8,sK9)))) ) | $spl6),
  inference(forward_demodulation,[],[f547,f323])).
fof(f547,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK8,sK9),bool),X0),s('type.pair.prod'(sK8,sK9),sK11)))) | s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))) = s('type.pair.prod'(sK8,sK9),sK15('const.pair.,_2'(s(sK8,sK13),s(sK9,sK14)),X0,'type.pair.prod'(sK8,sK9)))) ) | $spl6),
  inference(superposition,[],[f239,f323])).
fof(f950,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f193,f192,f191,f194,f324,f201,f209,f208,f213,f212,f240,f239,f238,f244,f241,f251,f250,f325,f264,f265,f268,f267,f266,f275,f274,f273,f272,f271,f270,f269,f279,f278,f277,f276,f283,f282,f281,f280,f287,f286,f285,f284,f291,f290,f289,f288,f298,f297,f296,f311,f294,f293,f292,f302,f301,f300,f299,f306,f305,f304,f303,f309,f310,f326,f327,f328,f329,f330,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f349,f352,f353,f355,f362,f375,f411,f420,f419,f430,f428,f442,f452,f458,f462,f467,f475,f479,f485,f506,f528,f549,f550,f575,f621,f649,f867,f871,f874,f877,f880,f883,f886,f889,f669,f891,f892,f893,f894,f895,f896,f897,f898,f899,f900,f901,f902,f903,f696,f911,f913,f915,f917,f919,f921,f923,f925,f869,f944,f949])).
fof(f946,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f928,f194])).
fof(f928,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f869,f324])).
fof(f945,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f193,f192,f191,f194,f324,f201,f209,f208,f213,f212,f240,f239,f238,f244,f241,f251,f250,f325,f264,f265,f268,f267,f266,f275,f274,f273,f272,f271,f270,f269,f279,f278,f277,f276,f283,f282,f281,f280,f287,f286,f285,f284,f291,f290,f289,f288,f298,f297,f296,f311,f294,f293,f292,f302,f301,f300,f299,f306,f305,f304,f303,f309,f310,f326,f327,f328,f329,f330,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f349,f352,f353,f355,f362,f375,f411,f420,f419,f430,f428,f442,f452,f458,f462,f467,f475,f479,f485,f506,f528,f549,f550,f575,f621,f649,f867,f871,f874,f877,f880,f883,f886,f889,f669,f891,f892,f893,f894,f895,f896,f897,f898,f899,f900,f901,f902,f903,f696,f911,f913,f915,f917,f919,f921,f923,f925,f869,f944])).
fof(f890,plain,(
  ( ! [X2,X1] : (sP6('T_0',X1,X2) | ~sP5('T_0',X1,X2)) ) | $spl24),
  inference(cnf_transformation,[],[f890_D])).
fof(f890_D,plain,(
  ( ! [X2,X1] : (sP6('T_0',X1,X2) | ~sP5('T_0',X1,X2)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f869,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f869_D])).
fof(f869_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f925,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP5('T_0',X1,X2) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f924])).
fof(f924,plain,(
  ( ! [X59,X57,X60,X58] : (~sP5('T_0',X58,X59) | sP6('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f715,f411])).
fof(f715,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~sP5('T_0',X58,X59) | sP6('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f311,f375])).
fof(f887,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f887_D])).
fof(f887_D,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f923,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f922])).
fof(f922,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP6('T_0',X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f714,f411])).
fof(f714,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP6('T_0',X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f306,f375])).
fof(f884,plain,(
  ( ! [X2,X1] : (sP6('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl20),
  inference(cnf_transformation,[],[f884_D])).
fof(f884_D,plain,(
  ( ! [X2,X1] : (sP6('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f921,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f920])).
fof(f920,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | sP6('T_0',X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f713,f411])).
fof(f713,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | sP6('T_0',X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f305,f375])).
fof(f881,plain,(
  ( ! [X2,X1] : (sP6('T_0',X2,X1) | p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f881_D])).
fof(f881_D,plain,(
  ( ! [X2,X1] : (sP6('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f919,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f918])).
fof(f918,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X46)) | sP6('T_0',X47,X46) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f712,f411])).
fof(f712,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | sP6('T_0',X47,X46) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f304,f375])).
fof(f878,plain,(
  ( ! [X2,X1] : (~sP5(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl16),
  inference(cnf_transformation,[],[f878_D])).
fof(f878_D,plain,(
  ( ! [X2,X1] : (~sP5(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f917,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X1,'T_0')) )),
  inference(condensation,[],[f916])).
fof(f916,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~sP5(X43,X42,'T_0') | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f711,f411])).
fof(f711,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,X43)) | ~sP5(X43,X42,'T_0') | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f299,f375])).
fof(f875,plain,(
  ( ! [X2,X1] : (~sP2(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f875_D])).
fof(f875_D,plain,(
  ( ! [X2,X1] : (~sP2(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f915,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X1,'T_0',X2)) )),
  inference(condensation,[],[f914])).
fof(f914,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~sP2(X38,'T_0',X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f710,f411])).
fof(f710,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~sP2(X38,'T_0',X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f280,f375])).
fof(f872,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f872_D])).
fof(f872_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f913,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f912])).
fof(f912,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | ~sP3(X34,'T_0',X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f709,f411])).
fof(f709,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP3(X34,'T_0',X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f276,f375])).
fof(f868,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1) | sP4(X1,X2,'T_0')) ) | $spl8),
  inference(cnf_transformation,[],[f868_D])).
fof(f868_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1) | sP4(X1,X2,'T_0')) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f911,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP4(X1,X2,'T_0') | ~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f910,f336])).
fof(f910,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP4(X1,X2,'T_0') | ~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1) | ~sP1('T_0',X2,X1)) )),
  inference(condensation,[],[f909])).
fof(f909,plain,(
  ( ! [X30,X31,X29,X32] : (sP4(X30,X31,'T_0') | ~sP3(X30,'T_0',X31) | ~sP2(X31,'T_0',X30) | ~sP1('T_0',X31,X30) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f708,f411])).
fof(f708,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X29)) | sP4(X30,X31,'T_0') | ~sP3(X30,'T_0',X31) | ~sP2(X31,'T_0',X30) | ~sP1('T_0',X31,X30) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f273,f375])).
fof(f889,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP5('T_0',X1,X2) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f888])).
fof(f888,plain,(
  ( ! [X59,X57,X60,X58] : (~sP5('T_0',X58,X59) | sP6('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f667,f411])).
fof(f667,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~sP5('T_0',X58,X59) | sP6('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f311,f375])).
fof(f886,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f885])).
fof(f885,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP6('T_0',X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f666,f411])).
fof(f666,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP6('T_0',X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f306,f375])).
fof(f883,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f882])).
fof(f882,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | sP6('T_0',X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f665,f411])).
fof(f665,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | sP6('T_0',X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f305,f375])).
fof(f880,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f879])).
fof(f879,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X46)) | sP6('T_0',X47,X46) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f664,f411])).
fof(f664,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | sP6('T_0',X47,X46) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f304,f375])).
fof(f877,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X1,'T_0')) )),
  inference(condensation,[],[f876])).
fof(f876,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~sP5(X43,X42,'T_0') | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f663,f411])).
fof(f663,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,X43)) | ~sP5(X43,X42,'T_0') | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f299,f375])).
fof(f874,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X1,'T_0',X2)) )),
  inference(condensation,[],[f873])).
fof(f873,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~sP2(X38,'T_0',X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f662,f411])).
fof(f662,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~sP2(X38,'T_0',X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f280,f375])).
fof(f871,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f870])).
fof(f870,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | ~sP3(X34,'T_0',X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f661,f411])).
fof(f661,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP3(X34,'T_0',X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f276,f375])).
fof(f867,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP4(X1,X2,'T_0') | ~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f866,f336])).
fof(f866,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP4(X1,X2,'T_0') | ~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1) | ~sP1('T_0',X2,X1)) )),
  inference(condensation,[],[f865])).
fof(f865,plain,(
  ( ! [X30,X31,X29,X32] : (sP4(X30,X31,'T_0') | ~sP3(X30,'T_0',X31) | ~sP2(X31,'T_0',X30) | ~sP1('T_0',X31,X30) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f660,f411])).
fof(f660,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X29)) | sP4(X30,X31,'T_0') | ~sP3(X30,'T_0',X31) | ~sP2(X31,'T_0',X30) | ~sP1('T_0',X31,X30) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f273,f375])).
fof(f317,plain,(
  s(sK7,sK10) = s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) | $spl0),
  inference(cnf_transformation,[],[f317_D])).
fof(f317_D,plain,(
  s(sK7,sK10) = s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f323,plain,(
  s('type.pair.prod'(sK8,sK9),sK11) = s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))) | $spl6),
  inference(cnf_transformation,[],[f323_D])).
fof(f323_D,plain,(
  s('type.pair.prod'(sK8,sK9),sK11) = s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f191,plain,(
  s('type.pair.prod'(sK8,sK9),sK11) = s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))) | s(sK7,sK10) = s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12)))),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  (s(sK7,sK10) = s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) | (s('type.pair.prod'(sK8,sK9),sK11) = s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,sK13),s(sK9,sK14))) & s(sK7,sK10) = s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))))) & (s(sK7,sK10) != s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) | ! [X6,X7] : (s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,X6),s(sK9,X7))) | s(sK7,sK10) != s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,X6))),s(sK9,X7)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9,sK10,sK11,sK12,sK13,sK14])],[f137])).
fof(f137,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X0,X3) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),X4),s(fun(X1,fun(X2,X0)),X5))) | ? [X8,X9] : (s('type.pair.prod'(X1,X2),X4) = s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X8),s(X2,X9))) & s(X0,X3) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X8))),s(X2,X9))))) & (s(X0,X3) != s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),X4),s(fun(X1,fun(X2,X0)),X5))) | ! [X6,X7] : (s('type.pair.prod'(X1,X2),X4) != s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7))) | s(X0,X3) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X6))),s(X2,X7))))))),
  inference(rectify,[],[f136])).
fof(f136,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X0,X3) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),X4),s(fun(X1,fun(X2,X0)),X5))) | ? [X6,X7] : (s('type.pair.prod'(X1,X2),X4) = s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7))) & s(X0,X3) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X6))),s(X2,X7))))) & (s(X0,X3) != s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),X4),s(fun(X1,fun(X2,X0)),X5))) | ! [X6,X7] : (s('type.pair.prod'(X1,X2),X4) != s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7))) | s(X0,X3) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X6))),s(X2,X7))))))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(X0,X3) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),X4),s(fun(X1,fun(X2,X0)),X5))) <~> ? [X6,X7] : (s('type.pair.prod'(X1,X2),X4) = s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7))) & s(X0,X3) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X6))),s(X2,X7)))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s(X0,X3) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),X4),s(fun(X1,fun(X2,X0)),X5))) <=> ? [X6,X7] : (s('type.pair.prod'(X1,X2),X4) = s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7))) & s(X0,X3) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X6))),s(X2,X7)))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5] : (s(X0,X3) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),X4),s(fun(X1,fun(X2,X0)),X5))) <=> ? [X6,X7] : (s('type.pair.prod'(X1,X2),X4) = s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X6),s(X2,X7))) & s(X0,X3) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X6))),s(X2,X7)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11] : (s(X6,X9) = s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X7,X8),X10),s(fun(X7,fun(X8,X6)),X11))) <=> ? [X12,X13] : (s('type.pair.prod'(X7,X8),X10) = s('type.pair.prod'(X7,X8),'const.pair.,_2'(s(X7,X12),s(X8,X13))) & s(X6,X9) = s(X6,happ(s(fun(X8,X6),happ(s(fun(X7,fun(X8,X6)),X11),s(X7,X12))),s(X8,X13)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11] : (s(X6,X9) = s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X7,X8),X10),s(fun(X7,fun(X8,X6)),X11))) <=> ? [X12,X13] : (s('type.pair.prod'(X7,X8),X10) = s('type.pair.prod'(X7,X8),'const.pair.,_2'(s(X7,X12),s(X8,X13))) & s(X6,X9) = s(X6,happ(s(fun(X8,X6),happ(s(fun(X7,fun(X8,X6)),X11),s(X7,X12))),s(X8,X13)))))),
  file('thm.clos_relationProps.clos_relationProps_9_prove_ltb','thm.clos_relationProps.clos_relationProps_9_prove_ltb')).
fof(f321,plain,(
  s(sK7,sK10) = s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))) | $spl4),
  inference(cnf_transformation,[],[f321_D])).
fof(f321_D,plain,(
  s(sK7,sK10) = s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f192,plain,(
  s(sK7,sK10) = s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,sK13))),s(sK9,sK14))) | s(sK7,sK10) = s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12)))),
  inference(cnf_transformation,[],[f138])).
fof(f318,plain,(
  s(sK7,sK10) != s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) | $spl1),
  inference(cnf_transformation,[],[f318_D])).
fof(f318_D,plain,(
  s(sK7,sK10) != s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f319,plain,(
  ( ! [X6,X7] : (s(sK7,sK10) != s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,X6))),s(sK9,X7))) | s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,X6),s(sK9,X7)))) ) | $spl2),
  inference(cnf_transformation,[],[f319_D])).
fof(f319_D,plain,(
  ( ! [X6,X7] : (s(sK7,sK10) != s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,X6))),s(sK9,X7))) | s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,X6),s(sK9,X7)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f193,plain,(
  ( ! [X6,X7] : (s(sK7,sK10) != s(sK7,happ(s(fun(sK9,sK7),happ(s(fun(sK8,fun(sK9,sK7)),sK12),s(sK8,X6))),s(sK9,X7))) | s('type.pair.prod'(sK8,sK9),sK11) != s('type.pair.prod'(sK8,sK9),'const.pair.,_2'(s(sK8,X6),s(sK9,X7))) | s(sK7,sK10) != s(sK7,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK8,sK9),sK11),s(fun(sK8,fun(sK9,sK7)),sK12)))) )),
  inference(cnf_transformation,[],[f138])).
% SZS output end Proof for thm.clos_relationProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 127
% Passive clauses: 248
% Generated clauses: 1827
% Final active clauses: 108
% Final passive clauses: 78
% Input formulas: 38
% Initial clauses: 120
% 
% Duplicate literals: 45
% Trivial inequalities: 17
% Fw subsumption resolutions: 101
% Fw demodulations: 281
% Bw demodulations: 20
% Condensations: 97
% Global subsumptions: 4
% 
% Simple tautologies: 51
% Equational tautologies: 44
% Forward subsumptions: 750
% Fw demodulations to eq. taut.: 118
% Bw demodulations to eq. taut.: 10
% 
% Unit resulting resolution: 486
% Factoring: 1
% Forward superposition: 316
% Backward superposition: 215
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 8
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 344
% SAT solver unit clauses: 47
% SAT solver binary clauses: 82
% 
% Sat splits: 19
% Sat splitting refutations: 52
% 
% Memory used [KB]: 2174
% Time elapsed: 0.065 s
% ------------------------------
% ------------------------------
% Success in time 1.418 s
